[1376372038] SERVICE ALERT: wdcpf001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376372058] HOST ALERT: wdcpf001;DOWN;SOFT;1;(Host Check Timed Out)
[1376372098] SERVICE ALERT: wdcpf001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1376372098] SERVICE ALERT: wdcps002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 6.53 ms
[1376372108] HOST ALERT: wdcpf001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1376372128] HOST ALERT: wdcps002;DOWN;SOFT;1;(Host Check Timed Out)
[1376372158] SERVICE ALERT: wdcps002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376372228] HOST ALERT: wdcps002;DOWN;SOFT;2;(Host Check Timed Out)
[1376372328] HOST ALERT: wdcps002;DOWN;SOFT;3;(Host Check Timed Out)
[1376372398] SERVICE ALERT: wdcps002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1376372398] HOST ALERT: wdcps002;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1376373598] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376373628] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1376373658] SERVICE ALERT: wdcps001;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 6.69 ms
[1376373698] HOST ALERT: wdcps001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1376373718] SERVICE ALERT: wdcps001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1376373718] SERVICE ALERT: wdcps003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376373758] HOST ALERT: wdcps003;DOWN;SOFT;1;(Host Check Timed Out)
[1376373778] SERVICE ALERT: wdcps003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376373828] HOST ALERT: wdcps003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1376373838] SERVICE ALERT: wdcps003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.51 ms
[1376381148] SERVICE ALERT: omnmdk11-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376381148] SERVICE ALERT: omnmdk11h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376381198] SERVICE ALERT: omnmdk11-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1376381198] SERVICE ALERT: omnmdk11h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1376390198] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376390198] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 6.59 ms
[1376390198] SERVICE ALERT: wdcms002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376390228] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1376390228] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1376390228] HOST ALERT: wdcms002;DOWN;SOFT;1;(Host Check Timed Out)
[1376390248] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1376390248] SERVICE ALERT: wdcms002;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1376390258] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376390258] HOST ALERT: wdcms001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1376390258] HOST ALERT: wdcms002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1376390298] HOST ALERT: wdcms003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1376390308] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1376410398] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410398] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410428] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1376410428] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376410458] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410458] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410498] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.49 ms
[1376410498] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.97 ms
[1376410508] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376410508] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376410578] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410578] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410618] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1376410618] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376410638] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410638] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410698] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.28 ms
[1376410698] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 2.36 ms
[1376410698] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1376410698] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376410698] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410698] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410728] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1376410728] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376410758] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410758] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410798] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376410798] HOST FLAPPING ALERT: nycnn-ipvx6;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376410798] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376410798] HOST FLAPPING ALERT: nycnn-ipvx6e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376410808] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376410808] SERVICE FLAPPING ALERT: nycnn-ipvx6;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376410808] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376410808] SERVICE FLAPPING ALERT: nycnn-ipvx6e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376410818] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410818] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376410858] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1376410858] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376410878] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410878] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376410938] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376410938] SERVICE FLAPPING ALERT: nycnn-ipvx7;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376410938] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1376410938] SERVICE FLAPPING ALERT: nycnn-ipvx7e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376410938] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.34 ms
[1376410938] HOST FLAPPING ALERT: nycnn-ipvx7;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376410938] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.27 ms
[1376410938] HOST FLAPPING ALERT: nycnn-ipvx7e1;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376411238] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411238] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411268] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1376411268] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376411298] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411298] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411298] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411298] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411328] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1376411328] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376411338] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.51 ms
[1376411338] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.53 ms
[1376411348] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1376411348] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1376411358] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411358] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411398] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376411398] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;2;(Host Check Timed Out)
[1376411408] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1376411408] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376411418] HOST ALERT: nycnn-ipvx4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1376411418] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376411478] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411478] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411518] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376411518] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1376411538] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411538] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411538] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411538] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411578] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1376411578] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376411588] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1376411588] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376411598] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411598] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376411598] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376411598] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411598] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411598] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411628] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1376411628] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376411638] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411638] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411648] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.42 ms
[1376411648] HOST FLAPPING ALERT: nycnn-ipvx3e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376411648] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 54.20 ms
[1376411648] SERVICE FLAPPING ALERT: nycnn-ipvx3;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376411648] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.33 ms
[1376411648] HOST FLAPPING ALERT: nycnn-ipvx3;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1376411648] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376411648] SERVICE FLAPPING ALERT: nycnn-ipvx3e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376411658] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411658] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411668] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1376411668] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376411698] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1376411698] HOST FLAPPING ALERT: nycnn-ipvx4;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376411698] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376411698] HOST FLAPPING ALERT: nycnn-ipvx4e1;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376411698] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411698] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411708] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376411708] SERVICE FLAPPING ALERT: nycnn-ipvx4;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376411708] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376411708] SERVICE FLAPPING ALERT: nycnn-ipvx4e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376411718] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411718] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376411748] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;2;(Host Check Timed Out)
[1376411748] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376411758] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376411758] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1376411778] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411778] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376411818] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376411818] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376411818] HOST ALERT: nycnn-ipvx10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376411818] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376411828] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1376411828] HOST FLAPPING ALERT: nycnn-ipvx8e1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376411828] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376411828] HOST FLAPPING ALERT: nycnn-ipvx8;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376411838] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376411838] SERVICE FLAPPING ALERT: nycnn-ipvx8e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376411838] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376411838] SERVICE FLAPPING ALERT: nycnn-ipvx8;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376411948] SERVICE FLAPPING ALERT: nycnn-ipvx6;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376411948] SERVICE FLAPPING ALERT: nycnn-ipvx6e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412018] HOST FLAPPING ALERT: nycnn-ipvx7;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376412018] HOST FLAPPING ALERT: nycnn-ipvx7e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376412078] SERVICE FLAPPING ALERT: nycnn-ipvx7;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412078] SERVICE FLAPPING ALERT: nycnn-ipvx7e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412118] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376412118] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376412128] HOST FLAPPING ALERT: nycnn-ipvx6;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412128] HOST FLAPPING ALERT: nycnn-ipvx6e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412148] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1376412148] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376412178] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376412178] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376412198] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.19 ms
[1376412198] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376412218] HOST ALERT: nycnn-ipvx10;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376412218] HOST FLAPPING ALERT: nycnn-ipvx10;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1376412218] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376412218] HOST FLAPPING ALERT: nycnn-ipvx10e1;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1376412228] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1376412228] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376412238] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376412238] SERVICE FLAPPING ALERT: nycnn-ipvx10;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1376412238] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376412238] SERVICE FLAPPING ALERT: nycnn-ipvx10e1;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1376412248] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376412258] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376412288] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376412288] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;2;(Host Check Timed Out)
[1376412358] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376412358] HOST ALERT: nycnn-ipvx3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376412368] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376412368] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376412728] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376412728] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376412768] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1376412768] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376412788] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376412788] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376412838] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376412838] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376412838] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376412838] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376412848] SERVICE FLAPPING ALERT: nycnn-ipvx4;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412848] HOST FLAPPING ALERT: nycnn-ipvx4;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412848] SERVICE FLAPPING ALERT: nycnn-ipvx4e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412848] HOST FLAPPING ALERT: nycnn-ipvx4e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412918] HOST FLAPPING ALERT: nycnn-ipvx8e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412918] HOST FLAPPING ALERT: nycnn-ipvx8;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412968] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376412968] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376412978] SERVICE FLAPPING ALERT: nycnn-ipvx8e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412978] SERVICE FLAPPING ALERT: nycnn-ipvx8;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376412988] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1376412988] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413028] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413028] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413028] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413028] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413058] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.54 ms
[1376413058] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.79 ms
[1376413068] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1376413068] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413078] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376413078] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376413088] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413088] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413138] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376413138] HOST FLAPPING ALERT: nycnn-ipvx20;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376413138] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376413138] HOST FLAPPING ALERT: nycnn-ipvx20e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376413138] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376413138] SERVICE FLAPPING ALERT: nycnn-ipvx20;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376413138] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376413138] SERVICE FLAPPING ALERT: nycnn-ipvx20e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376413148] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413148] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413178] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1376413178] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413208] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413208] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413208] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413208] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413248] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1376413248] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413248] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376413248] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1376413258] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1376413258] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376413268] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413268] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413318] HOST FLAPPING ALERT: nycnn-ipvx10;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376413318] HOST FLAPPING ALERT: nycnn-ipvx10e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376413318] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376413318] HOST FLAPPING ALERT: nycnn-ipvx21;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376413318] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376413318] HOST FLAPPING ALERT: nycnn-ipvx21e1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376413318] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376413318] SERVICE FLAPPING ALERT: nycnn-ipvx21;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376413318] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376413318] SERVICE FLAPPING ALERT: nycnn-ipvx21e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376413378] SERVICE FLAPPING ALERT: nycnn-ipvx10;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376413378] SERVICE FLAPPING ALERT: nycnn-ipvx10e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376413448] SERVICE FLAPPING ALERT: nycnn-ipvx3;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376413448] HOST FLAPPING ALERT: nycnn-ipvx3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376413448] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413448] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413448] SERVICE FLAPPING ALERT: nycnn-ipvx3e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376413448] HOST FLAPPING ALERT: nycnn-ipvx3e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376413478] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1376413478] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413508] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413508] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413518] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.19 ms
[1376413518] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413548] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376413548] HOST FLAPPING ALERT: nycnn-ipvx22;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376413548] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376413548] HOST FLAPPING ALERT: nycnn-ipvx22e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376413558] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1376413558] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413558] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376413558] SERVICE FLAPPING ALERT: nycnn-ipvx22;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376413558] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1376413558] SERVICE FLAPPING ALERT: nycnn-ipvx22e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376413578] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413578] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413638] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413638] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413648] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.31 ms
[1376413648] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 1.15 ms
[1376413678] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413678] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1376413688] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376413688] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376413698] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413698] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413748] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376413748] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376413758] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376413758] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376413818] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413818] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413858] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1376413858] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413878] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413878] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413878] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413878] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413918] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413918] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1376413928] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413938] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376413938] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413938] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413948] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.50 ms
[1376413948] HOST FLAPPING ALERT: nycnn-ipvx5;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1376413948] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.54 ms
[1376413948] HOST FLAPPING ALERT: nycnn-ipvx5e1;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1376413958] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1376413978] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376413988] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413988] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376413988] SERVICE FLAPPING ALERT: nycnn-ipvx5;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1376413998] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.71 ms
[1376413998] HOST FLAPPING ALERT: nycnn-ipvx9e1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376413998] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.25 ms
[1376413998] HOST FLAPPING ALERT: nycnn-ipvx9;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376413998] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376413998] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376413998] SERVICE FLAPPING ALERT: nycnn-ipvx5e1;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1376413998] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1376413998] SERVICE FLAPPING ALERT: nycnn-ipvx9e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376413998] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376413998] SERVICE FLAPPING ALERT: nycnn-ipvx9;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376414058] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.30 ms
[1376414058] SERVICE ALERT: nycnn-ipvx2e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 54%, RTA = 0.27 ms
[1376414058] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;2;(Host Check Timed Out)
[1376414098] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376414108] HOST ALERT: nycnn-ipvx2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376414108] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376414228] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414238] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414258] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1376414268] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376414278] SERVICE FLAPPING ALERT: nycnn-ipvx20;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414278] HOST FLAPPING ALERT: nycnn-ipvx20;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414278] SERVICE FLAPPING ALERT: nycnn-ipvx20e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414278] HOST FLAPPING ALERT: nycnn-ipvx20e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414288] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414298] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414328] HOST ALERT: nycnn-ipvx2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376414328] HOST FLAPPING ALERT: nycnn-ipvx2;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376414338] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376414338] HOST FLAPPING ALERT: nycnn-ipvx2e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376414338] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376414338] SERVICE FLAPPING ALERT: nycnn-ipvx2;PING;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1376414348] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376414348] SERVICE FLAPPING ALERT: nycnn-ipvx2e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376414348] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414348] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414378] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376414378] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1376414408] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414408] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414408] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414408] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414438] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376414438] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1376414448] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376414448] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376414458] SERVICE FLAPPING ALERT: nycnn-ipvx21;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414458] HOST FLAPPING ALERT: nycnn-ipvx21;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414458] SERVICE FLAPPING ALERT: nycnn-ipvx21e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414458] HOST FLAPPING ALERT: nycnn-ipvx21e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414468] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376414468] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376414468] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414468] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414508] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.45 ms
[1376414508] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.26 ms
[1376414528] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376414528] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376414588] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414588] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414628] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376414628] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1376414648] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414648] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414648] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414648] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414688] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376414688] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1376414698] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376414698] HOST FLAPPING ALERT: nycnn-ipvx23e1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376414698] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376414698] HOST FLAPPING ALERT: nycnn-ipvx23;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376414698] SERVICE FLAPPING ALERT: nycnn-ipvx22;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414698] HOST FLAPPING ALERT: nycnn-ipvx22;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414698] SERVICE FLAPPING ALERT: nycnn-ipvx22e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414698] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376414698] SERVICE FLAPPING ALERT: nycnn-ipvx23e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376414698] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376414698] SERVICE FLAPPING ALERT: nycnn-ipvx23;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376414708] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414708] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376414758] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376414758] HOST FLAPPING ALERT: nycnn-ipvx24e1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376414758] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376414758] HOST FLAPPING ALERT: nycnn-ipvx24;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376414768] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376414768] SERVICE FLAPPING ALERT: nycnn-ipvx24e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376414768] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376414768] SERVICE FLAPPING ALERT: nycnn-ipvx24;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376414788] HOST FLAPPING ALERT: nycnn-ipvx22e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376414948] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414948] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376414978] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1376414978] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415008] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415008] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415068] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415068] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415068] SERVICE FLAPPING ALERT: nycnn-ipvx5;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376415068] HOST FLAPPING ALERT: nycnn-ipvx5;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376415068] SERVICE FLAPPING ALERT: nycnn-ipvx5e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376415078] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;2;(Host Check Timed Out)
[1376415078] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376415108] HOST ALERT: nycnn-ipvx26;DOWN;SOFT;1;(Host Check Timed Out)
[1376415108] HOST ALERT: nycnn-ipvx26e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415128] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376415128] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376415128] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415128] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415138] SERVICE FLAPPING ALERT: nycnn-ipvx9e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415138] HOST FLAPPING ALERT: nycnn-ipvx9e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415138] SERVICE FLAPPING ALERT: nycnn-ipvx9;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415138] HOST FLAPPING ALERT: nycnn-ipvx9;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415138] HOST ALERT: nycnn-ipvx25;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376415138] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376415138] HOST ALERT: nycnn-ipvx26e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376415178] HOST ALERT: nycnn-ipvx26;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376415188] SERVICE ALERT: nycnn-ipvx26;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1376415188] SERVICE ALERT: nycnn-ipvx26e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376415188] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415188] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415198] HOST ALERT: nycnn-ipvx26e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376415218] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1376415218] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415248] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415248] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415248] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415248] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415248] HOST FLAPPING ALERT: nycnn-ipvx5e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376415288] HOST ALERT: nycnn-ipvx26;DOWN;SOFT;1;(Host Check Timed Out)
[1376415288] HOST ALERT: nycnn-ipvx26e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415308] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415308] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415308] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415308] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415308] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;2;(Host Check Timed Out)
[1376415308] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376415338] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1376415338] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415368] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376415368] SERVICE FLAPPING ALERT: nycnn-ipvx25;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376415368] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376415368] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376415368] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415368] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415368] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415368] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415378] HOST ALERT: nycnn-ipvx25;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376415378] HOST FLAPPING ALERT: nycnn-ipvx25;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1376415378] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376415378] HOST FLAPPING ALERT: nycnn-ipvx25e1;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1376415378] HOST ALERT: nycnn-ipvx26;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.26 ms
[1376415378] HOST FLAPPING ALERT: nycnn-ipvx26;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1376415378] HOST ALERT: nycnn-ipvx26e1;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.24 ms
[1376415378] HOST FLAPPING ALERT: nycnn-ipvx26e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376415408] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1376415408] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415428] SERVICE ALERT: nycnn-ipvx26;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376415428] SERVICE FLAPPING ALERT: nycnn-ipvx26;PING;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1376415428] SERVICE ALERT: nycnn-ipvx26e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376415428] SERVICE FLAPPING ALERT: nycnn-ipvx26e1;PING;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1376415428] SERVICE ALERT: nycnn-ipvx27;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 1.05 ms
[1376415428] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.48 ms
[1376415428] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.30 ms
[1376415428] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.30 ms
[1376415428] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415428] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415448] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376415448] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;2;(Host Check Timed Out)
[1376415478] SERVICE FLAPPING ALERT: nycnn-ipvx2;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415478] HOST FLAPPING ALERT: nycnn-ipvx2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415488] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376415488] SERVICE FLAPPING ALERT: nycnn-ipvx2e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415488] HOST FLAPPING ALERT: nycnn-ipvx2e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415518] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1376415518] HOST ALERT: nycnn-ipvx28;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376415548] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376415548] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376415608] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415608] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415608] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415608] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376415638] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1376415638] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415638] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1376415638] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376415668] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415668] SERVICE FLAPPING ALERT: nycnn-ipvx27;PING;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1376415668] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415668] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415668] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376415708] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376415708] HOST FLAPPING ALERT: nycnn-ipvx27;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376415708] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376415708] HOST FLAPPING ALERT: nycnn-ipvx27e1;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376415728] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376415728] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1376415728] SERVICE FLAPPING ALERT: nycnn-ipvx27e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376415738] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;2;(Host Check Timed Out)
[1376415738] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376415788] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376415788] SERVICE FLAPPING ALERT: nycnn-ipvx28;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376415788] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376415788] SERVICE FLAPPING ALERT: nycnn-ipvx28e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376415798] HOST ALERT: nycnn-ipvx28;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376415798] HOST FLAPPING ALERT: nycnn-ipvx28;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1376415798] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376415798] HOST FLAPPING ALERT: nycnn-ipvx28e1;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1376415838] SERVICE FLAPPING ALERT: nycnn-ipvx23e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415838] SERVICE FLAPPING ALERT: nycnn-ipvx23;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415848] HOST FLAPPING ALERT: nycnn-ipvx24e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415848] HOST FLAPPING ALERT: nycnn-ipvx24;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415908] SERVICE FLAPPING ALERT: nycnn-ipvx24e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415908] SERVICE FLAPPING ALERT: nycnn-ipvx24;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415938] HOST FLAPPING ALERT: nycnn-ipvx23e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376415938] HOST FLAPPING ALERT: nycnn-ipvx23;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376416128] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376416168] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;1;(Host Check Timed Out)
[1376416188] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376416268] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;2;(Host Check Timed Out)
[1376416298] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;HARD;1;Server is down!
[1376416308] SERVICE ALERT: nycnn-adbeame2;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.45 ms
[1376416308] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376416338] HOST ALERT: nycnn-adbeame2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1376416348] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;1;(Host Check Timed Out)
[1376416358] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1376416368] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376416388] HOST FLAPPING ALERT: nycnn-ipvx25;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1376416388] HOST FLAPPING ALERT: nycnn-ipvx25e1;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1376416418] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1376416438] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;2;(Host Check Timed Out)
[1376416448] SERVICE FLAPPING ALERT: nycnn-ipvx25;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416448] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416488] SERVICE ALERT: nycnn-adbeame3;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.50 ms
[1376416508] SERVICE FLAPPING ALERT: nycnn-ipvx26;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416508] HOST FLAPPING ALERT: nycnn-ipvx26;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416508] SERVICE FLAPPING ALERT: nycnn-ipvx26e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416508] HOST FLAPPING ALERT: nycnn-ipvx26e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416508] HOST ALERT: nycnn-adbeame3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1376416548] SERVICE ALERT: nycnn-adbeame3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1376416608] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376416648] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;1;(Host Check Timed Out)
[1376416668] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376416668] SERVICE FLAPPING ALERT: nycnn-adbeame2;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1376416668] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376416708] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;1;(Host Check Timed Out)
[1376416728] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376416728] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376416728] SERVICE FLAPPING ALERT: nycnn-adbeame3;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1376416748] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;2;(Host Check Timed Out)
[1376416768] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;1;(Host Check Timed Out)
[1376416778] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1376416788] HOST ALERT: nycnn-adbeame2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1376416788] HOST FLAPPING ALERT: nycnn-adbeame2;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376416788] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376416808] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;2;(Host Check Timed Out)
[1376416808] HOST FLAPPING ALERT: nycnn-ipvx27;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376416808] HOST FLAPPING ALERT: nycnn-ipvx27e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376416808] HOST FLAPPING ALERT: nycnn-ipvx28;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376416808] HOST FLAPPING ALERT: nycnn-ipvx28e1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376416848] SERVICE ALERT: nycnn-adbeame3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1376416848] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.28 ms
[1376416858] HOST ALERT: nycnn-adbeame3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1376416858] HOST FLAPPING ALERT: nycnn-adbeame3;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376416868] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;2;(Host Check Timed Out)
[1376416868] SERVICE FLAPPING ALERT: nycnn-ipvx27;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376416868] SERVICE FLAPPING ALERT: nycnn-ipvx27e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376416868] SERVICE FLAPPING ALERT: nycnn-ipvx28;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416868] SERVICE FLAPPING ALERT: nycnn-ipvx28e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376416878] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;1;(Host Check Timed Out)
[1376416898] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;HARD;1;AME is idle:
[1376416898] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1376416908] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376416938] HOST ALERT: nycnn-adbeame4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1376416958] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376416958] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;HARD;1;AME is idle:
[1376416958] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1376416968] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;2;(Host Check Timed Out)
[1376416968] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;1;(Host Check Timed Out)
[1376416968] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1376417018] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376417038] HOST ALERT: nycnn-adbeame5;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 0.47 ms
[1376417068] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;2;(Host Check Timed Out)
[1376417088] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1376417128] SERVICE ALERT: nycnn-adbeame6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1376417138] HOST ALERT: nycnn-adbeame6;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1376417148] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376417178] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;1;(Host Check Timed Out)
[1376417208] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376417208] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376417238] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;1;(Host Check Timed Out)
[1376417268] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376417278] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;2;(Host Check Timed Out)
[1376417318] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376417328] SERVICE ALERT: nycnn-adbeame4;PING;WARNING;HARD;2;PING WARNING - Packet loss = 54%, RTA = 0.35 ms
[1376417328] SERVICE FLAPPING ALERT: nycnn-adbeame4;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1376417338] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;2;(Host Check Timed Out)
[1376417348] HOST ALERT: nycnn-adbeame4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376417348] HOST FLAPPING ALERT: nycnn-adbeame4;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376417348] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;1;(Host Check Timed Out)
[1376417378] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376417388] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1376417418] HOST ALERT: nycnn-adbeame5;UP;SOFT;3;PING OK - Packet loss = 68%, RTA = 0.26 ms
[1376417418] HOST FLAPPING ALERT: nycnn-adbeame5;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376417438] SERVICE ALERT: nycnn-adbeame4;AME Job status;OK;HARD;1;AME is idle:
[1376417438] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;2;(Host Check Timed Out)
[1376417448] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376417448] SERVICE FLAPPING ALERT: nycnn-adbeame5;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1376417448] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;HARD;1;Server is down!
[1376417488] SERVICE ALERT: nycnn-adbeame6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1376417488] SERVICE FLAPPING ALERT: nycnn-adbeame6;PING;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1376417498] HOST ALERT: nycnn-adbeame6;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376417498] HOST FLAPPING ALERT: nycnn-adbeame6;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376417498] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;SOFT;1;Job  failed with error Destination path appears to be invalid.
[1376417738] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;HARD;1;AME is idle:
[1376417858] SERVICE FLAPPING ALERT: nycnn-adbeame2;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376417868] SERVICE ALERT: nycnn-adbeame5;AME Job status;OK;HARD;2;Job 09475766 progress is  62.95 :
[1376417928] SERVICE FLAPPING ALERT: nycnn-adbeame3;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376418468] SERVICE FLAPPING ALERT: nycnn-adbeame5;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376418528] SERVICE FLAPPING ALERT: nycnn-adbeame4;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376418568] SERVICE FLAPPING ALERT: nycnn-adbeame6;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376419478] HOST FLAPPING ALERT: nycnn-adbeame2;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376419548] HOST FLAPPING ALERT: nycnn-adbeame3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376420138] HOST FLAPPING ALERT: nycnn-adbeame4;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376420188] HOST FLAPPING ALERT: nycnn-adbeame6;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376420208] HOST FLAPPING ALERT: nycnn-adbeame5;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)