[1413356316] SERVICE ALERT: nycnn-ipvx18;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413356366] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413356376] SERVICE ALERT: nycnn-ipvx18;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.01%
[1413356426] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1413376016] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job 10519877 remains in a queued state
[1413376076] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1413376376] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;HARD;2;AME is idle:
[1413393990] HOST DOWNTIME ALERT: nycnn-ipvx1;STARTED; Host has entered a period of scheduled downtime
[1413394013] HOST DOWNTIME ALERT: nycnn-ipvx10;STARTED; Host has entered a period of scheduled downtime
[1413394244] HOST DOWNTIME ALERT: nycnn-ipvx11;STARTED; Host has entered a period of scheduled downtime
[1413394262] HOST DOWNTIME ALERT: nycnn-ipvx12;STARTED; Host has entered a period of scheduled downtime
[1413394277] HOST DOWNTIME ALERT: nycnn-ipvx13;STARTED; Host has entered a period of scheduled downtime
[1413394301] HOST DOWNTIME ALERT: nycnn-ipvx14;STARTED; Host has entered a period of scheduled downtime
[1413394319] HOST DOWNTIME ALERT: nycnn-ipvx15;STARTED; Host has entered a period of scheduled downtime
[1413394341] HOST DOWNTIME ALERT: nycnn-ipvx16;STARTED; Host has entered a period of scheduled downtime
[1413394357] HOST DOWNTIME ALERT: nycnn-ipvx17;STARTED; Host has entered a period of scheduled downtime
[1413394372] HOST DOWNTIME ALERT: nycnn-ipvx18;STARTED; Host has entered a period of scheduled downtime
[1413394393] HOST DOWNTIME ALERT: nycnn-ipvx19;STARTED; Host has entered a period of scheduled downtime
[1413394414] HOST DOWNTIME ALERT: nycnn-ipvx1e1;STARTED; Host has entered a period of scheduled downtime
[1413394434] HOST DOWNTIME ALERT: nycnn-ipvx2;STARTED; Host has entered a period of scheduled downtime
[1413394453] HOST DOWNTIME ALERT: nycnn-ipvx20;STARTED; Host has entered a period of scheduled downtime
[1413394477] HOST DOWNTIME ALERT: nycnn-ipvx21;STARTED; Host has entered a period of scheduled downtime
[1413394495] HOST DOWNTIME ALERT: nycnn-ipvx22;STARTED; Host has entered a period of scheduled downtime
[1413394512] HOST DOWNTIME ALERT: nycnn-ipvx23;STARTED; Host has entered a period of scheduled downtime
[1413394533] HOST DOWNTIME ALERT: nycnn-ipvx24;STARTED; Host has entered a period of scheduled downtime
[1413394551] HOST DOWNTIME ALERT: nycnn-ipvx25;STARTED; Host has entered a period of scheduled downtime
[1413394593] HOST DOWNTIME ALERT: nycnn-ipvx26;STARTED; Host has entered a period of scheduled downtime
[1413394620] HOST DOWNTIME ALERT: nycnn-ipvx27;STARTED; Host has entered a period of scheduled downtime
[1413394641] HOST DOWNTIME ALERT: nycnn-ipvx28;STARTED; Host has entered a period of scheduled downtime
[1413394671] HOST DOWNTIME ALERT: nycnn-ipvx29;STARTED; Host has entered a period of scheduled downtime
[1413394691] HOST DOWNTIME ALERT: nycnn-ipvx3;STARTED; Host has entered a period of scheduled downtime
[1413394714] HOST DOWNTIME ALERT: nycnn-ipvx30;STARTED; Host has entered a period of scheduled downtime
[1413394746] HOST DOWNTIME ALERT: nycnn-ipvx31;STARTED; Host has entered a period of scheduled downtime
[1413394770] HOST DOWNTIME ALERT: nycnn-ipvx32;STARTED; Host has entered a period of scheduled downtime
[1413394786] SERVICE ALERT: rt-nycchdagg-b;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 110.45 ms
[1413394822] HOST DOWNTIME ALERT: nycnn-ipvx33;STARTED; Host has entered a period of scheduled downtime
[1413394842] HOST DOWNTIME ALERT: nycnn-ipvx34;STARTED; Host has entered a period of scheduled downtime
[1413394846] SERVICE ALERT: rt-nycchdagg-b;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 28.32 ms
[1413394878] HOST DOWNTIME ALERT: nycnn-ipvx35;STARTED; Host has entered a period of scheduled downtime
[1413394896] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;SOFT;1;Job 10520881 remains in a queued state
[1413394900] HOST DOWNTIME ALERT: nycnn-ipvx4;STARTED; Host has entered a period of scheduled downtime
[1413394922] HOST DOWNTIME ALERT: nycnn-ipvx5;STARTED; Host has entered a period of scheduled downtime
[1413394940] HOST DOWNTIME ALERT: nycnn-ipvx6;STARTED; Host has entered a period of scheduled downtime
[1413394956] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;SOFT;2;AME is idle:
[1413394967] HOST DOWNTIME ALERT: nycnn-ipvx8;STARTED; Host has entered a period of scheduled downtime
[1413394984] HOST DOWNTIME ALERT: nycnn-ipvx9;STARTED; Host has entered a period of scheduled downtime
[1413396766] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413396786] SERVICE ALERT: nycnn-ipvx1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413396786] SERVICE ALERT: nycnn-ipvx1e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413396786] SERVICE ALERT: nycnn-ipvx1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413396796] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413396796] SERVICE ALERT: nycnn-ipvx1e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413396796] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1413396806] SERVICE ALERT: nycnn-ipvx1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413396816] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413396826] SERVICE ALERT: nycnn-ipvx1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413396826] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413396846] SERVICE ALERT: nycnn-ipvx1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413396846] SERVICE ALERT: nycnn-ipvx1e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413396846] SERVICE ALERT: nycnn-ipvx1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413396856] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413396856] SERVICE ALERT: nycnn-ipvx1e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413396856] SERVICE ALERT: nycnn-ipvx1e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413396886] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413396896] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.44 ms
[1413396896] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;2;(Host Check Timed Out)
[1413396906] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413396946] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413396956] HOST ALERT: nycnn-ipvx1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413397086] SERVICE ALERT: nycnn-ipvx1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413397106] SERVICE ALERT: nycnn-ipvx1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413397126] SERVICE ALERT: nycnn-ipvx1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413397136] SERVICE ALERT: nycnn-ipvx1e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.843GB (42%) - Free: 1.147GB (58%)
[1413397136] SERVICE ALERT: nycnn-ipvx1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.843GB (42%) - Free: 1.147GB (58%)
[1413397136] SERVICE ALERT: nycnn-ipvx1e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413397146] SERVICE ALERT: nycnn-ipvx1e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413397186] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413397836] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413397836] SERVICE ALERT: nycnn-ipvx3;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413397896] SERVICE ALERT: nycnn-ipvx5e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397906] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413397906] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1413397906] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413397906] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413397916] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413397916] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413397916] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397916] SERVICE ALERT: nycnn-ipvx3;Service: Apache;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1413397926] SERVICE ALERT: nycnn-ipvx5;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397926] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413397936] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1413397946] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397946] SERVICE ALERT: nycnn-ipvx5;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397946] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397946] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1413397956] SERVICE ALERT: nycnn-ipvx5;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397956] SERVICE ALERT: nycnn-ipvx5e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413397966] SERVICE ALERT: nycnn-ipvx3;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397966] SERVICE ALERT: nycnn-ipvx5;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397966] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413397966] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413397976] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413397976] SERVICE ALERT: nycnn-ipvx3;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397976] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413397976] SERVICE ALERT: nycnn-ipvx5e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397976] SERVICE ALERT: nycnn-ipvx3;Service: Apache;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1413397986] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397986] SERVICE ALERT: nycnn-ipvx5;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413397986] SERVICE ALERT: nycnn-ipvx5;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413397986] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.42 ms
[1413397996] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1413398006] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413398006] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398016] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1413398026] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1413398026] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413398026] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1413398026] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1413398196] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398236] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413398236] SERVICE ALERT: nycnn-ipvx5;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413398236] SERVICE ALERT: nycnn-ipvx5;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413398236] SERVICE ALERT: nycnn-ipvx5e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413398256] SERVICE ALERT: nycnn-ipvx3;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413398256] SERVICE ALERT: nycnn-ipvx5;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413398256] SERVICE ALERT: nycnn-ipvx3;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.833GB (42%) - Free: 1.157GB (58%)
[1413398256] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413398256] SERVICE ALERT: nycnn-ipvx5e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398256] SERVICE ALERT: nycnn-ipvx3;Service: Apache;OK;HARD;3;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413398266] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398266] SERVICE ALERT: nycnn-ipvx5;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.840GB (42%) - Free: 1.150GB (58%)
[1413398276] SERVICE ALERT: nycnn-ipvx5;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398296] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398316] SERVICE ALERT: nycnn-ipvx6e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413398366] SERVICE ALERT: nycnn-ipvx6;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398376] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398376] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413398386] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413398396] SERVICE ALERT: nycnn-ipvx6e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398396] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1413398416] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413398426] SERVICE ALERT: nycnn-ipvx6;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398426] SERVICE ALERT: nycnn-ipvx6e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398436] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398436] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413398446] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413398456] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398456] SERVICE ALERT: nycnn-ipvx6e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413398466] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1413398486] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413398486] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413398506] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1413398676] SERVICE ALERT: nycnn-ipvx6e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413398706] SERVICE ALERT: nycnn-ipvx6;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398716] SERVICE ALERT: nycnn-ipvx6e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.837GB (42%) - Free: 1.153GB (58%)
[1413398716] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413398736] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398756] SERVICE ALERT: nycnn-ipvx6e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413398756] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413398766] SERVICE ALERT: nycnn-ipvx8e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413398796] SERVICE ALERT: nycnn-ipvx8e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398806] SERVICE ALERT: nycnn-ipvx8e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398806] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413398816] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413398826] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398836] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413398836] SERVICE ALERT: nycnn-ipvx8;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398836] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398836] SERVICE ALERT: nycnn-ipvx8e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398856] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1413398856] SERVICE ALERT: nycnn-ipvx8e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398866] SERVICE ALERT: nycnn-ipvx8e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398866] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413398876] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413398876] SERVICE ALERT: nycnn-ipvx8e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413398886] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398896] SERVICE ALERT: nycnn-ipvx8;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413398916] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413398926] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.41 ms
[1413398926] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1413398926] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413398926] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1413398956] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413398976] SERVICE ALERT: nycnn-ipvx9e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413398986] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.19 ms
[1413398986] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 0.33 ms
[1413398986] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413398986] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399006] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399016] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399016] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1413399026] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399026] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1413399036] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399046] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399046] SERVICE ALERT: nycnn-ipvx9e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399046] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399046] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399046] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399056] SERVICE ALERT: nycnn-ipvx9;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399066] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399066] SERVICE ALERT: nycnn-ipvx10e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399076] SERVICE ALERT: nycnn-ipvx9;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399076] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399076] SERVICE ALERT: nycnn-ipvx9;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399096] SERVICE ALERT: nycnn-ipvx10;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399106] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413399106] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;2;(Host Check Timed Out)
[1413399116] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399126] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413399126] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;2;(Host Check Timed Out)
[1413399126] SERVICE ALERT: nycnn-ipvx8e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399126] SERVICE ALERT: nycnn-ipvx10e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413399136] SERVICE ALERT: nycnn-ipvx9;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399146] SERVICE ALERT: nycnn-ipvx8e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399156] SERVICE ALERT: nycnn-ipvx8e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.834GB (42%) - Free: 1.155GB (58%)
[1413399156] HOST ALERT: nycnn-ipvx10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413399156] SERVICE ALERT: nycnn-ipvx8e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399166] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413399166] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413399166] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413399166] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413399166] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399166] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399166] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413399176] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413399176] HOST ALERT: nycnn-ipvx9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413399176] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413399176] SERVICE ALERT: nycnn-ipvx8;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.837GB (42%) - Free: 1.153GB (58%)
[1413399186] SERVICE ALERT: nycnn-ipvx12e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399206] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1413399206] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399216] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399226] SERVICE ALERT: nycnn-ipvx12e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399226] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399226] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399246] SERVICE ALERT: nycnn-ipvx12e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399256] SERVICE ALERT: nycnn-ipvx12;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399266] SERVICE ALERT: nycnn-ipvx12e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399266] SERVICE ALERT: nycnn-ipvx12e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399286] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399286] SERVICE ALERT: nycnn-ipvx4;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413399296] SERVICE ALERT: nycnn-ipvx12;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413399296] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.47 ms
[1413399296] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.62 ms
[1413399316] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413399326] SERVICE ALERT: nycnn-ipvx11e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413399336] SERVICE ALERT: nycnn-ipvx9e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399336] SERVICE ALERT: nycnn-ipvx11;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413399336] SERVICE ALERT: nycnn-ipvx9;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399346] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1413399346] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413399346] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399346] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399346] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399346] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399356] SERVICE ALERT: nycnn-ipvx10e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413399356] SERVICE ALERT: nycnn-ipvx2;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413399356] SERVICE ALERT: nycnn-ipvx9;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413399366] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399366] SERVICE ALERT: nycnn-ipvx9;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399366] SERVICE ALERT: nycnn-ipvx4;Process: Siphon;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399366] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399376] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1413399376] SERVICE ALERT: nycnn-ipvx4e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399376] SERVICE ALERT: nycnn-ipvx4e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399376] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399376] SERVICE ALERT: nycnn-ipvx10;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399386] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399386] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399386] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1413399386] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399396] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1413399396] SERVICE ALERT: nycnn-ipvx11;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399396] SERVICE ALERT: nycnn-ipvx2;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399396] SERVICE ALERT: nycnn-ipvx2e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399396] SERVICE ALERT: nycnn-ipvx4;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399406] SERVICE ALERT: nycnn-ipvx11e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399406] SERVICE ALERT: nycnn-ipvx4e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399406] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399406] SERVICE ALERT: nycnn-ipvx11e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399406] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399406] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399406] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399416] SERVICE ALERT: nycnn-ipvx11;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399416] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399426] SERVICE ALERT: nycnn-ipvx4;Process: Siphon;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1413399426] SERVICE ALERT: nycnn-ipvx9;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399426] SERVICE ALERT: nycnn-ipvx10e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399426] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399436] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399436] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413399436] SERVICE ALERT: nycnn-ipvx4e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399436] SERVICE ALERT: nycnn-ipvx2;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399436] SERVICE ALERT: nycnn-ipvx4e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399436] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399436] SERVICE ALERT: nycnn-ipvx2e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399436] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;2;(Host Check Timed Out)
[1413399446] SERVICE ALERT: nycnn-ipvx2e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399446] SERVICE ALERT: nycnn-ipvx11e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399446] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399456] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399456] SERVICE ALERT: nycnn-ipvx2e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399466] SERVICE ALERT: nycnn-ipvx11;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399466] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1413399466] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413399466] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399476] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;2;(Host Check Timed Out)
[1413399476] SERVICE ALERT: nycnn-ipvx11e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399476] SERVICE ALERT: nycnn-ipvx2e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399476] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413399476] HOST ALERT: nycnn-ipvx4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1413399476] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399476] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413399486] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1413399486] SERVICE ALERT: nycnn-ipvx2e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399486] SERVICE ALERT: nycnn-ipvx2;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399496] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;2;(Host Check Timed Out)
[1413399496] HOST ALERT: nycnn-ipvx11;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413399506] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.42 ms
[1413399506] SERVICE ALERT: nycnn-ipvx2;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.49 ms
[1413399506] SERVICE ALERT: nycnn-ipvx12e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399526] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413399526] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1413399526] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1413399536] SERVICE ALERT: nycnn-ipvx12e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399536] SERVICE ALERT: nycnn-ipvx12;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399546] HOST ALERT: nycnn-ipvx2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413399546] SERVICE ALERT: nycnn-ipvx12e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.856GB (43%) - Free: 1.134GB (57%)
[1413399556] SERVICE ALERT: nycnn-ipvx12e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399566] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1413399576] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399586] SERVICE ALERT: nycnn-ipvx12;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399676] SERVICE ALERT: nycnn-ipvx11;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413399676] SERVICE ALERT: nycnn-ipvx2;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399686] SERVICE ALERT: nycnn-ipvx4;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399686] SERVICE ALERT: nycnn-ipvx11e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399696] SERVICE ALERT: nycnn-ipvx4e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399696] SERVICE ALERT: nycnn-ipvx11e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.844GB (42%) - Free: 1.146GB (58%)
[1413399696] SERVICE ALERT: nycnn-ipvx11;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399706] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399706] SERVICE ALERT: nycnn-ipvx4;Process: Siphon;OK;HARD;3;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399716] SERVICE ALERT: nycnn-ipvx4e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413399716] SERVICE ALERT: nycnn-ipvx2;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413399716] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399716] SERVICE ALERT: nycnn-ipvx4e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399726] SERVICE ALERT: nycnn-ipvx2e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399726] SERVICE ALERT: nycnn-ipvx2e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413399736] SERVICE ALERT: nycnn-ipvx13e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399736] SERVICE ALERT: nycnn-ipvx11e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413399736] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399736] SERVICE ALERT: nycnn-ipvx2e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399746] SERVICE ALERT: nycnn-ipvx11;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413399746] SERVICE ALERT: nycnn-ipvx13;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399756] SERVICE ALERT: nycnn-ipvx11e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399756] SERVICE ALERT: nycnn-ipvx2e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.832GB (42%) - Free: 1.158GB (58%)
[1413399756] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399766] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399766] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399766] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399776] SERVICE ALERT: nycnn-ipvx2e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413399776] SERVICE ALERT: nycnn-ipvx2;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.834GB (42%) - Free: 1.156GB (58%)
[1413399776] SERVICE ALERT: nycnn-ipvx2;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.14%
[1413399786] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1413399786] SERVICE ALERT: nycnn-ipvx2e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.12%
[1413399796] SERVICE ALERT: nycnn-ipvx13e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399806] SERVICE ALERT: nycnn-ipvx13;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399826] SERVICE ALERT: nycnn-ipvx13;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399826] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399826] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399836] SERVICE ALERT: nycnn-ipvx13;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399846] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.50 ms
[1413399846] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413399856] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413399856] HOST ALERT: nycnn-ipvx13;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413399876] SERVICE ALERT: nycnn-ipvx16;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413399886] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413399886] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413399906] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399906] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399916] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1413399916] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413399926] SERVICE ALERT: nycnn-ipvx16e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399936] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399946] SERVICE ALERT: nycnn-ipvx16;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399946] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399946] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1413399946] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413399956] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399966] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399966] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413399976] SERVICE ALERT: nycnn-ipvx14;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413399976] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413399976] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1413399986] SERVICE ALERT: nycnn-ipvx16e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413399996] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413400006] SERVICE ALERT: nycnn-ipvx15;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413400006] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413400006] SERVICE ALERT: nycnn-ipvx14;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400006] SERVICE ALERT: nycnn-ipvx15e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413400006] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413400016] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;2;(Host Check Timed Out)
[1413400016] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413400026] SERVICE ALERT: nycnn-ipvx14;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413400026] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413400026] SERVICE ALERT: nycnn-ipvx14e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400026] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 5.31 ms
[1413400026] SERVICE ALERT: nycnn-ipvx16e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 50%, RTA = 0.25 ms
[1413400026] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413400026] SERVICE ALERT: nycnn-ipvx16e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413400036] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400046] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400046] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413400046] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413400066] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413400066] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413400066] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413400066] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413400066] HOST ALERT: nycnn-ipvx16;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1413400066] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413400076] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400076] SERVICE ALERT: nycnn-ipvx13e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.842GB (42%) - Free: 1.148GB (58%)
[1413400086] SERVICE ALERT: nycnn-ipvx15;Process: Siphon;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400086] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413400086] SERVICE ALERT: nycnn-ipvx15e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400086] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400086] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413400086] SERVICE ALERT: nycnn-ipvx16e1;Service: Apache;OK;SOFT;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413400096] SERVICE ALERT: nycnn-ipvx13;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413400096] SERVICE ALERT: nycnn-ipvx15;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400096] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400106] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1413400106] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400116] SERVICE ALERT: nycnn-ipvx13;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.844GB (42%) - Free: 1.146GB (58%)
[1413400126] SERVICE ALERT: nycnn-ipvx13;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400126] SERVICE ALERT: nycnn-ipvx15e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400126] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413400136] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;2;(Host Check Timed Out)
[1413400136] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413400146] SERVICE ALERT: nycnn-ipvx15;Process: Siphon;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1413400146] SERVICE ALERT: nycnn-ipvx15;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413400146] SERVICE ALERT: nycnn-ipvx15;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413400146] SERVICE ALERT: nycnn-ipvx15e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 54%, RTA = 0.44 ms
[1413400146] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400186] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413400196] HOST ALERT: nycnn-ipvx15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413400206] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413400206] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413400236] SERVICE ALERT: nycnn-ipvx16;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413400236] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400246] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400266] SERVICE ALERT: nycnn-ipvx16e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413400296] SERVICE ALERT: nycnn-ipvx14;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400316] SERVICE ALERT: nycnn-ipvx14e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400326] SERVICE ALERT: nycnn-ipvx14;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.842GB (42%) - Free: 1.148GB (58%)
[1413400326] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400356] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400366] SERVICE ALERT: nycnn-ipvx15e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400376] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413400386] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413400386] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400406] SERVICE ALERT: nycnn-ipvx18e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413400406] SERVICE ALERT: nycnn-ipvx18e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413400406] SERVICE ALERT: nycnn-ipvx15e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413400426] SERVICE ALERT: nycnn-ipvx15;Process: Siphon;OK;HARD;3;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400426] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.23 ms
[1413400426] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413400436] SERVICE ALERT: nycnn-ipvx15;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413400446] SERVICE ALERT: nycnn-ipvx15;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.831GB (42%) - Free: 1.158GB (58%)
[1413400456] HOST ALERT: nycnn-ipvx18e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413400456] SERVICE ALERT: nycnn-ipvx18e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400456] HOST ALERT: nycnn-ipvx18;DOWN;SOFT;1;(Host Check Timed Out)
[1413400476] SERVICE ALERT: nycnn-ipvx18e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400486] SERVICE ALERT: nycnn-ipvx18e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413400486] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413400486] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413400496] SERVICE ALERT: nycnn-ipvx18;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400506] SERVICE ALERT: nycnn-ipvx18;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400546] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400556] HOST ALERT: nycnn-ipvx18e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413400556] HOST ALERT: nycnn-ipvx18;DOWN;SOFT;2;(Host Check Timed Out)
[1413400556] SERVICE ALERT: nycnn-ipvx18;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413400556] SERVICE ALERT: nycnn-ipvx18;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413400566] SERVICE ALERT: nycnn-ipvx18e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413400596] HOST ALERT: nycnn-ipvx18;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413400606] SERVICE ALERT: nycnn-ipvx18e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413400606] SERVICE ALERT: nycnn-ipvx18;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413400616] HOST ALERT: nycnn-ipvx18e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413400736] SERVICE ALERT: nycnn-ipvx18e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413400766] SERVICE ALERT: nycnn-ipvx18e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400766] SERVICE ALERT: nycnn-ipvx18e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400776] SERVICE ALERT: nycnn-ipvx18;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413400786] SERVICE ALERT: nycnn-ipvx18;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413400826] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413400836] SERVICE ALERT: nycnn-ipvx18;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1413400836] SERVICE ALERT: nycnn-ipvx18;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413400856] SERVICE ALERT: nycnn-ipvx18;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.796GB (40%) - Free: 1.194GB (60%)
[1413400866] SERVICE ALERT: nycnn-ipvx18e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.796GB (40%) - Free: 1.194GB (60%)
[1413400886] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.08%
[1413401256] SERVICE ALERT: nycnn-ipvx17e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413401266] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401286] SERVICE ALERT: nycnn-ipvx17e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401296] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413401296] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401326] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401326] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401326] SERVICE ALERT: nycnn-ipvx17e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401326] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401326] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1413401346] SERVICE ALERT: nycnn-ipvx17;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401346] SERVICE ALERT: nycnn-ipvx17e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401356] SERVICE ALERT: nycnn-ipvx17;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401356] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401366] SERVICE ALERT: nycnn-ipvx17e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401386] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401396] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413401396] SERVICE ALERT: nycnn-ipvx17e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413401396] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.55 ms
[1413401406] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413401446] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413401456] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1413401606] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413401616] SERVICE ALERT: nycnn-ipvx17e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413401636] SERVICE ALERT: nycnn-ipvx17;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413401636] SERVICE ALERT: nycnn-ipvx17e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413401636] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413401636] SERVICE ALERT: nycnn-ipvx17;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413401646] SERVICE ALERT: nycnn-ipvx17e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413401646] SERVICE ALERT: nycnn-ipvx19;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413401666] SERVICE ALERT: nycnn-ipvx19;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401666] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413401666] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401676] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1413401686] SERVICE ALERT: nycnn-ipvx21;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401686] SERVICE ALERT: nycnn-ipvx17e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.721GB (36%) - Free: 1.269GB (64%)
[1413401686] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401686] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401686] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401686] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401686] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413401696] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401706] SERVICE ALERT: nycnn-ipvx19;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401706] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401706] SERVICE ALERT: nycnn-ipvx20e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401706] SERVICE ALERT: nycnn-ipvx19e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401706] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401716] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401726] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1413401726] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413401726] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1413401726] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413401726] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413401726] SERVICE ALERT: nycnn-ipvx19;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401726] SERVICE ALERT: nycnn-ipvx19;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401726] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401746] SERVICE ALERT: nycnn-ipvx19;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401746] SERVICE ALERT: nycnn-ipvx21;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401746] SERVICE ALERT: nycnn-ipvx21;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401746] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401746] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401746] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401746] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401746] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413401756] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401766] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401766] SERVICE ALERT: nycnn-ipvx19e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401766] SERVICE ALERT: nycnn-ipvx20e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401766] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413401766] SERVICE ALERT: nycnn-ipvx21;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401776] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;2;(Host Check Timed Out)
[1413401776] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1413401776] SERVICE ALERT: nycnn-ipvx20;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401796] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.46 ms
[1413401796] SERVICE ALERT: nycnn-ipvx19;CPU Load;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413401796] SERVICE ALERT: nycnn-ipvx20e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401806] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413401806] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.48 ms
[1413401806] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.44 ms
[1413401806] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1413401806] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1413401816] SERVICE ALERT: nycnn-ipvx20;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413401816] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.46 ms
[1413401816] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.51 ms
[1413401836] SERVICE ALERT: nycnn-ipvx20e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413401846] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413401846] HOST ALERT: nycnn-ipvx19;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1413401866] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1413401866] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1413401886] SERVICE ALERT: nycnn-ipvx20e1;CPU Load;OK;SOFT;2;OK (Sample Period 360 sec) - Average CPU Utilisation 0.00%
[1413401986] SERVICE ALERT: nycnn-ipvx19;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.852GB (43%) - Free: 1.138GB (57%)
[1413402006] SERVICE ALERT: nycnn-ipvx19;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402016] SERVICE ALERT: nycnn-ipvx19;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413402026] SERVICE ALERT: nycnn-ipvx19;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413402026] SERVICE ALERT: nycnn-ipvx21;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402026] SERVICE ALERT: nycnn-ipvx21;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.843GB (42%) - Free: 1.146GB (58%)
[1413402036] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402046] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402046] SERVICE ALERT: nycnn-ipvx20e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413402046] SERVICE ALERT: nycnn-ipvx19e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413402056] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402056] SERVICE ALERT: nycnn-ipvx21;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413402056] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;OK;HARD;3;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402056] SERVICE ALERT: nycnn-ipvx20;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413402076] SERVICE ALERT: nycnn-ipvx20e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413402096] SERVICE ALERT: nycnn-ipvx19;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413402106] SERVICE ALERT: nycnn-ipvx20;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413402216] SERVICE ALERT: nycnn-ipvx22;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413402256] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402256] SERVICE ALERT: nycnn-ipvx22e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402266] SERVICE ALERT: nycnn-ipvx22e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402266] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413402286] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413402286] SERVICE ALERT: nycnn-ipvx22;CPU Load;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402296] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1413402296] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413402296] SERVICE ALERT: nycnn-ipvx22;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402296] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402316] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402316] SERVICE ALERT: nycnn-ipvx22e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402326] SERVICE ALERT: nycnn-ipvx22e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402326] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413402346] SERVICE ALERT: nycnn-ipvx22;CPU Load;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1413402346] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413402376] SERVICE ALERT: nycnn-ipvx22;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402376] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413402376] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.48 ms
[1413402376] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.41 ms
[1413402386] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1413402406] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413402516] SERVICE ALERT: nycnn-ipvx23;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413402526] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413402546] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413402546] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402546] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1413402566] SERVICE ALERT: nycnn-ipvx23;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402576] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402576] SERVICE ALERT: nycnn-ipvx22;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.845GB (42%) - Free: 1.145GB (58%)
[1413402576] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413402586] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.846GB (43%) - Free: 1.144GB (57%)
[1413402586] SERVICE ALERT: nycnn-ipvx23;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402586] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413402606] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402606] SERVICE ALERT: nycnn-ipvx22e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413402606] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413402606] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402606] SERVICE ALERT: nycnn-ipvx22e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402636] SERVICE ALERT: nycnn-ipvx22;CPU Load;OK;HARD;3;OK (Sample Period 720 sec) - Average CPU Utilisation 0.00%
[1413402646] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;2;(Host Check Timed Out)
[1413402646] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402646] SERVICE ALERT: nycnn-ipvx23;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402646] SERVICE ALERT: nycnn-ipvx23e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402656] SERVICE ALERT: nycnn-ipvx22;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413402666] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.27 ms
[1413402676] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413402706] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1413402716] HOST ALERT: nycnn-ipvx23;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413402716] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1413402806] SERVICE ALERT: nycnn-ipvx25;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413402826] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402826] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413402846] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1413402846] SERVICE ALERT: nycnn-ipvx23;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402856] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402856] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1413402866] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413402866] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402876] SERVICE ALERT: nycnn-ipvx23;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413402876] SERVICE ALERT: nycnn-ipvx25e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402876] SERVICE ALERT: nycnn-ipvx25;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402886] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413402886] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413402896] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413402896] SERVICE ALERT: nycnn-ipvx25;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402916] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413402926] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413402926] SERVICE ALERT: nycnn-ipvx23;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413402926] SERVICE ALERT: nycnn-ipvx23e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413402936] SERVICE ALERT: nycnn-ipvx23;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.13%
[1413402946] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;2;(Host Check Timed Out)
[1413402946] SERVICE ALERT: nycnn-ipvx25;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413402946] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 5.15 ms
[1413402946] SERVICE ALERT: nycnn-ipvx25;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 1.79 ms
[1413402956] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.15%
[1413402966] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413402976] HOST ALERT: nycnn-ipvx25;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413403006] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413403006] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413403036] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1413403036] SERVICE ALERT: nycnn-ipvx27;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403036] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403036] SERVICE ALERT: nycnn-ipvx27e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403036] SERVICE ALERT: nycnn-ipvx27;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403036] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403056] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413403076] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413403076] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1413403076] SERVICE ALERT: nycnn-ipvx27e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403086] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403086] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403096] SERVICE ALERT: nycnn-ipvx27;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403096] SERVICE ALERT: nycnn-ipvx27;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403096] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403096] SERVICE ALERT: nycnn-ipvx27;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403096] SERVICE ALERT: nycnn-ipvx27e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403096] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403096] SERVICE ALERT: nycnn-ipvx27;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403116] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403126] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403136] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403136] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;2;(Host Check Timed Out)
[1413403156] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413403166] SERVICE ALERT: nycnn-ipvx25e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403166] SERVICE ALERT: nycnn-ipvx25;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413403166] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413403176] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413403176] HOST ALERT: nycnn-ipvx27;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413403176] SERVICE ALERT: nycnn-ipvx25;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.848GB (43%) - Free: 1.142GB (57%)
[1413403186] SERVICE ALERT: nycnn-ipvx28e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413403206] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413403216] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1413403226] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413403226] SERVICE ALERT: nycnn-ipvx25;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403246] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1413403246] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413403256] SERVICE ALERT: nycnn-ipvx28e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403266] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413403276] SERVICE ALERT: nycnn-ipvx29e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403276] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403286] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1413403296] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413403296] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1413403306] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403316] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413403316] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403326] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413403326] SERVICE ALERT: nycnn-ipvx28;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403326] SERVICE ALERT: nycnn-ipvx29;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403326] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403336] SERVICE ALERT: nycnn-ipvx29;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403336] SERVICE ALERT: nycnn-ipvx28e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413403336] SERVICE ALERT: nycnn-ipvx29e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403336] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403346] SERVICE ALERT: nycnn-ipvx29;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403346] SERVICE ALERT: nycnn-ipvx30;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413403356] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403356] HOST ALERT: nycnn-ipvx28;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413403356] SERVICE ALERT: nycnn-ipvx27e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413403356] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413403366] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413403366] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1413403366] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403366] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.25 ms
[1413403366] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.21 ms
[1413403376] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413403376] SERVICE ALERT: nycnn-ipvx27;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413403376] SERVICE ALERT: nycnn-ipvx27;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413403386] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.842GB (42%) - Free: 1.148GB (58%)
[1413403386] SERVICE ALERT: nycnn-ipvx27;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.841GB (42%) - Free: 1.148GB (58%)
[1413403386] SERVICE ALERT: nycnn-ipvx27e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403386] SERVICE ALERT: nycnn-ipvx27;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413403386] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413403386] HOST ALERT: nycnn-ipvx29;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.48 ms
[1413403386] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.47 ms
[1413403386] SERVICE ALERT: nycnn-ipvx29;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 37%, RTA = 0.31 ms
[1413403396] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413403396] SERVICE ALERT: nycnn-ipvx30e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403406] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1413403406] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413403416] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403426] SERVICE ALERT: nycnn-ipvx30;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403426] SERVICE ALERT: nycnn-ipvx30;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403426] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403426] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403436] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403446] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1413403466] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403466] SERVICE ALERT: nycnn-ipvx30;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403486] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413403496] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;2;(Host Check Timed Out)
[1413403526] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403546] SERVICE ALERT: nycnn-ipvx28e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.849GB (43%) - Free: 1.141GB (57%)
[1413403546] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413403546] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413403556] HOST ALERT: nycnn-ipvx30;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413403556] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413403596] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403606] SERVICE ALERT: nycnn-ipvx28;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413403606] SERVICE ALERT: nycnn-ipvx29;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413403616] SERVICE ALERT: nycnn-ipvx29;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413403626] SERVICE ALERT: nycnn-ipvx29e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403626] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413403626] SERVICE ALERT: nycnn-ipvx29;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403636] SERVICE ALERT: nycnn-ipvx28e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413403686] SERVICE ALERT: nycnn-ipvx30e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413403706] SERVICE ALERT: nycnn-ipvx30;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413403716] SERVICE ALERT: nycnn-ipvx30;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403716] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403756] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413403756] SERVICE ALERT: nycnn-ipvx30;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413403766] SERVICE ALERT: nycnn-ipvx30;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.11%
[1413403816] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413403856] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.20 ms
[1413403866] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413403876] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413403876] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 4.81%
[1413403876] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403876] SERVICE ALERT: nycnn-ipvx32e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403896] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1413403896] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413403896] SERVICE ALERT: nycnn-ipvx32;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403906] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413403916] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413403916] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1413403916] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403926] SERVICE ALERT: nycnn-ipvx31;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403926] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403926] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403936] SERVICE ALERT: nycnn-ipvx32e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403936] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403936] SERVICE ALERT: nycnn-ipvx31e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403946] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413403946] SERVICE ALERT: nycnn-ipvx32e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403956] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413403956] SERVICE ALERT: nycnn-ipvx32;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413403966] SERVICE ALERT: nycnn-ipvx32e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403966] SERVICE ALERT: nycnn-ipvx32;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403976] SERVICE ALERT: nycnn-ipvx32;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413403996] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;2;(Host Check Timed Out)
[1413403996] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.43 ms
[1413403996] SERVICE ALERT: nycnn-ipvx32;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413404006] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413404016] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413404016] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413404026] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413404036] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1413404036] HOST ALERT: nycnn-ipvx31;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413404046] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1413404046] SERVICE ALERT: nycnn-ipvx32;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 1.990GB - Used: 0.834GB (42%) - Free: 1.156GB (58%)
[1413404046] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413404086] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413404086] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413404106] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404116] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413404116] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1413404146] SERVICE ALERT: nycnn-ipvx33;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404146] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413404146] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413404156] SERVICE ALERT: nycnn-ipvx33;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404156] SERVICE ALERT: nycnn-ipvx33e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404166] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413404186] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413404196] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;2;(Host Check Timed Out)
[1413404206] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1413404206] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1413404206] SERVICE ALERT: nycnn-ipvx33;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413404206] SERVICE ALERT: nycnn-ipvx31;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.857GB (43%) - Free: 1.133GB (57%)
[1413404216] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413404216] HOST ALERT: nycnn-ipvx33;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413404216] SERVICE ALERT: nycnn-ipvx32e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404216] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404226] SERVICE ALERT: nycnn-ipvx31e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.854GB (43%) - Free: 1.136GB (57%)
[1413404226] SERVICE ALERT: nycnn-ipvx33;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404236] SERVICE ALERT: nycnn-ipvx32e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413404236] SERVICE ALERT: nycnn-ipvx32;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404256] SERVICE ALERT: nycnn-ipvx32e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413404256] SERVICE ALERT: nycnn-ipvx32;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413404256] SERVICE ALERT: nycnn-ipvx32;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413404266] SERVICE ALERT: nycnn-ipvx33;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 1.990GB - Used: 0.857GB (43%) - Free: 1.132GB (57%)
[1413404266] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413404266] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413404296] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404306] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1413404306] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413404306] SERVICE ALERT: nycnn-ipvx34e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404326] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413404326] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413404346] SERVICE ALERT: nycnn-ipvx34e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404346] SERVICE ALERT: nycnn-ipvx34;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404356] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413404386] SERVICE ALERT: nycnn-ipvx34;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413404396] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.34 ms
[1413404396] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.22 ms
[1413404436] SERVICE ALERT: nycnn-ipvx33;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413404436] SERVICE ALERT: nycnn-ipvx33;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404436] SERVICE ALERT: nycnn-ipvx33e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413404446] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413404446] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413404456] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404506] SERVICE ALERT: nycnn-ipvx33;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413404596] SERVICE ALERT: nycnn-ipvx34e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404636] SERVICE ALERT: nycnn-ipvx34e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413404636] SERVICE ALERT: nycnn-ipvx34;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.841GB (42%) - Free: 1.149GB (58%)
[1413404636] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404686] SERVICE ALERT: nycnn-ipvx34;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413404926] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413404926] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404946] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404946] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413404946] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404956] SERVICE ALERT: nycnn-ipvx35;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413404956] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1413404986] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413404986] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413404986] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413405006] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405006] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413405006] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405016] SERVICE ALERT: nycnn-ipvx35;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405036] SERVICE ALERT: nycnn-ipvx26e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413405046] SERVICE ALERT: nycnn-ipvx35;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.25 ms
[1413405046] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.29 ms
[1413405046] SERVICE ALERT: nycnn-ipvx35e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413405056] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413405066] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1413405076] SERVICE ALERT: nycnn-ipvx26e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405086] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405096] SERVICE ALERT: nycnn-ipvx26;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405106] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413405106] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413405106] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413405106] HOST ALERT: nycnn-ipvx26e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413405116] SERVICE ALERT: nycnn-ipvx26e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405116] SERVICE ALERT: nycnn-ipvx26;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405136] HOST ALERT: nycnn-ipvx26;DOWN;SOFT;1;(Host Check Timed Out)
[1413405136] SERVICE ALERT: nycnn-ipvx26e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405146] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405146] SERVICE ALERT: nycnn-ipvx26e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405156] SERVICE ALERT: nycnn-ipvx26;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405166] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413405166] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413405176] SERVICE ALERT: nycnn-ipvx26;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413405186] SERVICE ALERT: nycnn-ipvx26;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405196] HOST ALERT: nycnn-ipvx26e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.46 ms
[1413405206] HOST ALERT: nycnn-ipvx26;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413405226] SERVICE ALERT: nycnn-ipvx26e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413405226] SERVICE ALERT: nycnn-ipvx26;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413405276] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413405296] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;OK;HARD;2;OK (Sample Period 659 sec) - Average CPU Utilisation 0.00%
[1413405296] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413405296] SERVICE ALERT: nycnn-ipvx35;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413405336] SERVICE ALERT: nycnn-ipvx35e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413405396] SERVICE ALERT: nycnn-ipvx26e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413405416] SERVICE ALERT: nycnn-ipvx26e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.857GB (43%) - Free: 1.133GB (57%)
[1413405426] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413405436] SERVICE ALERT: nycnn-ipvx26e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413405436] SERVICE ALERT: nycnn-ipvx26;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413405466] SERVICE ALERT: nycnn-ipvx26;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1413405466] SERVICE ALERT: nycnn-ipvx26;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413405486] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.22%
[1413405526] SERVICE ALERT: nycnn-ipvx26;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.16%
[1413405966] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405966] SERVICE ALERT: nycnn-ipvx24;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413405966] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413405976] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413406006] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413406006] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1413406006] SERVICE ALERT: nycnn-ipvx24;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413406006] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1413406026] SERVICE ALERT: nycnn-ipvx24;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413406026] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413406026] SERVICE ALERT: nycnn-ipvx24;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413406026] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413406036] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413406096] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.42 ms
[1413406096] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.44 ms
[1413406126] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413406146] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1413406286] SERVICE ALERT: nycnn-ipvx24;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413406306] SERVICE ALERT: nycnn-ipvx24;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.850GB (43%) - Free: 1.139GB (57%)
[1413406316] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413406316] SERVICE ALERT: nycnn-ipvx24;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413406316] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413408352] HOST DOWNTIME ALERT: nycnn-ipvx1;STOPPED; Host has exited from a period of scheduled downtime
[1413408402] HOST DOWNTIME ALERT: nycnn-ipvx10;STOPPED; Host has exited from a period of scheduled downtime
[1413408633] HOST DOWNTIME ALERT: nycnn-ipvx11;STOPPED; Host has exited from a period of scheduled downtime
[1413408653] HOST DOWNTIME ALERT: nycnn-ipvx12;STOPPED; Host has exited from a period of scheduled downtime
[1413408669] HOST DOWNTIME ALERT: nycnn-ipvx13;STOPPED; Host has exited from a period of scheduled downtime
[1413408691] HOST DOWNTIME ALERT: nycnn-ipvx14;STOPPED; Host has exited from a period of scheduled downtime
[1413408711] HOST DOWNTIME ALERT: nycnn-ipvx15;STOPPED; Host has exited from a period of scheduled downtime
[1413408733] HOST DOWNTIME ALERT: nycnn-ipvx16;STOPPED; Host has exited from a period of scheduled downtime
[1413408750] HOST DOWNTIME ALERT: nycnn-ipvx17;STOPPED; Host has exited from a period of scheduled downtime
[1413408764] HOST DOWNTIME ALERT: nycnn-ipvx18;STOPPED; Host has exited from a period of scheduled downtime
[1413408785] HOST DOWNTIME ALERT: nycnn-ipvx19;STOPPED; Host has exited from a period of scheduled downtime
[1413408806] HOST DOWNTIME ALERT: nycnn-ipvx1e1;STOPPED; Host has exited from a period of scheduled downtime
[1413408826] HOST DOWNTIME ALERT: nycnn-ipvx2;STOPPED; Host has exited from a period of scheduled downtime
[1413408846] HOST DOWNTIME ALERT: nycnn-ipvx20;STOPPED; Host has exited from a period of scheduled downtime
[1413408869] HOST DOWNTIME ALERT: nycnn-ipvx21;STOPPED; Host has exited from a period of scheduled downtime
[1413408887] HOST DOWNTIME ALERT: nycnn-ipvx22;STOPPED; Host has exited from a period of scheduled downtime
[1413408904] HOST DOWNTIME ALERT: nycnn-ipvx23;STOPPED; Host has exited from a period of scheduled downtime
[1413408925] HOST DOWNTIME ALERT: nycnn-ipvx24;STOPPED; Host has exited from a period of scheduled downtime
[1413408941] HOST DOWNTIME ALERT: nycnn-ipvx25;STOPPED; Host has exited from a period of scheduled downtime
[1413408984] HOST DOWNTIME ALERT: nycnn-ipvx26;STOPPED; Host has exited from a period of scheduled downtime
[1413409011] HOST DOWNTIME ALERT: nycnn-ipvx27;STOPPED; Host has exited from a period of scheduled downtime
[1413409033] HOST DOWNTIME ALERT: nycnn-ipvx28;STOPPED; Host has exited from a period of scheduled downtime
[1413409062] HOST DOWNTIME ALERT: nycnn-ipvx29;STOPPED; Host has exited from a period of scheduled downtime
[1413409084] HOST DOWNTIME ALERT: nycnn-ipvx3;STOPPED; Host has exited from a period of scheduled downtime
[1413409105] HOST DOWNTIME ALERT: nycnn-ipvx30;STOPPED; Host has exited from a period of scheduled downtime
[1413409124] HOST DOWNTIME ALERT: nycnn-ipvx31;STOPPED; Host has exited from a period of scheduled downtime
[1413409194] HOST DOWNTIME ALERT: nycnn-ipvx32;STOPPED; Host has exited from a period of scheduled downtime
[1413409213] HOST DOWNTIME ALERT: nycnn-ipvx33;STOPPED; Host has exited from a period of scheduled downtime
[1413409251] HOST DOWNTIME ALERT: nycnn-ipvx34;STOPPED; Host has exited from a period of scheduled downtime
[1413409270] HOST DOWNTIME ALERT: nycnn-ipvx35;STOPPED; Host has exited from a period of scheduled downtime
[1413409290] HOST DOWNTIME ALERT: nycnn-ipvx4;STOPPED; Host has exited from a period of scheduled downtime
[1413409314] HOST DOWNTIME ALERT: nycnn-ipvx5;STOPPED; Host has exited from a period of scheduled downtime
[1413409332] HOST DOWNTIME ALERT: nycnn-ipvx6;STOPPED; Host has exited from a period of scheduled downtime
[1413409356] HOST DOWNTIME ALERT: nycnn-ipvx8;STOPPED; Host has exited from a period of scheduled downtime
[1413409376] HOST DOWNTIME ALERT: nycnn-ipvx9;STOPPED; Host has exited from a period of scheduled downtime
[1413416826] SERVICE ALERT: nycnn-ipvx25;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416836] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413416846] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416866] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1413416866] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416866] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1413416866] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413416886] SERVICE ALERT: nycnn-ipvx25;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413416896] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413416906] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413416926] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;2;(Host Check Timed Out)
[1413416926] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413416956] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416956] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;2;(Host Check Timed Out)
[1413416976] SERVICE ALERT: nycnn-ipvx25e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416976] SERVICE ALERT: nycnn-ipvx25;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416986] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416996] SERVICE ALERT: nycnn-ipvx25;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413416996] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 0.50 ms
[1413417006] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413417026] HOST ALERT: nycnn-ipvx25;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413417046] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1413417166] SERVICE ALERT: nycnn-ipvx25;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413417196] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413417206] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413417236] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413417266] SERVICE ALERT: nycnn-ipvx25e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413417266] SERVICE ALERT: nycnn-ipvx25;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413417276] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 1.57%
[1413417276] SERVICE ALERT: nycnn-ipvx25;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.878GB (44%) - Free: 1.112GB (56%)
[1413417406] SERVICE ALERT: nycnn-ipvx25e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413417406] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413417406] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413417436] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413417446] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1413417446] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1413417466] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413417466] SERVICE ALERT: nycnn-ipvx25e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413417466] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413417486] SERVICE ALERT: nycnn-ipvx25;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413417496] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413417506] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413417526] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413417526] HOST ALERT: nycnn-ipvx25;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.41 ms
[1413417526] HOST FLAPPING ALERT: nycnn-ipvx25;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1413417526] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 1.51 ms
[1413417526] HOST FLAPPING ALERT: nycnn-ipvx25e1;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1413417526] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1413417526] SERVICE FLAPPING ALERT: nycnn-ipvx25;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1413417546] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1413417546] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;PING;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1413417756] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413417756] SERVICE ALERT: nycnn-ipvx25e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.878GB (44%) - Free: 1.112GB (56%)
[1413417766] SERVICE ALERT: nycnn-ipvx25;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413417766] SERVICE FLAPPING ALERT: nycnn-ipvx25;Process: XCode;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1413417796] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1413417796] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;Process: XCode;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1413417806] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1413417806] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;Service: Apache;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1413418666] SERVICE FLAPPING ALERT: nycnn-ipvx25;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1413418686] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1413419736] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;SOFT;1;Job 10522730 remains in a queued state
[1413419796] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;SOFT;2;Job 10522730 progress is  0 :
[1413420586] HOST FLAPPING ALERT: nycnn-ipvx25;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1413420606] HOST FLAPPING ALERT: nycnn-ipvx25e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1413423466] SERVICE FLAPPING ALERT: nycnn-ipvx25;Process: XCode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1413423496] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;Process: XCode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1413423506] SERVICE FLAPPING ALERT: nycnn-ipvx25e1;Service: Apache;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)