[1340512195] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1340512255] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1340512315] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1340515435] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.005 second response time
[1340518055] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518095] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;1;(Host Check Timed Out)
[1340518115] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518115] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518155] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1340518165] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;2;(Host Check Timed Out)
[1340518175] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518175] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518215] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;1;(Host Check Timed Out)
[1340518225] SERVICE ALERT: nycnn-ipvsphn4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1340518235] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518235] HOST ALERT: nycnn-ipvsphn4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1340518245] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1340518265] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.21 ms
[1340518295] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518295] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518295] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518295] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518295] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518295] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.34 ms
[1340518295] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.23 ms
[1340518295] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518305] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1340518305] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;2;(Host Check Timed Out)
[1340518325] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518325] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518335] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1340518335] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518335] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1340518335] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518335] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518335] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1340518335] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518335] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1340518345] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;3;(Host Check Timed Out)
[1340518345] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1340518355] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;2;(Host Check Timed Out)
[1340518355] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1340518355] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518355] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.24 ms
[1340518355] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.19 ms
[1340518355] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518355] HOST ALERT: nycnn-adbeame4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1340518355] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518385] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518395] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1340518395] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518395] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1340518395] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518395] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 1.30 ms
[1340518395] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;1;(Host Check Timed Out)
[1340518405] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1340518405] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.98 ms
[1340518405] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1340518405] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1340518405] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1340518405] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1340518415] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;1;(Host Check Timed Out)
[1340518415] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;1;(Host Check Timed Out)
[1340518415] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518415] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518415] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1340518415] HOST ALERT: nycnn-ipvx29;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.38 ms
[1340518415] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.34 ms
[1340518415] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.49 ms
[1340518415] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.49 ms
[1340518425] HOST ALERT: nycnn-ipvx3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1340518425] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.52 ms
[1340518425] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.50 ms
[1340518435] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1340518435] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.51 ms
[1340518435] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.56 ms
[1340518435] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1340518435] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1340518445] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518445] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1340518465] HOST ALERT: nycnn-ipvx15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1340518465] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1340518465] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1340518465] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1340518465] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1340518465] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1340518465] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1340518465] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1340518465] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1340518465] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1340518465] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1340518475] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518475] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518475] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518485] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.49 ms
[1340518485] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.48 ms
[1340518495] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;2;(Host Check Timed Out)
[1340518515] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;2;(Host Check Timed Out)
[1340518515] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;2;(Host Check Timed Out)
[1340518515] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1340518515] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.64 ms
[1340518525] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1340518525] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1340518525] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1340518525] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1340518525] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1340518575] HOST ALERT: nycnn-adbeame7;UP;SOFT;3;PING OK - Packet loss = 66%, RTA = 0.46 ms
[1340518585] HOST ALERT: nycnn-adbeame2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1340518585] HOST ALERT: nycnn-adbeame9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1340518585] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1340518585] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1340518585] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1340518585] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1340518595] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1340518835] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518835] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340518875] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340518875] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1340518895] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518895] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340518945] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1340518945] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1340518945] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1340518945] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms