[1402583046] SERVICE ALERT: nycnn-ipvapilic;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402583106] SERVICE ALERT: nycnn-ipvapilic;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.00%
[1402587356] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402587416] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.36%
[1402595255] HOST DOWNTIME ALERT: nycnn-ipvx9e1;STARTED; Host has entered a period of scheduled downtime
[1402595278] HOST DOWNTIME ALERT: nycnn-ipvx9;STARTED; Host has entered a period of scheduled downtime
[1402595302] HOST DOWNTIME ALERT: nycnn-ipvx1;STARTED; Host has entered a period of scheduled downtime
[1402595325] HOST DOWNTIME ALERT: nycnn-ipvx10;STARTED; Host has entered a period of scheduled downtime
[1402595343] HOST DOWNTIME ALERT: nycnn-ipvx10e1;STARTED; Host has entered a period of scheduled downtime
[1402595360] HOST DOWNTIME ALERT: nycnn-ipvx11;STARTED; Host has entered a period of scheduled downtime
[1402595380] HOST DOWNTIME ALERT: nycnn-ipvx11e1;STARTED; Host has entered a period of scheduled downtime
[1402595398] HOST DOWNTIME ALERT: nycnn-ipvx12;STARTED; Host has entered a period of scheduled downtime
[1402595414] HOST DOWNTIME ALERT: nycnn-ipvx12e1;STARTED; Host has entered a period of scheduled downtime
[1402595430] HOST DOWNTIME ALERT: nycnn-ipvx13;STARTED; Host has entered a period of scheduled downtime
[1402595445] HOST DOWNTIME ALERT: nycnn-ipvx13e1;STARTED; Host has entered a period of scheduled downtime
[1402595459] HOST DOWNTIME ALERT: nycnn-ipvx14;STARTED; Host has entered a period of scheduled downtime
[1402595486] HOST DOWNTIME ALERT: nycnn-ipvx14e1;STARTED; Host has entered a period of scheduled downtime
[1402595502] HOST DOWNTIME ALERT: nycnn-ipvx15;STARTED; Host has entered a period of scheduled downtime
[1402595521] HOST DOWNTIME ALERT: nycnn-ipvx15e1;STARTED; Host has entered a period of scheduled downtime
[1402595551] HOST DOWNTIME ALERT: nycnn-ipvx16;STARTED; Host has entered a period of scheduled downtime
[1402595568] HOST DOWNTIME ALERT: nycnn-ipvx16e1;STARTED; Host has entered a period of scheduled downtime
[1402595599] HOST DOWNTIME ALERT: nycnn-ipvx17;STARTED; Host has entered a period of scheduled downtime
[1402595616] HOST DOWNTIME ALERT: nycnn-ipvx17e1;STARTED; Host has entered a period of scheduled downtime
[1402595646] HOST DOWNTIME ALERT: nycnn-ipvx18;STARTED; Host has entered a period of scheduled downtime
[1402595665] HOST DOWNTIME ALERT: nycnn-ipvx18e1;STARTED; Host has entered a period of scheduled downtime
[1402595681] HOST DOWNTIME ALERT: nycnn-ipvx19;STARTED; Host has entered a period of scheduled downtime
[1402595695] HOST DOWNTIME ALERT: nycnn-ipvx19e1;STARTED; Host has entered a period of scheduled downtime
[1402595715] HOST DOWNTIME ALERT: nycnn-ipvx1e1;STARTED; Host has entered a period of scheduled downtime
[1402595730] HOST DOWNTIME ALERT: nycnn-ipvx2;STARTED; Host has entered a period of scheduled downtime
[1402595763] HOST DOWNTIME ALERT: nycnn-ipvx20;STARTED; Host has entered a period of scheduled downtime
[1402595778] HOST DOWNTIME ALERT: nycnn-ipvx20e1;STARTED; Host has entered a period of scheduled downtime
[1402595796] HOST DOWNTIME ALERT: nycnn-ipvx21;STARTED; Host has entered a period of scheduled downtime
[1402595813] HOST DOWNTIME ALERT: nycnn-ipvx21e1;STARTED; Host has entered a period of scheduled downtime
[1402595831] HOST DOWNTIME ALERT: nycnn-ipvx22;STARTED; Host has entered a period of scheduled downtime
[1402596266] SERVICE ALERT: nycnn-ipvx1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402596346] SERVICE ALERT: nycnn-ipvx1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596346] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402596346] SERVICE ALERT: nycnn-ipvx1;Service: Apache;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1402596346] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402596346] SERVICE ALERT: nycnn-ipvx1e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596366] SERVICE ALERT: nycnn-ipvx1e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596376] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1402596376] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402596406] SERVICE ALERT: nycnn-ipvx1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402596406] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402596406] SERVICE ALERT: nycnn-ipvx1;Service: Apache;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1402596406] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402596406] SERVICE ALERT: nycnn-ipvx1e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402596416] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402596426] SERVICE ALERT: nycnn-ipvx1e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402596446] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402596456] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402596456] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402596466] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402596497] HOST DOWNTIME ALERT: nycnn-ipvx3;STARTED; Host has entered a period of scheduled downtime
[1402596517] HOST DOWNTIME ALERT: nycnn-ipvx4;STARTED; Host has entered a period of scheduled downtime
[1402596533] HOST DOWNTIME ALERT: nycnn-ipvx5;STARTED; Host has entered a period of scheduled downtime
[1402596555] HOST DOWNTIME ALERT: nycnn-ipvx6;STARTED; Host has entered a period of scheduled downtime
[1402596575] HOST DOWNTIME ALERT: nycnn-ipvx5e1;STARTED; Host has entered a period of scheduled downtime
[1402596593] HOST DOWNTIME ALERT: nycnn-ipvx6e1;STARTED; Host has entered a period of scheduled downtime
[1402596636] SERVICE ALERT: nycnn-ipvx2;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402596636] SERVICE ALERT: nycnn-ipvx2e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402596636] SERVICE ALERT: nycnn-ipvx2;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402596646] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.23 ms
[1402596646] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.25 ms
[1402596676] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1402596676] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402596686] SERVICE ALERT: nycnn-ipvx1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402596686] SERVICE ALERT: nycnn-ipvx1;Service: Apache;OK;HARD;3;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402596696] SERVICE ALERT: nycnn-ipvx1e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402596706] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402596706] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402596706] SERVICE ALERT: nycnn-ipvx2;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402596706] SERVICE ALERT: nycnn-ipvx2e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402596706] SERVICE ALERT: nycnn-ipvx2;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402596706] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;2;(Host Check Timed Out)
[1402596716] SERVICE ALERT: nycnn-ipvx1e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402596726] SERVICE ALERT: nycnn-ipvx2e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596736] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402596736] SERVICE ALERT: nycnn-ipvx2;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596776] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596776] SERVICE ALERT: nycnn-ipvx2;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596786] HOST ALERT: nycnn-ipvx2;UP;SOFT;3;PING OK - Packet loss = 54%, RTA = 16.45 ms
[1402596786] SERVICE ALERT: nycnn-ipvx2e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402596806] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402596813] HOST DOWNTIME ALERT: nycnn-ipvx8;STARTED; Host has entered a period of scheduled downtime
[1402596816] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402596816] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402596845] HOST DOWNTIME ALERT: nycnn-ipvx8e1;STARTED; Host has entered a period of scheduled downtime
[1402596868] HOST DOWNTIME ALERT: nycnn-ipvx4e1;STARTED; Host has entered a period of scheduled downtime
[1402596887] HOST DOWNTIME ALERT: nycnn-ipvx3e1;STARTED; Host has entered a period of scheduled downtime
[1402596996] SERVICE ALERT: nycnn-ipvx2;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.815GB (41%) - Free: 1.175GB (59%)
[1402596996] SERVICE ALERT: nycnn-ipvx2;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402596996] SERVICE ALERT: nycnn-ipvx2e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402597006] SERVICE ALERT: nycnn-ipvx2e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.813GB (41%) - Free: 1.177GB (59%)
[1402597016] SERVICE ALERT: nycnn-ipvx2;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402597056] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597066] SERVICE ALERT: nycnn-ipvx2;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597066] SERVICE ALERT: nycnn-ipvx2e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597166] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597186] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597186] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597206] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1402597226] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402597226] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597246] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402597246] SERVICE ALERT: nycnn-ipvx3;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597246] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402597246] SERVICE ALERT: nycnn-ipvx3;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597246] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597246] SERVICE ALERT: nycnn-ipvx3e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597276] SERVICE ALERT: nycnn-ipvx3e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597286] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.41 ms
[1402597296] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402597296] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402597296] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402597506] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597536] SERVICE ALERT: nycnn-ipvx3;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402597536] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597536] SERVICE ALERT: nycnn-ipvx3;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402597536] SERVICE ALERT: nycnn-ipvx3e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597546] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597546] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597546] SERVICE ALERT: nycnn-ipvx4;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597556] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402597556] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402597556] SERVICE ALERT: nycnn-ipvx3e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402597566] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597576] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1402597576] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402597606] SERVICE ALERT: nycnn-ipvx4e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597606] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597606] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402597606] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597606] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402597606] SERVICE ALERT: nycnn-ipvx4;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597616] SERVICE ALERT: nycnn-ipvx4e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597616] SERVICE ALERT: nycnn-ipvx5e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597616] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597626] SERVICE ALERT: nycnn-ipvx5;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597626] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597626] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597636] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597646] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402597646] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1402597646] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 57.30 ms
[1402597646] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 4.54 ms
[1402597656] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402597656] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402597666] SERVICE ALERT: nycnn-ipvx5;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597666] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402597666] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402597676] SERVICE ALERT: nycnn-ipvx5e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597676] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597686] SERVICE ALERT: nycnn-ipvx5;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597686] SERVICE ALERT: nycnn-ipvx5;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597696] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1402597706] SERVICE ALERT: nycnn-ipvx5e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402597706] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 1.76 ms
[1402597716] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402597716] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402597716] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402597886] SERVICE ALERT: nycnn-ipvx4e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597896] SERVICE ALERT: nycnn-ipvx4;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402597896] SERVICE ALERT: nycnn-ipvx4e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402597916] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597916] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597916] SERVICE ALERT: nycnn-ipvx6;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402597946] SERVICE ALERT: nycnn-ipvx5;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402597956] SERVICE ALERT: nycnn-ipvx5e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402597966] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597966] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597966] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597966] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402597966] SERVICE ALERT: nycnn-ipvx5;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402597976] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402597976] SERVICE ALERT: nycnn-ipvx5;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.807GB (41%) - Free: 1.183GB (59%)
[1402597976] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;OK;HARD;3;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402597986] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1402597986] SERVICE ALERT: nycnn-ipvx6e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597996] SERVICE ALERT: nycnn-ipvx6;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402597996] SERVICE ALERT: nycnn-ipvx6;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402597996] SERVICE ALERT: nycnn-ipvx5e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.807GB (41%) - Free: 1.183GB (59%)
[1402598016] SERVICE ALERT: nycnn-ipvx6e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598026] SERVICE ALERT: nycnn-ipvx6e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598026] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598026] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402598026] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402598066] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.50 ms
[1402598066] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.40 ms
[1402598076] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1402598076] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402598086] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402598146] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402598146] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402598146] SERVICE ALERT: nycnn-ipvx8;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598166] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598176] SERVICE ALERT: nycnn-ipvx8e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598176] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598186] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402598186] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1402598186] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598206] SERVICE ALERT: nycnn-ipvx8;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598206] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402598206] SERVICE ALERT: nycnn-ipvx8;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598206] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402598216] SERVICE ALERT: nycnn-ipvx8;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598226] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1402598226] SERVICE ALERT: nycnn-ipvx8;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598236] SERVICE ALERT: nycnn-ipvx8e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598236] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598256] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402598256] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1402598256] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402598256] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1402598266] SERVICE ALERT: nycnn-ipvx6e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402598276] SERVICE ALERT: nycnn-ipvx6;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402598276] SERVICE ALERT: nycnn-ipvx6;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402598296] SERVICE ALERT: nycnn-ipvx6e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402598306] SERVICE ALERT: nycnn-ipvx6e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.808GB (41%) - Free: 1.182GB (59%)
[1402598306] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402598476] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402598486] SERVICE ALERT: nycnn-ipvx8;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402598496] SERVICE ALERT: nycnn-ipvx8;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402598496] SERVICE ALERT: nycnn-ipvx8;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402598506] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;OK;HARD;3;OK (Sample Period 720 sec) - Average CPU Utilisation 0.00%
[1402598516] SERVICE ALERT: nycnn-ipvx8;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 1.027GB (52%) - Free: 0.963GB (48%)
[1402598516] SERVICE ALERT: nycnn-ipvx8e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402598516] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402598636] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402598636] SERVICE ALERT: nycnn-ipvx10;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402598656] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598676] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402598676] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1402598686] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402598686] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402598716] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598716] SERVICE ALERT: nycnn-ipvx10;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598716] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598736] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402598736] SERVICE ALERT: nycnn-ipvx10e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598746] SERVICE ALERT: nycnn-ipvx10;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598746] SERVICE ALERT: nycnn-ipvx10e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598746] SERVICE ALERT: nycnn-ipvx10e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598766] SERVICE ALERT: nycnn-ipvx10;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598766] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;2;(Host Check Timed Out)
[1402598786] SERVICE ALERT: nycnn-ipvx10;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402598796] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402598796] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402598806] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402598806] HOST ALERT: nycnn-ipvx10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402598846] SERVICE ALERT: nycnn-ipvx9;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402598866] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402598866] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402598866] SERVICE ALERT: nycnn-ipvx9e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598886] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1402598896] SERVICE ALERT: nycnn-ipvx9;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598896] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402598916] SERVICE ALERT: nycnn-ipvx9e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598926] SERVICE ALERT: nycnn-ipvx9;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598926] SERVICE ALERT: nycnn-ipvx9;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598926] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402598926] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402598926] SERVICE ALERT: nycnn-ipvx9e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402598966] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402598966] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402598986] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;2;(Host Check Timed Out)
[1402598996] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402598996] SERVICE ALERT: nycnn-ipvx10;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599006] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599026] SERVICE ALERT: nycnn-ipvx10e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402599036] SERVICE ALERT: nycnn-ipvx10;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.822GB (41%) - Free: 1.168GB (59%)
[1402599036] SERVICE ALERT: nycnn-ipvx10e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599036] SERVICE ALERT: nycnn-ipvx10e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402599036] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402599036] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402599036] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402599046] HOST ALERT: nycnn-ipvx9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402599046] SERVICE ALERT: nycnn-ipvx10;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402599086] SERVICE ALERT: nycnn-ipvx10;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402599106] SERVICE ALERT: nycnn-ipvx11;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599106] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402599106] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402599116] SERVICE ALERT: nycnn-ipvx11e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599136] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1402599136] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402599146] SERVICE ALERT: nycnn-ipvx11e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599166] SERVICE ALERT: nycnn-ipvx11;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599166] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402599166] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402599176] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;2;(Host Check Timed Out)
[1402599176] SERVICE ALERT: nycnn-ipvx11e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599176] SERVICE ALERT: nycnn-ipvx11e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599176] SERVICE ALERT: nycnn-ipvx9;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599186] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599196] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599196] SERVICE ALERT: nycnn-ipvx9e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402599206] SERVICE ALERT: nycnn-ipvx11;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599206] SERVICE ALERT: nycnn-ipvx11;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599206] SERVICE ALERT: nycnn-ipvx9;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402599206] SERVICE ALERT: nycnn-ipvx9;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402599216] SERVICE ALERT: nycnn-ipvx9e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599216] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.95 ms
[1402599216] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402599216] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1402599226] HOST ALERT: nycnn-ipvx11;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402599246] SERVICE ALERT: nycnn-ipvx12;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402599256] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402599286] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402599286] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402599306] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1402599316] SERVICE ALERT: nycnn-ipvx12;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599316] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402599346] SERVICE ALERT: nycnn-ipvx12e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599346] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402599346] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402599346] SERVICE ALERT: nycnn-ipvx12e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599346] SERVICE ALERT: nycnn-ipvx12;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599356] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599366] SERVICE ALERT: nycnn-ipvx12;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599406] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;2;(Host Check Timed Out)
[1402599406] SERVICE ALERT: nycnn-ipvx12;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402599406] SERVICE ALERT: nycnn-ipvx12e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599406] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.32 ms
[1402599416] SERVICE ALERT: nycnn-ipvx12e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402599436] SERVICE ALERT: nycnn-ipvx11e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402599436] HOST ALERT: nycnn-ipvx12;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402599446] SERVICE ALERT: nycnn-ipvx11;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402599456] SERVICE ALERT: nycnn-ipvx11e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402599456] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402599456] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402599466] SERVICE ALERT: nycnn-ipvx11e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599476] SERVICE ALERT: nycnn-ipvx12e1;Service: Apache;OK;SOFT;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402599476] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599476] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599486] SERVICE ALERT: nycnn-ipvx11;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.837GB (42%) - Free: 1.153GB (58%)
[1402599496] SERVICE ALERT: nycnn-ipvx11;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402599576] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job 10211798 remains in a queued state
[1402599606] SERVICE ALERT: nycnn-ipvx12;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599626] SERVICE ALERT: nycnn-ipvx12e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402599636] SERVICE ALERT: nycnn-ipvx13;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402599636] SERVICE ALERT: nycnn-ipvx12e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599636] SERVICE ALERT: nycnn-ipvx12;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402599636] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;SOFT;2;Job 10211798 progress is  90.91 :
[1402599646] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402599656] SERVICE ALERT: nycnn-ipvx13e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402599656] SERVICE ALERT: nycnn-ipvx12;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402599686] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599686] SERVICE ALERT: nycnn-ipvx12e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402599696] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402599706] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402599706] SERVICE ALERT: nycnn-ipvx12;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.821GB (41%) - Free: 1.169GB (59%)
[1402599706] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402599706] SERVICE ALERT: nycnn-ipvx13;Memory Usage;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599706] SERVICE ALERT: nycnn-ipvx13e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599716] SERVICE ALERT: nycnn-ipvx14;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402599716] SERVICE ALERT: nycnn-ipvx13e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599726] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1402599726] SERVICE ALERT: nycnn-ipvx13e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599736] SERVICE ALERT: nycnn-ipvx14;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599736] SERVICE ALERT: nycnn-ipvx13e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599746] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599746] SERVICE ALERT: nycnn-ipvx13;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599746] SERVICE ALERT: nycnn-ipvx14e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599756] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1402599766] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402599766] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402599766] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402599766] SERVICE ALERT: nycnn-ipvx13;Memory Usage;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1402599776] SERVICE ALERT: nycnn-ipvx13;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599786] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402599786] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402599786] SERVICE ALERT: nycnn-ipvx14e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599786] SERVICE ALERT: nycnn-ipvx13;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599786] SERVICE ALERT: nycnn-ipvx14;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599786] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;2;(Host Check Timed Out)
[1402599786] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;2;(Host Check Timed Out)
[1402599796] SERVICE ALERT: nycnn-ipvx14;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599796] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402599796] SERVICE ALERT: nycnn-ipvx13;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402599806] SERVICE ALERT: nycnn-ipvx14;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599806] SERVICE ALERT: nycnn-ipvx14e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402599806] SERVICE ALERT: nycnn-ipvx14e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402599816] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402599816] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402599826] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402599826] HOST ALERT: nycnn-ipvx13;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402599826] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402599856] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.47 ms
[1402599856] HOST ALERT: nycnn-ipvx14;UP;SOFT;3;PING OK - Packet loss = 28%, RTA = 0.44 ms
[1402599876] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402599876] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402599996] SERVICE ALERT: nycnn-ipvx13e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402600006] SERVICE ALERT: nycnn-ipvx13e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.829GB (42%) - Free: 1.161GB (58%)
[1402600016] SERVICE ALERT: nycnn-ipvx13e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600016] SERVICE ALERT: nycnn-ipvx13e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402600026] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402600036] SERVICE ALERT: nycnn-ipvx13;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402600056] SERVICE ALERT: nycnn-ipvx13;Memory Usage;OK;HARD;3;OK - Physical Memory: Total: 1.990GB - Used: 0.826GB (42%) - Free: 1.164GB (58%)
[1402600066] SERVICE ALERT: nycnn-ipvx13;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402600066] SERVICE ALERT: nycnn-ipvx14e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402600066] SERVICE ALERT: nycnn-ipvx13;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402600076] SERVICE ALERT: nycnn-ipvx14;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600076] SERVICE ALERT: nycnn-ipvx14;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402600086] SERVICE ALERT: nycnn-ipvx14;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.817GB (41%) - Free: 1.173GB (59%)
[1402600096] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402600096] SERVICE ALERT: nycnn-ipvx13;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600096] SERVICE ALERT: nycnn-ipvx14e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600096] SERVICE ALERT: nycnn-ipvx14e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402600176] SERVICE ALERT: nycnn-ipvx15e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402600176] SERVICE ALERT: nycnn-ipvx15;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402600216] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402600216] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1402600246] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402600246] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402600246] SERVICE ALERT: nycnn-ipvx15e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600246] SERVICE ALERT: nycnn-ipvx15;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600246] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600266] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600306] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402600306] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402600306] SERVICE ALERT: nycnn-ipvx16e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600306] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600306] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402600306] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;2;(Host Check Timed Out)
[1402600316] SERVICE ALERT: nycnn-ipvx16;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600316] SERVICE ALERT: nycnn-ipvx15;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402600316] SERVICE ALERT: nycnn-ipvx16e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600326] SERVICE ALERT: nycnn-ipvx16;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600346] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1402600346] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402600356] SERVICE ALERT: nycnn-ipvx16;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600356] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1402600356] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402600366] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1402600366] HOST ALERT: nycnn-ipvx15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1402600366] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402600366] SERVICE ALERT: nycnn-ipvx16e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600366] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402600376] SERVICE ALERT: nycnn-ipvx16;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600376] SERVICE ALERT: nycnn-ipvx16e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600376] SERVICE ALERT: nycnn-ipvx16e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600386] SERVICE ALERT: nycnn-ipvx16;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600406] SERVICE ALERT: nycnn-ipvx16e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600406] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 1.97 ms
[1402600416] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1402600416] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402600416] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402600536] SERVICE ALERT: nycnn-ipvx15e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402600536] SERVICE ALERT: nycnn-ipvx15;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600536] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402600556] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600579] HOST DOWNTIME ALERT: nycnn-ipvx22e1;STARTED; Host has entered a period of scheduled downtime
[1402600596] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402600597] HOST DOWNTIME ALERT: nycnn-ipvx23;STARTED; Host has entered a period of scheduled downtime
[1402600606] SERVICE ALERT: nycnn-ipvx15;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.815GB (41%) - Free: 1.175GB (59%)
[1402600614] HOST DOWNTIME ALERT: nycnn-ipvx23e1;STARTED; Host has entered a period of scheduled downtime
[1402600636] SERVICE ALERT: nycnn-ipvx16;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402600644] HOST DOWNTIME ALERT: nycnn-ipvx24;STARTED; Host has entered a period of scheduled downtime
[1402600656] SERVICE ALERT: nycnn-ipvx16e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402600656] SERVICE ALERT: nycnn-ipvx16;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402600660] HOST DOWNTIME ALERT: nycnn-ipvx24e1;STARTED; Host has entered a period of scheduled downtime
[1402600666] SERVICE ALERT: nycnn-ipvx16e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402600666] SERVICE ALERT: nycnn-ipvx16e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.839GB (42%) - Free: 1.151GB (58%)
[1402600673] HOST DOWNTIME ALERT: nycnn-ipvx25;STARTED; Host has entered a period of scheduled downtime
[1402600676] SERVICE ALERT: nycnn-ipvx16;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600696] SERVICE ALERT: nycnn-ipvx16e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600716] SERVICE ALERT: nycnn-ipvx17;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402600733] HOST DOWNTIME ALERT: nycnn-ipvx25e1;STARTED; Host has entered a period of scheduled downtime
[1402600746] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1402600756] SERVICE ALERT: nycnn-ipvx17e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600786] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402600786] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402600786] SERVICE ALERT: nycnn-ipvx17;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600786] SERVICE ALERT: nycnn-ipvx17;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600786] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402600816] SERVICE ALERT: nycnn-ipvx17e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600846] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;2;(Host Check Timed Out)
[1402600846] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402600846] SERVICE ALERT: nycnn-ipvx17e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600846] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600846] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600866] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.52 ms
[1402600866] SERVICE ALERT: nycnn-ipvx17e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600876] SERVICE ALERT: nycnn-ipvx19;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402600896] SERVICE ALERT: nycnn-ipvx19e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402600896] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1402600896] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402600896] HOST DOWNTIME ALERT: nycnn-ipvx26;STARTED; Host has entered a period of scheduled downtime
[1402600906] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.21 ms
[1402600906] HOST ALERT: nycnn-ipvx17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402600906] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.21 ms
[1402600911] HOST DOWNTIME ALERT: nycnn-ipvx26e1;STARTED; Host has entered a period of scheduled downtime
[1402600916] SERVICE ALERT: nycnn-ipvx19;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600916] SERVICE ALERT: nycnn-ipvx17e1;Service: Apache;OK;SOFT;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402600926] SERVICE ALERT: nycnn-ipvx19e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600926] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402600936] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1402600945] HOST DOWNTIME ALERT: nycnn-ipvx27;STARTED; Host has entered a period of scheduled downtime
[1402600946] SERVICE ALERT: nycnn-ipvx19;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600956] SERVICE ALERT: nycnn-ipvx19e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600963] HOST DOWNTIME ALERT: nycnn-ipvx27e1;STARTED; Host has entered a period of scheduled downtime
[1402600966] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402600966] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402600966] SERVICE ALERT: nycnn-ipvx19e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600976] SERVICE ALERT: nycnn-ipvx19;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600976] SERVICE ALERT: nycnn-ipvx19;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402600986] SERVICE ALERT: nycnn-ipvx19e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402600990] HOST DOWNTIME ALERT: nycnn-ipvx28;STARTED; Host has entered a period of scheduled downtime
[1402601009] HOST DOWNTIME ALERT: nycnn-ipvx28e1;STARTED; Host has entered a period of scheduled downtime
[1402601026] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402601026] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;2;(Host Check Timed Out)
[1402601046] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402601076] SERVICE ALERT: nycnn-ipvx17;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402601076] SERVICE ALERT: nycnn-ipvx17;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402601076] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402601076] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402601086] HOST ALERT: nycnn-ipvx19;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402601086] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1402601106] SERVICE ALERT: nycnn-ipvx17e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601136] SERVICE ALERT: nycnn-ipvx17e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.638GB (32%) - Free: 1.352GB (68%)
[1402601136] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601136] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601196] SERVICE ALERT: nycnn-ipvx20e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601206] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402601206] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402601226] SERVICE ALERT: nycnn-ipvx20;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601236] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402601236] SERVICE ALERT: nycnn-ipvx19;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601236] SERVICE ALERT: nycnn-ipvx19e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402601246] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1402601256] SERVICE ALERT: nycnn-ipvx19e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402601256] SERVICE ALERT: nycnn-ipvx20;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601256] SERVICE ALERT: nycnn-ipvx19;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402601256] SERVICE ALERT: nycnn-ipvx20e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601266] SERVICE ALERT: nycnn-ipvx19;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601266] SERVICE ALERT: nycnn-ipvx19e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.831GB (42%) - Free: 1.159GB (58%)
[1402601266] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402601266] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402601286] SERVICE ALERT: nycnn-ipvx20;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601286] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601286] SERVICE ALERT: nycnn-ipvx20;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601306] SERVICE ALERT: nycnn-ipvx20;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402601306] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.44 ms
[1402601306] SERVICE ALERT: nycnn-ipvx20e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601316] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402601316] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402601316] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402601346] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601356] SERVICE ALERT: nycnn-ipvx20e1;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 1.990GB - Used: 0.831GB (42%) - Free: 1.159GB (58%)
[1402601436] SERVICE ALERT: nycnn-ipvx21;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402601436] SERVICE ALERT: nycnn-ipvx21;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402601446] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402601446] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402601466] SERVICE ALERT: nycnn-ipvx21;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601476] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1402601476] SERVICE ALERT: nycnn-ipvx21e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601486] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402601506] SERVICE ALERT: nycnn-ipvx21;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601506] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402601506] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402601506] SERVICE ALERT: nycnn-ipvx21;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601506] SERVICE ALERT: nycnn-ipvx21;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601526] SERVICE ALERT: nycnn-ipvx21e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601526] SERVICE ALERT: nycnn-ipvx21;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601536] SERVICE ALERT: nycnn-ipvx21e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601536] SERVICE ALERT: nycnn-ipvx20;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402601546] SERVICE ALERT: nycnn-ipvx20e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601566] SERVICE ALERT: nycnn-ipvx20;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.843GB (42%) - Free: 1.147GB (58%)
[1402601576] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;2;(Host Check Timed Out)
[1402601576] SERVICE ALERT: nycnn-ipvx21e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402601576] SERVICE ALERT: nycnn-ipvx20;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402601576] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601576] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 35.62 ms
[1402601576] SERVICE ALERT: nycnn-ipvx21;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601596] SERVICE ALERT: nycnn-ipvx20;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402601616] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402601616] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402601626] HOST ALERT: nycnn-ipvx21;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402601676] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402601676] SERVICE ALERT: nycnn-ipvx22e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402601716] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job 10211806 remains in a queued state
[1402601746] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402601746] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402601746] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601746] SERVICE ALERT: nycnn-ipvx22;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601746] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601746] SERVICE ALERT: nycnn-ipvx22;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601746] SERVICE ALERT: nycnn-ipvx22e1;Service: Apache;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601746] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601766] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;SOFT;2;AME is idle:
[1402601786] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1402601786] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402601796] SERVICE ALERT: nycnn-ipvx21;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402601796] SERVICE ALERT: nycnn-ipvx21;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.824GB (41%) - Free: 1.166GB (59%)
[1402601796] SERVICE ALERT: nycnn-ipvx21;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402601806] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402601806] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402601806] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402601806] SERVICE ALERT: nycnn-ipvx21e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.822GB (41%) - Free: 1.168GB (59%)
[1402601806] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601806] SERVICE ALERT: nycnn-ipvx22;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601806] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601806] SERVICE ALERT: nycnn-ipvx22;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402601806] SERVICE ALERT: nycnn-ipvx22e1;Service: Apache;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1402601816] SERVICE ALERT: nycnn-ipvx21;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402601826] SERVICE ALERT: nycnn-ipvx21e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402601836] SERVICE ALERT: nycnn-ipvx22;CPU Load;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402601856] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402601856] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1402601856] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402601856] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402601866] SERVICE ALERT: nycnn-ipvx21;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402601876] SERVICE ALERT: nycnn-ipvx21e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402602036] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602086] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602096] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.814GB (41%) - Free: 1.176GB (59%)
[1402602096] SERVICE ALERT: nycnn-ipvx22;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602096] SERVICE ALERT: nycnn-ipvx22;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402602096] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402602096] SERVICE ALERT: nycnn-ipvx22e1;Service: Apache;OK;HARD;3;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402602106] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402602106] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402602106] SERVICE ALERT: nycnn-ipvx24e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402602106] SERVICE ALERT: nycnn-ipvx23;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602106] SERVICE ALERT: nycnn-ipvx23;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602126] SERVICE ALERT: nycnn-ipvx23;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602126] SERVICE ALERT: nycnn-ipvx23e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602136] SERVICE ALERT: nycnn-ipvx22;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402602136] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1402602136] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402602156] SERVICE ALERT: nycnn-ipvx23;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602156] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602161] HOST DOWNTIME ALERT: nycnn-ipvx35e1;STARTED; Host has entered a period of scheduled downtime
[1402602166] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402602166] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402602166] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402602166] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402602166] SERVICE ALERT: nycnn-ipvx23;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602166] SERVICE ALERT: nycnn-ipvx23;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602176] SERVICE ALERT: nycnn-ipvx23e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602176] SERVICE ALERT: nycnn-ipvx24e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602186] SERVICE ALERT: nycnn-ipvx24e1;Memory Usage;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602186] SERVICE ALERT: nycnn-ipvx23;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602186] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602186] SERVICE ALERT: nycnn-ipvx23e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602186] SERVICE ALERT: nycnn-ipvx23e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602186] SERVICE ALERT: nycnn-ipvx24;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602196] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402602196] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1402602206] SERVICE ALERT: nycnn-ipvx23e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602206] SERVICE ALERT: nycnn-ipvx24;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602216] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602226] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402602226] SERVICE ALERT: nycnn-ipvx23;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.48 ms
[1402602226] SERVICE ALERT: nycnn-ipvx23e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.42 ms
[1402602226] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.27 ms
[1402602226] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.32 ms
[1402602226] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402602236] SERVICE ALERT: nycnn-ipvx24e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602246] SERVICE ALERT: nycnn-ipvx24e1;Memory Usage;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1402602246] SERVICE ALERT: nycnn-ipvx24;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602266] SERVICE ALERT: nycnn-ipvx24;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602276] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 20.91 ms
[1402602276] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.23 ms
[1402602276] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402602276] SERVICE ALERT: nycnn-ipvx24e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402602276] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1402602276] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402602276] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402602276] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402602308] HOST DOWNTIME ALERT: nycnn-ipvx35;STARTED; Host has entered a period of scheduled downtime
[1402602336] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602336] SERVICE ALERT: nycnn-ipvx24e1;Service: Apache;OK;SOFT;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402602336] HOST DOWNTIME ALERT: nycnn-ipvx34e1;STARTED; Host has entered a period of scheduled downtime
[1402602358] HOST DOWNTIME ALERT: nycnn-ipvx34;STARTED; Host has entered a period of scheduled downtime
[1402602379] HOST DOWNTIME ALERT: nycnn-ipvx33e1;STARTED; Host has entered a period of scheduled downtime
[1402602395] HOST DOWNTIME ALERT: nycnn-ipvx33;STARTED; Host has entered a period of scheduled downtime
[1402602396] HOST DOWNTIME ALERT: nycnn-ipvx9e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602436] SERVICE ALERT: nycnn-ipvx23;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402602456] SERVICE ALERT: nycnn-ipvx23;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402602456] SERVICE ALERT: nycnn-ipvx23;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602466] SERVICE ALERT: nycnn-ipvx23e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602466] SERVICE ALERT: nycnn-ipvx23;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.818GB (41%) - Free: 1.172GB (59%)
[1402602466] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602472] HOST DOWNTIME ALERT: nycnn-ipvx9;STOPPED; Host has exited from a period of scheduled downtime
[1402602476] SERVICE ALERT: nycnn-ipvx23e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402602476] SERVICE ALERT: nycnn-ipvx23e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402602495] HOST DOWNTIME ALERT: nycnn-ipvx1;STOPPED; Host has exited from a period of scheduled downtime
[1402602496] SERVICE ALERT: nycnn-ipvx23e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.820GB (41%) - Free: 1.170GB (59%)
[1402602496] SERVICE ALERT: nycnn-ipvx24;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402602496] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402602520] HOST DOWNTIME ALERT: nycnn-ipvx10;STOPPED; Host has exited from a period of scheduled downtime
[1402602526] SERVICE ALERT: nycnn-ipvx24e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602526] SERVICE ALERT: nycnn-ipvx24e1;Memory Usage;OK;HARD;3;OK - Physical Memory: Total: 1.990GB - Used: 0.827GB (42%) - Free: 1.163GB (58%)
[1402602536] SERVICE ALERT: nycnn-ipvx24;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402602537] HOST DOWNTIME ALERT: nycnn-ipvx10e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602546] SERVICE ALERT: nycnn-ipvx24;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402602553] HOST DOWNTIME ALERT: nycnn-ipvx11;STOPPED; Host has exited from a period of scheduled downtime
[1402602571] HOST DOWNTIME ALERT: nycnn-ipvx11e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602592] HOST DOWNTIME ALERT: nycnn-ipvx12;STOPPED; Host has exited from a period of scheduled downtime
[1402602609] HOST DOWNTIME ALERT: nycnn-ipvx12e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602626] HOST DOWNTIME ALERT: nycnn-ipvx13;STOPPED; Host has exited from a period of scheduled downtime
[1402602640] HOST DOWNTIME ALERT: nycnn-ipvx13e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602653] HOST DOWNTIME ALERT: nycnn-ipvx14;STOPPED; Host has exited from a period of scheduled downtime
[1402602680] HOST DOWNTIME ALERT: nycnn-ipvx14e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602696] HOST DOWNTIME ALERT: nycnn-ipvx15;STOPPED; Host has exited from a period of scheduled downtime
[1402602731] HOST DOWNTIME ALERT: nycnn-ipvx15e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602746] HOST DOWNTIME ALERT: nycnn-ipvx16;STOPPED; Host has exited from a period of scheduled downtime
[1402602778] HOST DOWNTIME ALERT: nycnn-ipvx16e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602794] HOST DOWNTIME ALERT: nycnn-ipvx17;STOPPED; Host has exited from a period of scheduled downtime
[1402602806] SERVICE ALERT: nycnn-ipvx34e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402602816] SERVICE ALERT: nycnn-ipvx34e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402602826] HOST DOWNTIME ALERT: nycnn-ipvx17e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602826] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.23 ms
[1402602826] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.20 ms
[1402602841] HOST DOWNTIME ALERT: nycnn-ipvx18;STOPPED; Host has exited from a period of scheduled downtime
[1402602856] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602856] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1402602856] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402602860] HOST DOWNTIME ALERT: nycnn-ipvx18e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602866] SERVICE ALERT: nycnn-ipvx34;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602866] SERVICE ALERT: nycnn-ipvx34;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602875] HOST DOWNTIME ALERT: nycnn-ipvx19;STOPPED; Host has exited from a period of scheduled downtime
[1402602876] SERVICE ALERT: nycnn-ipvx34e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602886] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402602886] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402602886] SERVICE ALERT: nycnn-ipvx34e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602891] HOST DOWNTIME ALERT: nycnn-ipvx19e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602910] HOST DOWNTIME ALERT: nycnn-ipvx1e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602916] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402602943] HOST DOWNTIME ALERT: nycnn-ipvx2;STOPPED; Host has exited from a period of scheduled downtime
[1402602946] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402602946] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;2;(Host Check Timed Out)
[1402602946] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402602946] SERVICE ALERT: nycnn-ipvx34e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402602946] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402602956] SERVICE ALERT: nycnn-ipvx34;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402602956] SERVICE ALERT: nycnn-ipvx34e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402602957] HOST DOWNTIME ALERT: nycnn-ipvx20;STOPPED; Host has exited from a period of scheduled downtime
[1402602973] HOST DOWNTIME ALERT: nycnn-ipvx20e1;STOPPED; Host has exited from a period of scheduled downtime
[1402602988] HOST DOWNTIME ALERT: nycnn-ipvx21;STOPPED; Host has exited from a period of scheduled downtime
[1402602996] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402602996] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402603006] HOST ALERT: nycnn-ipvx34;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402603006] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402603006] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402603006] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402603008] HOST DOWNTIME ALERT: nycnn-ipvx21e1;STOPPED; Host has exited from a period of scheduled downtime
[1402603025] HOST DOWNTIME ALERT: nycnn-ipvx22;STOPPED; Host has exited from a period of scheduled downtime
[1402603026] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402603026] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402603026] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402603026] SERVICE ALERT: nycnn-ipvx33;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603036] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1402603056] SERVICE ALERT: nycnn-ipvx33e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603066] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402603066] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402603086] SERVICE ALERT: nycnn-ipvx33;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603086] SERVICE ALERT: nycnn-ipvx33;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402603096] SERVICE ALERT: nycnn-ipvx33;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603116] SERVICE ALERT: nycnn-ipvx33;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603116] SERVICE ALERT: nycnn-ipvx33e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603116] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.44 ms
[1402603116] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 69.76 ms
[1402603116] SERVICE ALERT: nycnn-ipvx33;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402603116] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402603116] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1402603156] SERVICE ALERT: nycnn-ipvx34;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402603156] SERVICE ALERT: nycnn-ipvx34;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402603166] SERVICE ALERT: nycnn-ipvx34e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402603176] SERVICE ALERT: nycnn-ipvx33;CPU Load;OK;SOFT;2;OK (Sample Period 360 sec) - Average CPU Utilisation 0.00%
[1402603176] SERVICE ALERT: nycnn-ipvx34e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402603196] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402603236] SERVICE ALERT: nycnn-ipvx34e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402603256] SERVICE ALERT: nycnn-ipvx34e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402603256] SERVICE ALERT: nycnn-ipvx34;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402603273] HOST DOWNTIME ALERT: nycnn-ipvx29;STARTED; Host has entered a period of scheduled downtime
[1402603289] HOST DOWNTIME ALERT: nycnn-ipvx29e1;STARTED; Host has entered a period of scheduled downtime
[1402603305] HOST DOWNTIME ALERT: nycnn-ipvx30;STARTED; Host has entered a period of scheduled downtime
[1402603306] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402603306] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402603327] HOST DOWNTIME ALERT: nycnn-ipvx30e1;STARTED; Host has entered a period of scheduled downtime
[1402603341] HOST DOWNTIME ALERT: nycnn-ipvx31;STARTED; Host has entered a period of scheduled downtime
[1402603346] SERVICE ALERT: nycnn-ipvx33e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402603356] HOST DOWNTIME ALERT: nycnn-ipvx31e1;STARTED; Host has entered a period of scheduled downtime
[1402603366] SERVICE ALERT: nycnn-ipvx33;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402603376] SERVICE ALERT: nycnn-ipvx33;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.830GB (42%) - Free: 1.160GB (58%)
[1402603376] SERVICE ALERT: nycnn-ipvx33;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402603377] HOST DOWNTIME ALERT: nycnn-ipvx32e1;STARTED; Host has entered a period of scheduled downtime
[1402603396] SERVICE ALERT: nycnn-ipvx33;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402603396] SERVICE ALERT: nycnn-ipvx33e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.829GB (42%) - Free: 1.161GB (58%)
[1402603399] HOST DOWNTIME ALERT: nycnn-ipvx32;STARTED; Host has entered a period of scheduled downtime
[1402603556] SERVICE ALERT: nycnn-ipvx25;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402603586] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1402603606] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402603606] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402603626] SERVICE ALERT: nycnn-ipvx25;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402603636] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402603656] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603666] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402603676] SERVICE ALERT: nycnn-ipvx25;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603686] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;2;(Host Check Timed Out)
[1402603691] HOST DOWNTIME ALERT: nycnn-ipvx3;STOPPED; Host has exited from a period of scheduled downtime
[1402603696] SERVICE ALERT: nycnn-ipvx25;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402603696] SERVICE ALERT: nycnn-ipvx25;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402603706] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1402603712] HOST DOWNTIME ALERT: nycnn-ipvx4;STOPPED; Host has exited from a period of scheduled downtime
[1402603716] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402603716] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402603726] HOST ALERT: nycnn-ipvx25;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402603728] HOST DOWNTIME ALERT: nycnn-ipvx5;STOPPED; Host has exited from a period of scheduled downtime
[1402603749] HOST DOWNTIME ALERT: nycnn-ipvx6;STOPPED; Host has exited from a period of scheduled downtime
[1402603771] HOST DOWNTIME ALERT: nycnn-ipvx5e1;STOPPED; Host has exited from a period of scheduled downtime
[1402603788] HOST DOWNTIME ALERT: nycnn-ipvx6e1;STOPPED; Host has exited from a period of scheduled downtime
[1402603916] SERVICE ALERT: nycnn-ipvx25;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402603946] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402603966] SERVICE ALERT: nycnn-ipvx25;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402603976] SERVICE ALERT: nycnn-ipvx25;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402603996] SERVICE ALERT: nycnn-ipvx25;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.831GB (42%) - Free: 1.159GB (58%)
[1402604008] HOST DOWNTIME ALERT: nycnn-ipvx8;STOPPED; Host has exited from a period of scheduled downtime
[1402604036] SERVICE ALERT: nycnn-ipvx28;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402604041] HOST DOWNTIME ALERT: nycnn-ipvx8e1;STOPPED; Host has exited from a period of scheduled downtime
[1402604064] HOST DOWNTIME ALERT: nycnn-ipvx4e1;STOPPED; Host has exited from a period of scheduled downtime
[1402604066] SERVICE ALERT: nycnn-ipvx27e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604066] SERVICE ALERT: nycnn-ipvx28e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604076] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402604076] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604083] HOST DOWNTIME ALERT: nycnn-ipvx3e1;STOPPED; Host has exited from a period of scheduled downtime
[1402604086] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402604086] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402604086] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402604086] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402604096] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402604106] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1402604116] SERVICE ALERT: nycnn-ipvx28;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604126] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1402604126] SERVICE ALERT: nycnn-ipvx27e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604126] SERVICE ALERT: nycnn-ipvx28e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604136] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604146] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402604146] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402604146] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402604146] SERVICE ALERT: nycnn-ipvx28e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604146] SERVICE ALERT: nycnn-ipvx27;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604146] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604146] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604146] SERVICE ALERT: nycnn-ipvx28e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604166] SERVICE ALERT: nycnn-ipvx27;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402604166] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402604166] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.48 ms
[1402604166] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.46 ms
[1402604176] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402604176] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;2;(Host Check Timed Out)
[1402604196] SERVICE ALERT: nycnn-ipvx28;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402604196] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402604196] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402604196] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1402604196] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1402604206] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402604206] HOST ALERT: nycnn-ipvx28;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1402604396] SERVICE ALERT: nycnn-ipvx28;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402604406] SERVICE ALERT: nycnn-ipvx27e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402604406] SERVICE ALERT: nycnn-ipvx28e1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402604426] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604436] SERVICE ALERT: nycnn-ipvx28e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402604436] SERVICE ALERT: nycnn-ipvx27;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.820GB (41%) - Free: 1.170GB (59%)
[1402604436] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604436] SERVICE ALERT: nycnn-ipvx28e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402604436] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604446] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402604446] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402604446] SERVICE ALERT: nycnn-ipvx29e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604446] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604456] SERVICE ALERT: nycnn-ipvx29;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604456] SERVICE ALERT: nycnn-ipvx27;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402604466] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402604466] SERVICE ALERT: nycnn-ipvx29;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604476] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402604486] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1402604486] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604486] SERVICE ALERT: nycnn-ipvx28;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604496] SERVICE ALERT: nycnn-ipvx30;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1402604506] SERVICE ALERT: nycnn-ipvx29e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604506] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402604506] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402604506] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.27 ms
[1402604506] SERVICE ALERT: nycnn-ipvx29e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604506] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604506] SERVICE ALERT: nycnn-ipvx29e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604506] SERVICE ALERT: nycnn-ipvx29;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604516] SERVICE ALERT: nycnn-ipvx29;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604526] SERVICE ALERT: nycnn-ipvx29;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604526] SERVICE ALERT: nycnn-ipvx29;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604536] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1402604536] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402604536] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;2;(Host Check Timed Out)
[1402604556] SERVICE ALERT: nycnn-ipvx29e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402604556] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.52 ms
[1402604556] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1402604556] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1402604566] HOST ALERT: nycnn-ipvx29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402604566] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1402604566] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402604566] SERVICE ALERT: nycnn-ipvx30;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402604576] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604596] SERVICE ALERT: nycnn-ipvx30;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604616] SERVICE ALERT: nycnn-ipvx30e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604636] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;2;(Host Check Timed Out)
[1402604636] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402604636] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1402604666] SERVICE ALERT: nycnn-ipvx30;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402604676] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1402604676] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402604686] HOST ALERT: nycnn-ipvx30;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402604686] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1402604766] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402604786] SERVICE ALERT: nycnn-ipvx29e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604796] SERVICE ALERT: nycnn-ipvx29e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.815GB (41%) - Free: 1.175GB (59%)
[1402604796] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604796] SERVICE ALERT: nycnn-ipvx29e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604796] SERVICE ALERT: nycnn-ipvx29;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402604796] SERVICE ALERT: nycnn-ipvx29;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1402604806] SERVICE ALERT: nycnn-ipvx29;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604816] SERVICE ALERT: nycnn-ipvx29;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402604846] SERVICE ALERT: nycnn-ipvx29e1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402604856] SERVICE ALERT: nycnn-ipvx30;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604866] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604876] SERVICE ALERT: nycnn-ipvx30;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402604896] SERVICE ALERT: nycnn-ipvx30e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.828GB (42%) - Free: 1.162GB (58%)
[1402604936] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402604956] SERVICE ALERT: nycnn-ipvx30;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402605046] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402605046] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402605046] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402605046] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1402605046] SERVICE ALERT: nycnn-ipvx31e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605046] SERVICE ALERT: nycnn-ipvx32e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605046] SERVICE ALERT: nycnn-ipvx32e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605046] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605066] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605076] SERVICE ALERT: nycnn-ipvx32;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605076] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402605086] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1402605086] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1402605086] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1402605106] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402605106] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402605106] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402605106] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1402605106] SERVICE ALERT: nycnn-ipvx31e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402605106] SERVICE ALERT: nycnn-ipvx32e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402605106] SERVICE ALERT: nycnn-ipvx32e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402605106] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402605106] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;2;(Host Check Timed Out)
[1402605116] SERVICE ALERT: nycnn-ipvx32e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605116] SERVICE ALERT: nycnn-ipvx31;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605126] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402605126] SERVICE ALERT: nycnn-ipvx32e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605126] SERVICE ALERT: nycnn-ipvx31e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605136] SERVICE ALERT: nycnn-ipvx32e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605136] SERVICE ALERT: nycnn-ipvx32;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1402605156] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;2;(Host Check Timed Out)
[1402605156] HOST ALERT: nycnn-ipvx31;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402605156] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1402605156] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402605156] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402605156] SERVICE ALERT: nycnn-ipvx32;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1402605176] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1402605206] HOST ALERT: nycnn-ipvx32;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402605216] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1402605216] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1402605396] SERVICE ALERT: nycnn-ipvx31e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402605396] SERVICE ALERT: nycnn-ipvx32e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.820GB (41%) - Free: 1.170GB (59%)
[1402605396] SERVICE ALERT: nycnn-ipvx32e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402605396] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402605396] SERVICE ALERT: nycnn-ipvx32e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402605396] SERVICE ALERT: nycnn-ipvx31;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1402605406] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402605406] SERVICE ALERT: nycnn-ipvx32e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1402605406] SERVICE ALERT: nycnn-ipvx31e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.828GB (42%) - Free: 1.162GB (58%)
[1402605416] SERVICE ALERT: nycnn-ipvx32e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402605416] SERVICE ALERT: nycnn-ipvx32;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1402605446] SERVICE ALERT: nycnn-ipvx32;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1402606896] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;SOFT;1;Job 10211776 remains in a queued state
[1402606956] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;SOFT;2;Job 10211776 progress is  1.266 :
[1402607773] HOST DOWNTIME ALERT: nycnn-ipvx22e1;STOPPED; Host has exited from a period of scheduled downtime
[1402607791] HOST DOWNTIME ALERT: nycnn-ipvx23;STOPPED; Host has exited from a period of scheduled downtime
[1402607808] HOST DOWNTIME ALERT: nycnn-ipvx23e1;STOPPED; Host has exited from a period of scheduled downtime
[1402607839] HOST DOWNTIME ALERT: nycnn-ipvx24;STOPPED; Host has exited from a period of scheduled downtime
[1402607855] HOST DOWNTIME ALERT: nycnn-ipvx24e1;STOPPED; Host has exited from a period of scheduled downtime
[1402607868] HOST DOWNTIME ALERT: nycnn-ipvx25;STOPPED; Host has exited from a period of scheduled downtime
[1402607887] HOST DOWNTIME ALERT: nycnn-ipvx25e1;STOPPED; Host has exited from a period of scheduled downtime
[1402608091] HOST DOWNTIME ALERT: nycnn-ipvx26;STOPPED; Host has exited from a period of scheduled downtime
[1402608125] HOST DOWNTIME ALERT: nycnn-ipvx26e1;STOPPED; Host has exited from a period of scheduled downtime
[1402608139] HOST DOWNTIME ALERT: nycnn-ipvx27;STOPPED; Host has exited from a period of scheduled downtime
[1402608158] HOST DOWNTIME ALERT: nycnn-ipvx27e1;STOPPED; Host has exited from a period of scheduled downtime
[1402608181] HOST DOWNTIME ALERT: nycnn-ipvx28;STOPPED; Host has exited from a period of scheduled downtime
[1402608204] HOST DOWNTIME ALERT: nycnn-ipvx28e1;STOPPED; Host has exited from a period of scheduled downtime
[1402609316] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;SOFT;1;Job 10212570 remains in a queued state
[1402609355] HOST DOWNTIME ALERT: nycnn-ipvx35e1;STOPPED; Host has exited from a period of scheduled downtime
[1402609376] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;HARD;2;Job 10212570 remains in a queued state
[1402609501] HOST DOWNTIME ALERT: nycnn-ipvx35;STOPPED; Host has exited from a period of scheduled downtime
[1402609531] HOST DOWNTIME ALERT: nycnn-ipvx34e1;STOPPED; Host has exited from a period of scheduled downtime
[1402609552] HOST DOWNTIME ALERT: nycnn-ipvx34;STOPPED; Host has exited from a period of scheduled downtime
[1402609573] HOST DOWNTIME ALERT: nycnn-ipvx33e1;STOPPED; Host has exited from a period of scheduled downtime
[1402609590] HOST DOWNTIME ALERT: nycnn-ipvx33;STOPPED; Host has exited from a period of scheduled downtime
[1402609676] SERVICE ALERT: nycnn-adbeame4;AME Job status;OK;HARD;2;Job 10212570 progress is  3.7 :
[1402610469] HOST DOWNTIME ALERT: nycnn-ipvx29;STOPPED; Host has exited from a period of scheduled downtime
[1402610485] HOST DOWNTIME ALERT: nycnn-ipvx29e1;STOPPED; Host has exited from a period of scheduled downtime
[1402610501] HOST DOWNTIME ALERT: nycnn-ipvx30;STOPPED; Host has exited from a period of scheduled downtime
[1402610523] HOST DOWNTIME ALERT: nycnn-ipvx30e1;STOPPED; Host has exited from a period of scheduled downtime
[1402610537] HOST DOWNTIME ALERT: nycnn-ipvx31;STOPPED; Host has exited from a period of scheduled downtime
[1402610552] HOST DOWNTIME ALERT: nycnn-ipvx31e1;STOPPED; Host has exited from a period of scheduled downtime
[1402610573] HOST DOWNTIME ALERT: nycnn-ipvx32e1;STOPPED; Host has exited from a period of scheduled downtime
[1402610594] HOST DOWNTIME ALERT: nycnn-ipvx32;STOPPED; Host has exited from a period of scheduled downtime
[1402627656] SERVICE ALERT: nycnn-ipvx26;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402627716] SERVICE ALERT: nycnn-ipvx26;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1402627746] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402627806] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.01%