[1410168008] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1410168008] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1410168058] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;SOFT;1;(Service Check Timed Out)
[1410168058] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;SOFT;1;(Service Check Timed Out)
[1410168058] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;SOFT;1;(Service Check Timed Out)
[1410168068] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1410168068] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1410168128] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1410168137] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1410168173] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;SOFT;2;(Service Check Timed Out)
[1410168173] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;SOFT;2;(Service Check Timed Out)
[1410168173] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;SOFT;2;(Service Check Timed Out)
[1410168293] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;HARD;3;(Service Check Timed Out)
[1410168307] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;HARD;3;(Service Check Timed Out)
[1410168319] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;HARD;3;(Service Check Timed Out)
[1410168771] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;3;OK: all MS2_EVENT queues have depth 1 or less, all MS2_EVENT queues in RUNNING state
[1410168781] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;OK;HARD;3;DISK OK - free space: /oracle 900191 MB (92% inode=99%):
[1410168791] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1410168793] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;OK;HARD;3;DISK OK - free space: /oracle 862830 MB (88% inode=99%):
[1410168799] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;OK;HARD;3;DISK OK - free space: /oracle 777717 MB (79% inode=99%):
[1410192587] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1410192587] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1410192677] HOST ALERT: omnmgsm-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1410192707] SERVICE ALERT: omnmgsm-ny;PING;WARNING;HARD;1;PING WARNING - Packet loss = 37%, RTA = 0.36 ms
[1410192747] HOST ALERT: omnmgsm-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1410192767] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1410199717] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;SOFT;1;Job 10425453 remains in a queued state
[1410199777] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;SOFT;2;AME is idle:
[1410200037] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;SOFT;1;Job 10425447 remains in a queued state
[1410200087] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;SOFT;2;AME is idle:
[1410202934] HOST DOWNTIME ALERT: nycnn-ipvsphn1;STARTED; Host has entered a period of scheduled downtime
[1410202949] HOST DOWNTIME ALERT: nycnn-ipvsphn2;STARTED; Host has entered a period of scheduled downtime
[1410202959] HOST DOWNTIME ALERT: nycnn-ipvsphn3;STARTED; Host has entered a period of scheduled downtime
[1410202970] HOST DOWNTIME ALERT: nycnn-ipvsphn4;STARTED; Host has entered a period of scheduled downtime
[1410202988] HOST DOWNTIME ALERT: nycnn-ipvsphn5;STARTED; Host has entered a period of scheduled downtime
[1410202999] HOST DOWNTIME ALERT: nycnn-ipvstrm1;STARTED; Host has entered a period of scheduled downtime
[1410203010] HOST DOWNTIME ALERT: nycnn-ipvstrm2;STARTED; Host has entered a period of scheduled downtime
[1410203021] HOST DOWNTIME ALERT: nycnn-ipvstrm3;STARTED; Host has entered a period of scheduled downtime
[1410203033] HOST DOWNTIME ALERT: nycnn-ipvstrm4;STARTED; Host has entered a period of scheduled downtime
[1410203517] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1410203567] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1410203597] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1410203597] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1410203627] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1410203627] SERVICE ALERT: nycnn-ipvsphn1;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410203657] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1410203697] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;2;(Host Check Timed Out)
[1410203697] SERVICE ALERT: nycnn-ipvsphn1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1410203697] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1410203737] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1410203747] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1410203827] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1410203837] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1410203837] SERVICE ALERT: nycnn-ipvsphn3;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1410203867] HOST ALERT: nycnn-ipvsphn5;DOWN;SOFT;1;(Host Check Timed Out)
[1410203867] SERVICE ALERT: nycnn-ipvsphn3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1410203877] HOST ALERT: nycnn-ipvsphn3;DOWN;SOFT;1;(Host Check Timed Out)
[1410203887] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1410203897] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410203897] SERVICE ALERT: nycnn-ipvsphn3;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410203917] SERVICE ALERT: nycnn-ipvsphn1;C:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=1.62GB (1.2%), Free=135.07GB (98.8%)
[1410203927] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 70%, RTA = 0.26 ms
[1410203927] SERVICE ALERT: nycnn-ipvsphn3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1410203937] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;WARNING;HARD;3;WARNING (Sample Period 720 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410203947] SERVICE ALERT: nycnn-ipvsphn5;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1410203947] HOST ALERT: nycnn-ipvsphn5;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.37 ms
[1410203947] SERVICE ALERT: nycnn-ipvsphn5;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 54%, RTA = 0.25 ms
[1410203967] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1410203977] HOST ALERT: nycnn-ipvsphn3;DOWN;SOFT;2;(Host Check Timed Out)
[1410203987] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1410203987] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410203987] SERVICE ALERT: nycnn-ipvsphn1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1410203987] SERVICE ALERT: nycnn-ipvsphn5;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410203997] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1410204007] SERVICE ALERT: nycnn-ipvsphn5;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1410204007] SERVICE ALERT: nycnn-ipvstrm1;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204007] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204007] SERVICE ALERT: nycnn-ipvstrm1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204047] SERVICE ALERT: nycnn-ipvsphn5;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.64%
[1410204057] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1410204057] SERVICE ALERT: nycnn-ipvsphn3;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204057] SERVICE ALERT: nycnn-ipvsphn3;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204057] SERVICE ALERT: nycnn-ipvsphn3;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204057] SERVICE ALERT: nycnn-ipvsphn3;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204057] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204057] SERVICE ALERT: nycnn-ipvstrm1;Process: SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204057] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204067] HOST ALERT: nycnn-ipvsphn3;UP;SOFT;3;PING WARNING - Packet loss = 80%, RTA = 0.19 ms
[1410204097] SERVICE ALERT: nycnn-ipvsphn3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1410204127] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1410204137] SERVICE ALERT: nycnn-ipvsphn4;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204137] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204157] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1410204167] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1410204167] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;1;(Host Check Timed Out)
[1410204177] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410204187] SERVICE ALERT: nycnn-ipvsphn3;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1410204197] SERVICE ALERT: nycnn-ipvsphn4;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410204197] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410204227] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1410204237] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 3.26%
[1410204237] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;OK;HARD;3;OK (Sample Period 300 sec) - Average CPU Utilisation 0.20%
[1410204247] SERVICE ALERT: nycnn-ipvsphn5;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1410204267] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;2;(Host Check Timed Out)
[1410204277] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1410204277] SERVICE ALERT: nycnn-ipvsphn4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1410204287] HOST ALERT: nycnn-ipvsphn4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1410204287] SERVICE ALERT: nycnn-ipvstrm1;C:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1121.75GB (80.3%), Free=275.44GB (19.7%)
[1410204287] SERVICE ALERT: nycnn-ipvstrm1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.705GB (18%) - Free: 3.295GB (82%)
[1410204287] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1410204297] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204337] SERVICE ALERT: nycnn-ipvstrm2;Process: SVServer;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1410204337] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 4.000GB - Used: 0.475GB (12%) - Free: 3.524GB (88%)
[1410204337] SERVICE ALERT: nycnn-ipvstrm2;Process: IPS_FTP;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1410204337] SERVICE ALERT: nycnn-ipvsphn3;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1410204337] SERVICE ALERT: nycnn-ipvsphn3;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.578GB (14%) - Free: 3.422GB (86%)
[1410204337] SERVICE ALERT: nycnn-ipvsphn3;D:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.62GB (100.0%)
[1410204337] SERVICE ALERT: nycnn-ipvsphn3;C:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.62GB (100.0%)
[1410204337] SERVICE ALERT: nycnn-ipvstrm1;Process: SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1410204337] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1121.75GB (80.3%), Free=275.44GB (19.7%)
[1410204337] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1410204377] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1410204377] SERVICE ALERT: nycnn-ipvstrm2;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204407] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1410204417] SERVICE ALERT: nycnn-ipvstrm2;Process: SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410204417] SERVICE ALERT: nycnn-ipvstrm2;Process: IPS_FTP;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410204427] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204437] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204437] SERVICE ALERT: nycnn-ipvstrm2;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204437] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204437] SERVICE ALERT: nycnn-ipvstrm2;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204447] SERVICE ALERT: nycnn-ipvstrm2;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410204477] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1410204477] SERVICE ALERT: nycnn-ipvsphn4;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1410204477] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410204537] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 2.39%
[1410204557] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 54%, RTA = 0.38 ms
[1410204577] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1410204657] SERVICE ALERT: nycnn-ipvstrm2;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1410204697] SERVICE ALERT: nycnn-ipvstrm2;Process: SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1410204697] SERVICE ALERT: nycnn-ipvstrm2;Process: IPS_FTP;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1410204717] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1410204727] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1410204727] SERVICE ALERT: nycnn-ipvstrm2;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1410204727] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1116.52GB (79.9%), Free=280.67GB (20.1%)
[1410204727] SERVICE ALERT: nycnn-ipvstrm2;C:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1116.52GB (79.9%), Free=280.67GB (20.1%)
[1410204737] SERVICE ALERT: nycnn-ipvstrm2;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.706GB (18%) - Free: 3.293GB (82%)
[1410205187] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1410205207] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205207] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205217] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1410205247] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1410205257] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205267] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410205267] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1410205307] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1410205337] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205337] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205337] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205347] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205367] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1410205387] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING OK - Packet loss = 58%, RTA = 0.27 ms
[1410205417] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1410205537] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1410205547] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;2;OK - Proxy Total=1397.19GB, Used=1118.91GB (80.1%), Free=278.28GB (19.9%)
[1410205547] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;OK;HARD;2;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1410205627] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1410205627] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1410205627] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1410205627] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1118.91GB (80.1%), Free=278.28GB (19.9%)
[1410205647] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.797GB (20%) - Free: 3.203GB (80%)
[1410210115] HOST DOWNTIME ALERT: nycnn-ipvsphn1;STOPPED; Host has exited from a period of scheduled downtime
[1410210145] HOST DOWNTIME ALERT: nycnn-ipvsphn2;STOPPED; Host has exited from a period of scheduled downtime
[1410210156] HOST DOWNTIME ALERT: nycnn-ipvsphn3;STOPPED; Host has exited from a period of scheduled downtime
[1410210167] HOST DOWNTIME ALERT: nycnn-ipvsphn4;STOPPED; Host has exited from a period of scheduled downtime
[1410210184] HOST DOWNTIME ALERT: nycnn-ipvsphn5;STOPPED; Host has exited from a period of scheduled downtime
[1410210195] HOST DOWNTIME ALERT: nycnn-ipvstrm1;STOPPED; Host has exited from a period of scheduled downtime
[1410210206] HOST DOWNTIME ALERT: nycnn-ipvstrm2;STOPPED; Host has exited from a period of scheduled downtime
[1410210217] HOST DOWNTIME ALERT: nycnn-ipvstrm3;STOPPED; Host has exited from a period of scheduled downtime
[1410210229] HOST DOWNTIME ALERT: nycnn-ipvstrm4;STOPPED; Host has exited from a period of scheduled downtime
[1410219897] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 10429297 remains in a queued state
[1410219957] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;AME is idle:
[1410227867] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410227917] SERVICE ALERT: nycnn-ipvx23;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410227927] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1410227977] SERVICE ALERT: nycnn-ipvx23;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.00%
[1410233737] SERVICE ALERT: nycnn-ipvx30;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410233797] SERVICE ALERT: nycnn-ipvx30;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1410233987] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1410234047] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.03%