[1375593448] SERVICE ALERT: rt-nycchdagg-a;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 108.17 ms
[1375593508] SERVICE ALERT: rt-nycchdagg-a;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.26 ms
[1375596338] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.20 ms
[1375596338] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.31 ms
[1375596338] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375596338] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.30 ms
[1375596338] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375596368] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;1;(Host Check Timed Out)
[1375596368] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;1;(Host Check Timed Out)
[1375596368] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;1;(Host Check Timed Out)
[1375596368] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;1;(Host Check Timed Out)
[1375596368] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;1;(Host Check Timed Out)
[1375596398] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596398] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375596398] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596398] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596398] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375596398] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375596398] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596398] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596428] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;1;(Host Check Timed Out)
[1375596428] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;1;(Host Check Timed Out)
[1375596428] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;1;(Host Check Timed Out)
[1375596458] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1375596458] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1375596458] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1375596458] SERVICE ALERT: nycnn-adbeame7;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1375596458] SERVICE ALERT: nycnn-adbeame9;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1375596458] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596458] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596458] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375596468] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;2;(Host Check Timed Out)
[1375596468] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;2;(Host Check Timed Out)
[1375596468] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;2;(Host Check Timed Out)
[1375596468] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;2;(Host Check Timed Out)
[1375596468] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;2;(Host Check Timed Out)
[1375596508] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1375596518] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1375596518] HOST ALERT: nycnn-adbeame4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1375596528] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;2;(Host Check Timed Out)
[1375596528] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;2;(Host Check Timed Out)
[1375596528] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;2;(Host Check Timed Out)
[1375596538] HOST ALERT: nycnn-adbeame5;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1375596538] HOST ALERT: nycnn-adbeame9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1375596538] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;HARD;1;Server is down!
[1375596538] HOST ALERT: nycnn-adbeame1;UP;SOFT;3;PING OK - Packet loss = 28%, RTA = 1.32 ms
[1375596548] HOST ALERT: nycnn-adbeame7;UP;SOFT;3;PING OK - Packet loss = 54%, RTA = 0.51 ms
[1375596558] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;HARD;1;Server is down!
[1375596568] SERVICE ALERT: nycnn-adbeame1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1375596568] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1375596578] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1375596578] SERVICE ALERT: nycnn-adbeame6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1375596578] SERVICE ALERT: nycnn-adbeame8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1375596578] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1375596578] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1375596578] HOST ALERT: nycnn-adbeame2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1375596588] HOST ALERT: nycnn-adbeame6;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1375596588] HOST ALERT: nycnn-adbeame8;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1375596698] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;HARD;1;AME is idle:
[1375596698] SERVICE ALERT: nycnn-adbeame5;AME Job status;OK;HARD;1;AME is idle:
[1375596698] SERVICE ALERT: nycnn-adbeame4;AME Job status;OK;HARD;1;AME is idle:
[1375596698] SERVICE ALERT: nycnn-adbeame7;AME Job status;OK;HARD;1;AME is idle:
[1375596698] SERVICE ALERT: nycnn-adbeame9;AME Job status;OK;HARD;1;AME is idle:
[1375596758] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;HARD;1;AME is idle:
[1375596818] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;HARD;1;AME is idle:
[1375596818] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;HARD;1;AME is idle:
[1375598568] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 119.86 ms
[1375598628] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.93 ms
[1375621792] HOST DOWNTIME ALERT: omnclda01-ny;STARTED; Host has entered a period of scheduled downtime
[1375621792] HOST DOWNTIME ALERT: omnclda01e1-ny;STARTED; Host has entered a period of scheduled downtime
[1375621792] HOST DOWNTIME ALERT: omnclda02-ny;STARTED; Host has entered a period of scheduled downtime
[1375621792] HOST DOWNTIME ALERT: omnclda02e1-ny;STARTED; Host has entered a period of scheduled downtime
[1375621792] HOST DOWNTIME ALERT: omnclda03-ny;STARTED; Host has entered a period of scheduled downtime
[1375621792] HOST DOWNTIME ALERT: omnclda03e1-ny;STARTED; Host has entered a period of scheduled downtime
[1375621801] SERVICE DOWNTIME ALERT: omnclda01-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375621801] SERVICE DOWNTIME ALERT: omnclda01e1-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375621801] SERVICE DOWNTIME ALERT: omnclda02-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375621801] SERVICE DOWNTIME ALERT: omnclda02e1-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375621801] SERVICE DOWNTIME ALERT: omnclda03-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375621801] SERVICE DOWNTIME ALERT: omnclda03e1-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375624298] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1375624348] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;SOFT;2;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1375627898] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375627928] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375627958] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375628018] HOST ALERT: omnmgsm-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375628078] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.14 ms
[1375628088] HOST ALERT: omnmgsm-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1375628986] HOST DOWNTIME ALERT: omnclda01-ny;STOPPED; Host has exited from a period of scheduled downtime
[1375628986] HOST DOWNTIME ALERT: omnclda01e1-ny;STOPPED; Host has exited from a period of scheduled downtime
[1375628986] HOST DOWNTIME ALERT: omnclda02-ny;STOPPED; Host has exited from a period of scheduled downtime
[1375628986] HOST DOWNTIME ALERT: omnclda02e1-ny;STOPPED; Host has exited from a period of scheduled downtime
[1375628986] HOST DOWNTIME ALERT: omnclda03-ny;STOPPED; Host has exited from a period of scheduled downtime
[1375628986] HOST DOWNTIME ALERT: omnclda03e1-ny;STOPPED; Host has exited from a period of scheduled downtime
[1375628995] SERVICE DOWNTIME ALERT: omnclda01-ny;PING;STOPPED; Service has exited from a period of scheduled downtime
[1375628995] SERVICE DOWNTIME ALERT: omnclda01e1-ny;PING;STOPPED; Service has exited from a period of scheduled downtime
[1375628995] SERVICE DOWNTIME ALERT: omnclda02-ny;PING;STOPPED; Service has exited from a period of scheduled downtime
[1375628995] SERVICE DOWNTIME ALERT: omnclda02e1-ny;PING;STOPPED; Service has exited from a period of scheduled downtime
[1375628995] SERVICE DOWNTIME ALERT: omnclda03-ny;PING;STOPPED; Service has exited from a period of scheduled downtime
[1375628995] SERVICE DOWNTIME ALERT: omnclda03e1-ny;PING;STOPPED; Service has exited from a period of scheduled downtime
[1375630578] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375630608] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375630638] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375630708] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375630778] HOST ALERT: omnhbcba01-ny;UP;SOFT;3;PING OK - Packet loss = 44%, RTA = 0.23 ms
[1375630808] SERVICE ALERT: omnhbcba01-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1375634838] SERVICE ALERT: omnhbcba02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375634838] SERVICE ALERT: omnhbcba05-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375634878] HOST ALERT: omnhbcba02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375634878] HOST ALERT: omnhbcba05-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375634898] SERVICE ALERT: omnhbcba02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375634898] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375634898] SERVICE ALERT: omnhbcba05-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375634898] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375634938] HOST ALERT: omnhbcba07-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375634938] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375634958] HOST ALERT: omnhbcba02-ny;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.25 ms
[1375634958] SERVICE ALERT: omnhbcba02-ny;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.17 ms
[1375634958] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375634958] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375634968] HOST ALERT: omnhbcba05-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375635018] SERVICE ALERT: omnhbcba05-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1375635028] HOST ALERT: omnhbcba05-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1375635028] HOST ALERT: omnhbcba07-ny;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.22 ms
[1375635028] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375635078] SERVICE ALERT: omnhbcba07-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1375635078] SERVICE ALERT: omnhbcba09-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1375635088] HOST ALERT: omnhbcba09-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1375638498] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375638498] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375638508] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375638538] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375638558] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375638558] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375638558] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375638558] SERVICE ALERT: omnhbcba08-ny;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1375638598] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375638598] HOST ALERT: omnhbcba08-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375638598] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375638618] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375638618] SERVICE ALERT: omnhbcba08-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375638618] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375638638] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375638658] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375638668] SERVICE ALERT: omnhbcba03-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1375638678] SERVICE ALERT: omnhbcba10-ny;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.16 ms
[1375638678] HOST ALERT: omnhbcba03-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1375638678] HOST ALERT: omnhbcba10-ny;UP;SOFT;3;PING OK - Packet loss = 66%, RTA = 0.19 ms
[1375638678] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375638698] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375638698] HOST ALERT: omnhbcba08-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375638738] SERVICE ALERT: omnhbcba06-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1375638738] SERVICE ALERT: omnhbcba10-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1375638738] SERVICE ALERT: omnhbcba08-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1375638748] HOST ALERT: omnhbcba06-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1375638748] HOST ALERT: omnhbcba08-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1375638748] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375638798] SERVICE ALERT: omnhbcba12-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1375638808] HOST ALERT: omnhbcba12-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba01-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba02-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba03-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba04-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba05-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba06-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba07-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba08-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba09-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba10-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba11-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] SERVICE DOWNTIME ALERT: omnhbcba12-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba01-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba02-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba03-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba04-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba05-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba06-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba07-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba08-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba09-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba10-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba11-ny;STARTED; Host has entered a period of scheduled downtime
[1375638949] HOST DOWNTIME ALERT: omnhbcba12-ny;STARTED; Host has entered a period of scheduled downtime
[1375641078] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375641078] SERVICE ALERT: omnhbcba11-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1375641108] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375641108] HOST ALERT: omnhbcba11-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1375641138] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375641138] SERVICE ALERT: omnhbcba11-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1375641208] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375641208] HOST ALERT: omnhbcba11-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1375641248] SERVICE ALERT: omnhbcba04-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1375641258] SERVICE ALERT: omnhbcba11-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1375641258] HOST ALERT: omnhbcba04-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1375641268] HOST ALERT: omnhbcba11-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1375643231] HOST DOWNTIME ALERT: omnhbcba12-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643241] HOST DOWNTIME ALERT: omnhbcba01-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643247] HOST DOWNTIME ALERT: omnhbcba02-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643251] HOST DOWNTIME ALERT: omnhbcba05-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643256] HOST DOWNTIME ALERT: omnhbcba11-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643260] HOST DOWNTIME ALERT: omnhbcba10-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643264] SERVICE DOWNTIME ALERT: omnhbcba08-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643268] HOST DOWNTIME ALERT: omnhbcba06-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643290] HOST DOWNTIME ALERT: omnhbcba09-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643294] HOST DOWNTIME ALERT: omnhbcba08-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643298] HOST DOWNTIME ALERT: omnhbcba07-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643301] HOST DOWNTIME ALERT: omnhbcba03-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643306] SERVICE DOWNTIME ALERT: omnhbcba12-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643310] SERVICE DOWNTIME ALERT: omnhbcba11-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643314] SERVICE DOWNTIME ALERT: omnhbcba09-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643317] SERVICE DOWNTIME ALERT: omnhbcba03-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643321] HOST DOWNTIME ALERT: omnhbcba04-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1375643324] SERVICE DOWNTIME ALERT: omnhbcba10-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643328] SERVICE DOWNTIME ALERT: omnhbcba07-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643331] SERVICE DOWNTIME ALERT: omnhbcba05-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643335] SERVICE DOWNTIME ALERT: omnhbcba02-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643343] SERVICE DOWNTIME ALERT: omnhbcba06-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643346] SERVICE DOWNTIME ALERT: omnhbcba04-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1375643349] SERVICE DOWNTIME ALERT: omnhbcba01-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.