[1384504510] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job 09702407 remains in a queued state
[1384504570] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;SOFT;2;AME is idle:
[1384507230] SERVICE ALERT: rt-nycchdagg-b;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 102.67 ms
[1384507290] SERVICE ALERT: rt-nycchdagg-b;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1384527703] SERVICE ALERT: localhost;Root Partition;OK;HARD;4;OK - / space used=17%
[1384531333] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 117.87 ms
[1384531393] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 68.72 ms
[1384532043] SERVICE ALERT: deimos;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1384532063] HOST ALERT: deimos;DOWN;SOFT;1;(Host Check Timed Out)
[1384532093] SERVICE ALERT: deimos;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1384532103] HOST ALERT: deimos;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1384533003] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;SOFT;1;Job 09703193 remains in a queued state
[1384533063] SERVICE ALERT: nycnn-adbeame5;AME Job status;OK;SOFT;2;AME is idle:
[1384533483] SERVICE ALERT: deimos;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1384533523] HOST ALERT: deimos;DOWN;SOFT;1;(Host Check Timed Out)
[1384533543] SERVICE ALERT: deimos;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 0.21 ms
[1384533593] HOST ALERT: deimos;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1384533593] SERVICE ALERT: deimos;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1384533843] SERVICE ALERT: deimos;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1384533883] HOST ALERT: deimos;DOWN;SOFT;1;(Host Check Timed Out)
[1384533903] SERVICE ALERT: deimos;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1384533933] HOST ALERT: deimos;DOWN;SOFT;2;(Host Check Timed Out)
[1384534023] HOST ALERT: deimos;DOWN;SOFT;3;(Host Check Timed Out)
[1384534123] HOST ALERT: deimos;DOWN;SOFT;4;(Host Check Timed Out)
[1384534213] HOST ALERT: deimos;DOWN;SOFT;5;(Host Check Timed Out)
[1384534313] HOST ALERT: deimos;DOWN;SOFT;6;(Host Check Timed Out)
[1384534413] HOST ALERT: deimos;DOWN;SOFT;7;(Host Check Timed Out)
[1384534513] HOST ALERT: deimos;DOWN;SOFT;8;(Host Check Timed Out)
[1384534613] HOST ALERT: deimos;DOWN;SOFT;9;(Host Check Timed Out)
[1384534713] HOST ALERT: deimos;DOWN;HARD;10;(Host Check Timed Out)
[1384535018] HOST ALERT: deimos;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1384535038] SERVICE ALERT: deimos;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1384537508] SERVICE ALERT: wdcps002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1384537538] HOST ALERT: wdcps002;DOWN;SOFT;1;(Host Check Timed Out)
[1384537568] SERVICE ALERT: wdcps002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1384537628] HOST ALERT: wdcps002;DOWN;SOFT;2;(Host Check Timed Out)
[1384537718] HOST ALERT: wdcps002;DOWN;SOFT;3;(Host Check Timed Out)
[1384537818] HOST ALERT: wdcps002;DOWN;SOFT;4;(Host Check Timed Out)
[1384537888] HOST ALERT: wdcps002;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 5.39 ms
[1384537918] SERVICE ALERT: wdcps002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 5.42 ms
[1384538358] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 147.45 ms
[1384538418] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 86.21 ms
[1384538478] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 117.94 ms
[1384538538] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 58.58 ms
[1384539792] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384539792] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384539852] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384539852] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384539912] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;HARD;3;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384539912] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;HARD;3;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384539952] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384540012] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384540072] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;HARD;3;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384540301] SERVICE ALERT: localhost;Root Partition;UNKNOWN;SOFT;1;Unknown argument
[1384540361] SERVICE ALERT: localhost;Root Partition;UNKNOWN;SOFT;2;Unknown argument
[1384540421] SERVICE ALERT: localhost;Root Partition;UNKNOWN;SOFT;3;Unknown argument
[1384540481] SERVICE ALERT: localhost;Root Partition;UNKNOWN;HARD;4;Unknown argument
[1384543218] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;SOFT;1;Job 09704010 remains in a queued state
[1384543278] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;SOFT;2;AME is idle:
[1384545552] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 109.60 ms
[1384545572] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 115.59 ms
[1384545612] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 100.59 ms
[1384545632] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 96.52 ms
[1384545672] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 128.51 ms
[1384545792] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 92.88 ms
[1384546326] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1384546446] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1384546486] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1384546486] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1384546566] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;HARD;3;(Return code of 127 is out of bounds - plugin may be missing)
[1384546606] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1384546606] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1384546864] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1384546864] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1384546984] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1384546984] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1384547034] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1384547104] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;HARD;3;(Return code of 127 is out of bounds - plugin may be missing)
[1384547104] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;HARD;3;(Return code of 127 is out of bounds - plugin may be missing)
[1384547154] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1384547274] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;HARD;3;(Return code of 127 is out of bounds - plugin may be missing)
[1384547779] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384547779] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384547899] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384547899] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384547939] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1384549344] SERVICE ALERT: nycnn-splunk;Disk Space on /opt;UNKNOWN;SOFT;1;Remote command execution failed: bash: /usr/local/nagios/libexec/check_disk: No such file or directory
[1384549464] SERVICE ALERT: nycnn-splunk;Disk Space on /opt;UNKNOWN;SOFT;2;Remote command execution failed: bash: /usr/local/nagios/libexec/check_disk: No such file or directory
[1384549584] SERVICE ALERT: nycnn-splunk;Disk Space on /opt;UNKNOWN;HARD;3;Remote command execution failed: bash: /usr/local/nagios/libexec/check_disk: No such file or directory
[1384549919] SERVICE ALERT: nycnn-splunk;Disk Space on /opt;UNKNOWN;SOFT;1;Remote command execution failed: bash: /usr/local/nagios/libexec/check_disk: No such file or directory
[1384549969] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 09704038 remains in a queued state
[1384550029] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;AME is idle:
[1384550039] SERVICE ALERT: nycnn-splunk;Disk Space on /opt;UNKNOWN;SOFT;2;Remote command execution failed: bash: /usr/local/nagios/libexec/check_disk: No such file or directory
[1384550159] SERVICE ALERT: nycnn-splunk;Disk Space on /opt;UNKNOWN;HARD;3;Remote command execution failed: bash: /usr/local/nagios/libexec/check_disk: No such file or directory
[1384551859] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job 09704715 remains in a queued state
[1384551909] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;SOFT;2;AME is idle:
[1384552769] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.36 ms
[1384552819] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 101.43 ms
[1384552829] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 124.09 ms
[1384552879] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 118.28 ms
[1384552889] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 50.14 ms
[1384552939] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 96.36 ms
[1384553329] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 09704697 remains in a queued state
[1384553389] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;Job 09704697 progress is  6.891 :
[1384558429] SERVICE ALERT: nycnn-gridpushpull1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1384558449] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;1;(Host Check Timed Out)
[1384558489] SERVICE ALERT: nycnn-gridpushpull1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1384558549] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;2;(Host Check Timed Out)
[1384558649] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;3;(Host Check Timed Out)
[1384558739] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;4;(Host Check Timed Out)
[1384558829] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;5;(Host Check Timed Out)
[1384558929] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;6;(Host Check Timed Out)
[1384559029] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;7;(Host Check Timed Out)
[1384559119] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;8;(Host Check Timed Out)
[1384559219] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;9;(Host Check Timed Out)
[1384559319] HOST ALERT: nycnn-gridpushpull1;DOWN;HARD;10;(Host Check Timed Out)
[1384559909] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.69 ms
[1384559969] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 72.08 ms
[1384560579] SERVICE ALERT: nycnn-gridpushpull1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.43 ms
[1384560639] HOST ALERT: nycnn-gridpushpull1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1384560639] SERVICE ALERT: nycnn-gridpushpull1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1384562029] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1384562059] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;1;(Host Check Timed Out)
[1384562089] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1384562149] SERVICE ALERT: nycnn-arvtrtape;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1384562159] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;2;(Host Check Timed Out)
[1384562259] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;3;(Host Check Timed Out)
[1384562329] SERVICE ALERT: nycnn-arvtrtape;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1384562329] HOST ALERT: nycnn-arvtrtape;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1384562439] SERVICE ALERT: nycnn-arvtrtape;SSH;OK;HARD;1;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1384574599] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 108.09 ms
[1384574659] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.07 ms