[1376581528] SERVICE ALERT: wdcms002;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 100.72 ms
[1376581588] SERVICE ALERT: wdcms002;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1376583788] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376583788] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376583828] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376583828] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1376583848] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376583848] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376583908] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376583908] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376583908] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.23 ms
[1376583908] HOST ALERT: nycnn-ipvx15;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.25 ms
[1376583938] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376583938] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1376583958] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376583958] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376583968] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376583968] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376583968] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376583968] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584008] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584008] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1376584028] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584028] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584028] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.26 ms
[1376584028] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.23 ms
[1376584078] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376584078] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376584088] HOST ALERT: nycnn-ipvx18;DOWN;SOFT;1;(Host Check Timed Out)
[1376584088] HOST ALERT: nycnn-ipvx18e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584088] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1376584088] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1376584088] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.25 ms
[1376584088] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.31 ms
[1376584138] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376584138] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376584148] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584148] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584158] HOST ALERT: nycnn-ipvx18;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.47 ms
[1376584158] HOST ALERT: nycnn-ipvx18e1;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.29 ms
[1376584178] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584178] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1376584198] SERVICE ALERT: nycnn-ipvx18;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376584198] SERVICE ALERT: nycnn-ipvx18e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376584208] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584208] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584208] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584208] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584248] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584248] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1376584248] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376584248] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;2;(Host Check Timed Out)
[1376584258] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376584258] SERVICE FLAPPING ALERT: nycnn-ipvx15e1;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376584258] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376584258] SERVICE FLAPPING ALERT: nycnn-ipvx15;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376584268] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584268] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584268] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584268] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584268] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376584268] HOST FLAPPING ALERT: nycnn-ipvx15e1;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1376584268] HOST ALERT: nycnn-ipvx15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376584268] HOST FLAPPING ALERT: nycnn-ipvx15;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1376584308] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584308] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1376584318] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376584318] HOST FLAPPING ALERT: nycnn-ipvx17e1;STARTED; Host appears to have started flapping (23.8% change > 20.0% threshold)
[1376584318] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376584318] HOST FLAPPING ALERT: nycnn-ipvx17;STARTED; Host appears to have started flapping (23.8% change > 20.0% threshold)
[1376584318] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376584318] SERVICE FLAPPING ALERT: nycnn-ipvx17e1;PING;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1376584318] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376584318] SERVICE FLAPPING ALERT: nycnn-ipvx17;PING;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1376584318] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584318] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584328] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584328] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584348] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1376584348] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584358] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.24 ms
[1376584358] HOST FLAPPING ALERT: nycnn-ipvx16e1;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376584358] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.27 ms
[1376584358] HOST FLAPPING ALERT: nycnn-ipvx16;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376584378] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376584378] SERVICE FLAPPING ALERT: nycnn-ipvx16e1;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376584378] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376584378] SERVICE FLAPPING ALERT: nycnn-ipvx16;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376584378] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584378] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584378] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584388] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584408] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1376584428] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584428] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376584428] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;2;(Host Check Timed Out)
[1376584438] SERVICE ALERT: nycnn-ipvx11;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.35 ms
[1376584438] SERVICE ALERT: nycnn-ipvx11e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.31 ms
[1376584438] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584448] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584448] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584448] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584488] HOST ALERT: nycnn-ipvx18;DOWN;SOFT;1;(Host Check Timed Out)
[1376584488] HOST ALERT: nycnn-ipvx18e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584498] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376584498] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376584498] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376584498] HOST ALERT: nycnn-ipvx11;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376584508] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;2;(Host Check Timed Out)
[1376584508] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1376584508] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584508] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.30 ms
[1376584558] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376584558] HOST ALERT: nycnn-ipvx18;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376584558] HOST FLAPPING ALERT: nycnn-ipvx18;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1376584558] HOST ALERT: nycnn-ipvx18e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376584558] HOST FLAPPING ALERT: nycnn-ipvx18e1;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1376584558] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376584558] SERVICE ALERT: nycnn-ipvx18;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376584558] SERVICE FLAPPING ALERT: nycnn-ipvx18;PING;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1376584558] SERVICE ALERT: nycnn-ipvx18e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376584558] SERVICE FLAPPING ALERT: nycnn-ipvx18e1;PING;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1376584568] HOST ALERT: nycnn-ipvx13;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376584798] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584798] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584828] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1376584828] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584858] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584858] SERVICE FLAPPING ALERT: nycnn-ipvx11;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1376584858] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584858] SERVICE FLAPPING ALERT: nycnn-ipvx11e1;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1376584858] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584868] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584898] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1376584908] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376584918] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584918] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;2;(Host Check Timed Out)
[1376584918] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376584928] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376584978] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376584978] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376584978] HOST ALERT: nycnn-ipvx13;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.43 ms
[1376584978] HOST FLAPPING ALERT: nycnn-ipvx13;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1376584978] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376584978] SERVICE FLAPPING ALERT: nycnn-ipvx13;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1376584978] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376584978] HOST FLAPPING ALERT: nycnn-ipvx13e1;STARTED; Host appears to have started flapping (22.8% change > 20.0% threshold)
[1376584978] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376584978] SERVICE FLAPPING ALERT: nycnn-ipvx13e1;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1376584988] HOST ALERT: nycnn-ipvx11;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376584988] HOST FLAPPING ALERT: nycnn-ipvx11;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1376584988] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376584988] HOST FLAPPING ALERT: nycnn-ipvx11e1;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1376584998] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376584998] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1376585018] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1376585018] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585058] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585058] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585058] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585058] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585088] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1376585088] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376585098] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585098] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1376585108] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376585108] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376585118] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585118] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585178] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585178] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585178] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.24 ms
[1376585178] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.29 ms
[1376585208] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1376585208] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585228] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376585228] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376585238] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585238] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585238] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 80%, RTA = 0.23 ms
[1376585238] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 0.19 ms
[1376585268] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585268] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1376585298] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585298] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585308] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;2;(Host Check Timed Out)
[1376585308] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376585338] HOST FLAPPING ALERT: nycnn-ipvx15e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376585338] HOST FLAPPING ALERT: nycnn-ipvx15;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376585348] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;2;(Host Check Timed Out)
[1376585348] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376585348] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376585348] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376585358] HOST ALERT: nycnn-ipvx31;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376585358] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376585398] SERVICE FLAPPING ALERT: nycnn-ipvx15e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585398] SERVICE FLAPPING ALERT: nycnn-ipvx15;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585408] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376585408] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376585418] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376585418] HOST ALERT: nycnn-ipvx33;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376585458] SERVICE FLAPPING ALERT: nycnn-ipvx17e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585458] HOST FLAPPING ALERT: nycnn-ipvx17e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585458] SERVICE FLAPPING ALERT: nycnn-ipvx17;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585458] HOST FLAPPING ALERT: nycnn-ipvx17;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585478] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585478] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585478] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585478] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585518] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1376585518] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585518] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1376585518] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585518] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585518] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1376585518] SERVICE FLAPPING ALERT: nycnn-ipvx16e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585518] HOST FLAPPING ALERT: nycnn-ipvx16e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585518] SERVICE FLAPPING ALERT: nycnn-ipvx16;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585518] HOST FLAPPING ALERT: nycnn-ipvx16;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585538] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585538] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585538] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585538] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1376585538] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1376585538] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585538] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585538] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585568] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585568] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1376585588] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376585588] HOST FLAPPING ALERT: nycnn-ipvx30;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1376585588] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376585588] HOST FLAPPING ALERT: nycnn-ipvx30e1;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1376585588] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376585588] SERVICE FLAPPING ALERT: nycnn-ipvx30;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1376585588] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1376585588] SERVICE FLAPPING ALERT: nycnn-ipvx30e1;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1376585598] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585598] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585598] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585598] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585608] HOST ALERT: nycnn-ipvx31;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.25 ms
[1376585608] HOST FLAPPING ALERT: nycnn-ipvx31;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1376585608] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.20 ms
[1376585608] HOST FLAPPING ALERT: nycnn-ipvx31e1;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1376585618] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376585618] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;2;(Host Check Timed Out)
[1376585638] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1376585638] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585648] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376585648] SERVICE FLAPPING ALERT: nycnn-ipvx31;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1376585648] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376585648] SERVICE FLAPPING ALERT: nycnn-ipvx32e1;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1376585648] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376585648] SERVICE FLAPPING ALERT: nycnn-ipvx32;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1376585648] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376585648] SERVICE FLAPPING ALERT: nycnn-ipvx31e1;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1376585658] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585658] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376585658] HOST FLAPPING ALERT: nycnn-ipvx32e1;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1376585658] HOST ALERT: nycnn-ipvx32;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376585658] HOST FLAPPING ALERT: nycnn-ipvx32;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1376585658] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585658] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.23 ms
[1376585658] HOST FLAPPING ALERT: nycnn-ipvx33e1;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376585658] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.24 ms
[1376585658] HOST FLAPPING ALERT: nycnn-ipvx33;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376585698] SERVICE FLAPPING ALERT: nycnn-ipvx18;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585698] HOST FLAPPING ALERT: nycnn-ipvx18;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585698] SERVICE FLAPPING ALERT: nycnn-ipvx18e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585698] HOST FLAPPING ALERT: nycnn-ipvx18e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376585708] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1376585708] SERVICE FLAPPING ALERT: nycnn-ipvx33e1;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1376585708] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376585708] SERVICE FLAPPING ALERT: nycnn-ipvx33;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1376585718] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585718] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585728] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 38.35 ms
[1376585728] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.27 ms
[1376585748] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585748] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1376585768] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585768] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585768] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376585768] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376585778] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585778] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585798] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1376585798] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585818] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376585818] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1376585828] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585828] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585828] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376585828] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376585878] SERVICE ALERT: nycnn-ipvx14;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.25 ms
[1376585878] SERVICE ALERT: nycnn-ipvx14e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.31 ms
[1376585878] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.25 ms
[1376585878] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.25 ms
[1376585888] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585888] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585928] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585928] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1376585938] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376585938] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376585958] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585958] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376585958] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585958] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585958] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585958] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376585988] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1376585988] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585988] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376585988] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1376585998] HOST FLAPPING ALERT: nycnn-ipvx11;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1376585998] HOST FLAPPING ALERT: nycnn-ipvx11e1;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1376585998] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376585998] HOST ALERT: nycnn-ipvx29;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376586008] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586008] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586008] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376586008] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1376586018] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586018] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586018] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586018] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586038] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.45 ms
[1376586038] HOST FLAPPING ALERT: nycnn-ipvx34;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376586038] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.42 ms
[1376586038] HOST FLAPPING ALERT: nycnn-ipvx34e1;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1376586048] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1376586048] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376586058] SERVICE FLAPPING ALERT: nycnn-ipvx11;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586058] SERVICE FLAPPING ALERT: nycnn-ipvx11e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586068] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586068] SERVICE FLAPPING ALERT: nycnn-ipvx14;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376586068] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586068] SERVICE FLAPPING ALERT: nycnn-ipvx14e1;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376586068] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376586068] SERVICE FLAPPING ALERT: nycnn-ipvx34;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376586068] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376586068] SERVICE FLAPPING ALERT: nycnn-ipvx34e1;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376586078] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376586078] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;2;(Host Check Timed Out)
[1376586118] SERVICE FLAPPING ALERT: nycnn-ipvx13;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586118] HOST FLAPPING ALERT: nycnn-ipvx13;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586118] SERVICE FLAPPING ALERT: nycnn-ipvx13e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586118] HOST FLAPPING ALERT: nycnn-ipvx13e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586128] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 26.04 ms
[1376586128] HOST FLAPPING ALERT: nycnn-ipvx14;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1376586128] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.28 ms
[1376586128] HOST FLAPPING ALERT: nycnn-ipvx14e1;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1376586128] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376586128] SERVICE FLAPPING ALERT: nycnn-ipvx35e1;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376586128] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376586128] SERVICE FLAPPING ALERT: nycnn-ipvx35;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1376586138] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376586138] HOST FLAPPING ALERT: nycnn-ipvx35e1;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1376586138] HOST ALERT: nycnn-ipvx35;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376586138] HOST FLAPPING ALERT: nycnn-ipvx35;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1376586178] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376586178] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376586558] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586558] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586598] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376586598] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1376586618] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586618] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586618] HOST FLAPPING ALERT: nycnn-ipvx32e1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376586618] HOST FLAPPING ALERT: nycnn-ipvx32;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376586638] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;2;(Host Check Timed Out)
[1376586638] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;2;(Host Check Timed Out)
[1376586658] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586658] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586668] SERVICE ALERT: nycnn-ipvx29e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.26 ms
[1376586668] SERVICE FLAPPING ALERT: nycnn-ipvx29e1;PING;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1376586668] SERVICE ALERT: nycnn-ipvx29;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.31 ms
[1376586668] SERVICE FLAPPING ALERT: nycnn-ipvx29;PING;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1376586688] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1376586688] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376586708] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376586708] HOST FLAPPING ALERT: nycnn-ipvx29e1;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1376586708] HOST ALERT: nycnn-ipvx29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376586708] HOST FLAPPING ALERT: nycnn-ipvx29;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1376586718] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586718] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586728] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376586728] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376586728] SERVICE FLAPPING ALERT: nycnn-ipvx30;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586728] HOST FLAPPING ALERT: nycnn-ipvx30;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586728] SERVICE FLAPPING ALERT: nycnn-ipvx30e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586728] HOST FLAPPING ALERT: nycnn-ipvx30e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376586728] SERVICE FLAPPING ALERT: nycnn-ipvx31;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586728] HOST FLAPPING ALERT: nycnn-ipvx31;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586728] SERVICE FLAPPING ALERT: nycnn-ipvx32e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586728] SERVICE FLAPPING ALERT: nycnn-ipvx32;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586728] SERVICE FLAPPING ALERT: nycnn-ipvx31e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586728] HOST FLAPPING ALERT: nycnn-ipvx31e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586758] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 1.13 ms
[1376586758] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.42 ms
[1376586778] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376586778] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376586788] SERVICE FLAPPING ALERT: nycnn-ipvx33e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586788] HOST FLAPPING ALERT: nycnn-ipvx33e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586788] SERVICE FLAPPING ALERT: nycnn-ipvx33;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586788] HOST FLAPPING ALERT: nycnn-ipvx33;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376586898] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586898] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376586928] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1376586928] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1376586958] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586958] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376586998] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Server is down!
[1376587018] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.40 ms
[1376587018] HOST FLAPPING ALERT: nycnn-ipvx12;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376587018] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.22 ms
[1376587018] HOST FLAPPING ALERT: nycnn-ipvx12e1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376587018] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376587018] SERVICE FLAPPING ALERT: nycnn-ipvx12;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376587018] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376587018] SERVICE FLAPPING ALERT: nycnn-ipvx12e1;PING;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1376587058] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376587088] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;1;(Host Check Timed Out)
[1376587118] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376587118] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;HARD;2;(Service Check Timed Out)
[1376587148] HOST FLAPPING ALERT: nycnn-ipvx35e1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376587148] HOST FLAPPING ALERT: nycnn-ipvx35;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376587158] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;2;(Host Check Timed Out)
[1376587208] SERVICE FLAPPING ALERT: nycnn-ipvx34;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376587208] HOST FLAPPING ALERT: nycnn-ipvx34;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376587208] SERVICE FLAPPING ALERT: nycnn-ipvx34e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376587208] HOST FLAPPING ALERT: nycnn-ipvx34e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376587208] SERVICE FLAPPING ALERT: nycnn-ipvx35e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587208] SERVICE FLAPPING ALERT: nycnn-ipvx35;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587228] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376587238] HOST ALERT: nycnn-adbeame8;UP;SOFT;3;PING OK - Packet loss = 68%, RTA = 0.44 ms
[1376587248] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376587258] SERVICE FLAPPING ALERT: nycnn-ipvx14;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587258] HOST FLAPPING ALERT: nycnn-ipvx14;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587258] SERVICE FLAPPING ALERT: nycnn-ipvx14e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587258] HOST FLAPPING ALERT: nycnn-ipvx14e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587258] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;1;(Host Check Timed Out)
[1376587288] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1376587288] SERVICE ALERT: nycnn-adbeame8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1376587298] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.25 ms
[1376587298] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1376587298] SERVICE ALERT: nycnn-adbeame9;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1376587308] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376587338] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;1;(Host Check Timed Out)
[1376587358] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;2;(Host Check Timed Out)
[1376587358] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376587358] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1376587358] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;HARD;2;AME is idle:
[1376587358] SERVICE ALERT: nycnn-adbeame7;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1376587378] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1376587408] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1376587418] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376587418] HOST ALERT: nycnn-adbeame9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1376587428] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;2;(Host Check Timed Out)
[1376587448] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1376587458] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;1;(Host Check Timed Out)
[1376587478] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376587478] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376587498] HOST ALERT: nycnn-adbeame7;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1376587528] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1376587538] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376587558] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;2;(Host Check Timed Out)
[1376587568] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;1;(Host Check Timed Out)
[1376587598] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376587598] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;HARD;1;AME is idle:
[1376587638] HOST ALERT: nycnn-adbeame8;UP;SOFT;3;PING OK - Packet loss = 64%, RTA = 0.51 ms
[1376587638] HOST FLAPPING ALERT: nycnn-adbeame8;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376587648] SERVICE ALERT: nycnn-adbeame8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1376587648] SERVICE FLAPPING ALERT: nycnn-adbeame8;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1376587658] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376587658] SERVICE ALERT: nycnn-adbeame7;AME Job status;CRITICAL;SOFT;1;(Service Check Timed Out)
[1376587668] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;2;(Host Check Timed Out)
[1376587668] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376587688] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;1;(Host Check Timed Out)
[1376587698] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1376587708] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1376587708] SERVICE FLAPPING ALERT: nycnn-adbeame9;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1376587718] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376587718] HOST ALERT: nycnn-adbeame9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1376587718] HOST FLAPPING ALERT: nycnn-adbeame9;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376587728] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376587788] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1376587788] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;2;(Host Check Timed Out)
[1376587808] HOST FLAPPING ALERT: nycnn-ipvx29e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587808] HOST FLAPPING ALERT: nycnn-ipvx29;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376587828] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1376587828] SERVICE FLAPPING ALERT: nycnn-adbeame7;PING;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1376587838] HOST ALERT: nycnn-adbeame7;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1376587838] HOST FLAPPING ALERT: nycnn-adbeame7;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376587838] SERVICE ALERT: nycnn-adbeame9;AME Job status;OK;HARD;1;AME is idle:
[1376587868] SERVICE FLAPPING ALERT: nycnn-ipvx29e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376587868] SERVICE FLAPPING ALERT: nycnn-ipvx29;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376587868] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 75%, RTA = 0.44 ms
[1376587868] HOST FLAPPING ALERT: nycnn-adbeame10;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1376587898] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1376587898] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1376587968] SERVICE ALERT: nycnn-adbeame7;AME Job status;OK;HARD;2;AME is idle:
[1376588158] SERVICE FLAPPING ALERT: nycnn-ipvx12;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376588158] HOST FLAPPING ALERT: nycnn-ipvx12;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376588158] SERVICE FLAPPING ALERT: nycnn-ipvx12e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376588158] HOST FLAPPING ALERT: nycnn-ipvx12e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376588668] SERVICE FLAPPING ALERT: nycnn-adbeame8;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376588788] SERVICE FLAPPING ALERT: nycnn-adbeame9;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376588908] SERVICE FLAPPING ALERT: nycnn-adbeame7;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376588918] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1376590408] HOST FLAPPING ALERT: nycnn-adbeame9;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376590428] HOST FLAPPING ALERT: nycnn-adbeame8;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376590528] HOST FLAPPING ALERT: nycnn-adbeame7;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376590658] HOST FLAPPING ALERT: nycnn-adbeame10;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376591448] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job 09481591 remains in a queued state
[1376591498] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;SOFT;2;AME is idle:
[1376604458] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;SOFT;1;Job 09482533 remains in a queued state
[1376604528] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;SOFT;2;Job 09482533 progress is  27.59 :
[1376620658] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 09482993 remains in a queued state
[1376620718] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;AME is idle: