[1388905532] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1388905572] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905572] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905582] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905592] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 0.21 ms
[1388905592] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.20 ms
[1388905592] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905592] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1388905602] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905602] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905602] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905602] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1388905612] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905612] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905612] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905612] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905612] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905612] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905612] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905612] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905612] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905612] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.20 ms
[1388905622] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905622] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905622] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905622] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905622] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx35;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905632] SERVICE ALERT: nycnn-ipvx9;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905632] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905632] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1388905632] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905642] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905642] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905642] SERVICE ALERT: nycnn-ipvx15;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905642] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905642] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905642] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1388905642] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1388905652] SERVICE ALERT: nycnn-ipvx5;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905652] SERVICE ALERT: nycnn-ipvx2;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905652] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905652] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905652] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905662] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905662] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905662] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905662] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905672] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905672] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;2;(Host Check Timed Out)
[1388905672] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905672] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905672] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905672] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905672] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905672] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905672] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905672] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905672] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905672] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905682] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905682] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905682] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905682] SERVICE ALERT: nycnn-ipvx27;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905682] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905682] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905682] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905682] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1388905682] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905682] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905682] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1388905682] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905682] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1388905692] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905692] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905692] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905692] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905692] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905692] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1388905692] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1388905692] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1388905702] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;2;(Host Check Timed Out)
[1388905702] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905702] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905702] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905702] SERVICE ALERT: nycnn-ipvx16;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905702] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905702] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905702] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.25 ms
[1388905702] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905702] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905702] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1388905702] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905702] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1388905702] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905702] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1388905702] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905702] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905702] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 1.50 ms
[1388905702] SERVICE ALERT: nycnn-ipvx19;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905702] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905702] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.45 ms
[1388905702] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.27 ms
[1388905702] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.46 ms
[1388905702] HOST ALERT: nycnn-ipvx19;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.45 ms
[1388905702] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905702] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1388905712] SERVICE ALERT: nycnn-ipvx14e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx14;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx14e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.36 ms
[1388905712] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.55 ms
[1388905712] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1388905712] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905712] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1388905712] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1388905712] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905712] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905712] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1388905712] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905712] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1388905712] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1388905712] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905712] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1388905712] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1388905712] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1388905712] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905712] SERVICE ALERT: nycnn-ipvx10e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx19;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] SERVICE ALERT: nycnn-ipvx1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1388905712] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1388905712] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905712] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1388905712] SERVICE ALERT: nycnn-ipvx10;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1388905712] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388905712] SERVICE ALERT: nycnn-ipvx23e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] SERVICE ALERT: nycnn-ipvx21;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] SERVICE ALERT: nycnn-ipvx12;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx11;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx13;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905712] SERVICE ALERT: nycnn-ipvx21e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 1.29 ms
[1388905712] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 1.40 ms
[1388905712] SERVICE ALERT: nycnn-ipvx34e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 50%, RTA = 0.28 ms
[1388905712] SERVICE ALERT: nycnn-ipvx34;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 50%, RTA = 0.31 ms
[1388905712] HOST ALERT: nycnn-ipvx28;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.49 ms
[1388905712] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905712] SERVICE ALERT: nycnn-ipvx7e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.47 ms
[1388905712] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1388905712] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1388905712] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.30 ms
[1388905712] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.83 ms
[1388905712] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] SERVICE ALERT: nycnn-ipvx21;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905712] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.27 ms
[1388905712] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.29 ms
[1388905712] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.34 ms
[1388905712] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.43 ms
[1388905722] HOST ALERT: nycnn-ipvx29;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.50 ms
[1388905722] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.29 ms
[1388905722] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.25 ms
[1388905722] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905722] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905722] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.34 ms
[1388905722] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905722] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905722] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905722] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388905722] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905722] HOST ALERT: nycnn-ipvx25;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 27.01 ms
[1388905722] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905722] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.31 ms
[1388905722] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905722] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.39 ms
[1388905722] HOST ALERT: nycnn-ipvx15;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.48 ms
[1388905722] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905722] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905722] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905722] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905722] SERVICE ALERT: nycnn-ipvx22e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905722] SERVICE ALERT: nycnn-ipvx2e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905722] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905722] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388905722] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905722] SERVICE ALERT: nycnn-ipvx20;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905722] SERVICE ALERT: nycnn-ipvx1e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905732] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1388905732] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905732] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;2;(Host Check Timed Out)
[1388905732] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;2;(Host Check Timed Out)
[1388905732] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905732] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905732] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905732] HOST ALERT: nycnn-ipvx2;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.27 ms
[1388905732] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1388905732] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905732] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905732] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905732] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905732] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1388905732] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905732] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.24 ms
[1388905732] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905742] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;2;(Host Check Timed Out)
[1388905742] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905742] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.44 ms
[1388905742] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905742] HOST ALERT: nycnn-ipvx3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905742] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.41 ms
[1388905742] SERVICE ALERT: nycnn-ipvx24e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905742] SERVICE ALERT: nycnn-ipvx8e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905742] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905742] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1388905742] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1388905742] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905742] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1388905752] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;2;(Host Check Timed Out)
[1388905752] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905752] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1388905752] SERVICE ALERT: nycnn-ipvx27;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 0.48 ms
[1388905752] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.28 ms
[1388905752] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.31 ms
[1388905752] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905752] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905752] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905752] HOST ALERT: nycnn-ipvx10;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1388905752] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905752] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1388905752] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905752] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905752] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905752] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388905752] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388905762] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.45 ms
[1388905762] HOST ALERT: nycnn-ipvx9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905762] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905762] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1388905762] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.28 ms
[1388905762] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388905762] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905772] SERVICE ALERT: nycnn-ipvx8;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388905772] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1388905772] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905772] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905772] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905772] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905772] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905772] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905772] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905772] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905772] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905772] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905772] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388905772] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1388905772] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1388905772] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388905772] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388905772] SERVICE ALERT: nycnn-ipvx19;XCode Service;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905772] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1388905772] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1388905772] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905772] SERVICE ALERT: nycnn-ipvx23e1;XCode Service;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905772] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 1.68 ms
[1388905772] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1388905782] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905782] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;2;(Host Check Timed Out)
[1388905782] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388905782] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905782] HOST ALERT: nycnn-ipvx31;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1388905782] HOST ALERT: nycnn-ipvx13;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.41 ms
[1388905792] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905792] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905802] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1388905802] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1388905812] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905812] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1388905812] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905812] SERVICE ALERT: nycnn-ipvx8;XCode Service;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905822] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1388905832] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388905882] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905892] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905922] SERVICE ALERT: nycnn-ipvx35;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905922] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905922] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905922] SERVICE ALERT: nycnn-ipvx9;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905922] SERVICE ALERT: nycnn-ipvx15;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905922] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905932] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905932] SERVICE ALERT: nycnn-ipvx5;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905932] SERVICE ALERT: nycnn-ipvx2;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905942] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905952] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905962] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905962] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905962] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905962] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905962] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905962] SERVICE ALERT: nycnn-ipvx27;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905972] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905972] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905972] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905972] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905982] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx14e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx10e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx10;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388905992] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx11;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx12;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx14;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx13;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx16;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx14e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx19;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906002] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx21;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx22e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx1e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx20;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx21;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx21e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906012] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906022] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906022] SERVICE ALERT: nycnn-ipvx2e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906032] SERVICE ALERT: nycnn-ipvx24e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906042] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906042] SERVICE ALERT: nycnn-ipvx8e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1388906052] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388906072] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388909092] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909112] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388909112] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388909112] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388909112] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388909132] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1388909142] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1388909152] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1388909152] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1388909152] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1388909152] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909172] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1388909172] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388909172] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388909172] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388909172] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1388909172] SERVICE ALERT: nycnn-ipvstrm1;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909182] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909182] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909212] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;1;(Host Check Timed Out)
[1388909222] SERVICE ALERT: nycnn-ipvstrm1;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909222] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909232] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1388909232] SERVICE ALERT: nycnn-ipvstrm1;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909232] SERVICE ALERT: nycnn-ipvstrm3;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909232] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909232] SERVICE ALERT: nycnn-ipvstrm2;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909232] SERVICE ALERT: nycnn-ipvstrm4;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909232] SERVICE ALERT: nycnn-ipvstrm3;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909232] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1388909232] SERVICE ALERT: nycnn-ipvstrm4;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909242] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1388909242] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1388909252] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1388909282] HOST ALERT: nycnn-ipvapilic;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388909282] SERVICE ALERT: nycnn-ipvapilic;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388909282] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1388909282] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1388909292] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1388909292] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388909292] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1388909312] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1388909322] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING OK - Packet loss = 37%, RTA = 0.26 ms
[1388909342] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1388909342] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1388909432] SERVICE ALERT: nycnn-ipvstrm2;SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1388909432] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1388909462] SERVICE ALERT: nycnn-ipvstrm1;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1388909462] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1126.92GB (80.7%), Free=270.27GB (19.3%)
[1388909462] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1126.70GB (80.6%), Free=270.49GB (19.4%)
[1388909512] SERVICE ALERT: nycnn-ipvstrm1;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1388909512] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1388909512] SERVICE ALERT: nycnn-ipvstrm1;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388909512] SERVICE ALERT: nycnn-ipvstrm3;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1388909512] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1388909512] SERVICE ALERT: nycnn-ipvstrm2;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388909512] SERVICE ALERT: nycnn-ipvstrm4;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1388909512] SERVICE ALERT: nycnn-ipvstrm3;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388909522] SERVICE ALERT: nycnn-ipvstrm4;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1388909582] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1135.68GB (81.3%), Free=261.51GB (18.7%)