[1332893614] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5240 bytes in 0.016 second response time
[1332893614] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;UNKNOWN;SOFT;1;Remote command execution failed: Host key verification failed.
[1332893654] SERVICE ALERT: nyhdscheduler;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 400 Bad Request - 480 bytes in 0.001 second response time
[1332893674] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5240 bytes in 0.001 second response time
[1332893674] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;UNKNOWN;SOFT;2;Remote command execution failed: Host key verification failed.
[1332893674] SERVICE ALERT: omnmas2app-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1332893714] HOST ALERT: omnmas2app-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1332893714] SERVICE ALERT: nyhdscheduler;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 400 Bad Request - 480 bytes in 0.001 second response time
[1332893734] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5240 bytes in 0.001 second response time
[1332893734] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;UNKNOWN;HARD;3;Remote command execution failed: Host key verification failed.
[1332893734] SERVICE ALERT: omnmas2app-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1332893774] HOST ALERT: omnmas2app-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1332893774] SERVICE ALERT: nyhdscheduler;HTTP;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 480 bytes in 0.001 second response time
[1332893794] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5240 bytes in 0.001 second response time
[1332893874] HOST ALERT: omnmas2app-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1332893974] HOST ALERT: omnmas2app-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1332894074] HOST ALERT: omnmas2app-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1332894084] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;3;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1332894174] HOST ALERT: omnmas2app-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1332894274] HOST ALERT: omnmas2app-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1332894374] HOST ALERT: omnmas2app-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1332894474] HOST ALERT: omnmas2app-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1332894574] HOST ALERT: omnmas2app-ny;DOWN;HARD;10;(Host Check Timed Out)