[1342707267] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 28%, RTA = 99.35 ms
[1342707307] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 227.33 ms
[1342707327] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 189.96 ms
[1342707367] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 213.66 ms
[1342707387] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 230.61 ms
[1342707427] SERVICE ALERT: nyvideodrop;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 173.26 ms
[1342708747] SERVICE ALERT: nyvideodrop;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342708767] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342709167] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 239.85 ms
[1342709187] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 188.72 ms
[1342709227] SERVICE ALERT: nyvideodrop;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342709247] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342723387] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 189.51 ms
[1342723407] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 235.46 ms
[1342723447] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 232.83 ms
[1342723467] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 204.47 ms
[1342723507] SERVICE ALERT: nyvideodrop;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 223.74 ms
[1342723527] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 227.14 ms
[1342723647] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 40.16 ms
[1342723687] SERVICE ALERT: nyvideodrop;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.90 ms
[1342723867] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 185.37 ms
[1342723887] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 236.83 ms
[1342723927] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 193.34 ms
[1342723947] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 230.77 ms
[1342723987] SERVICE ALERT: nyvideodrop;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 182.88 ms
[1342724007] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 16%, RTA = 199.26 ms
[1342725067] SERVICE ALERT: nyvideodrop;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342725087] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342725147] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 199.16 ms
[1342725187] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 203.77 ms
[1342725207] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 208.75 ms
[1342725247] SERVICE ALERT: nyvideodrop;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342725267] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342739647] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 226.89 ms
[1342739707] SERVICE ALERT: nyvideodrop;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342742487] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 16%, RTA = 239.60 ms
[1342742527] SERVICE ALERT: nyvideodrop;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 237.49 ms
[1342742547] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 186.73 ms
[1342742587] SERVICE ALERT: nyvideodrop;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342742607] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms