[1383289548] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 09666275 remains in a queued state
[1383289608] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;AME is idle:
[1383328438] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job 09667783 remains in a queued state
[1383328498] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;SOFT;2;Job 09667783 progress is  17.02 :
[1383329159] HOST DOWNTIME ALERT: nycnn-adbeame10;STARTED; Host has entered a period of scheduled downtime
[1383329182] HOST DOWNTIME ALERT: nycnn-adbeame1;STARTED; Host has entered a period of scheduled downtime
[1383329210] HOST DOWNTIME ALERT: nycnn-adbeame2;STARTED; Host has entered a period of scheduled downtime
[1383329233] HOST DOWNTIME ALERT: nycnn-adbeame3;STARTED; Host has entered a period of scheduled downtime
[1383329257] HOST DOWNTIME ALERT: nycnn-adbeame4;STARTED; Host has entered a period of scheduled downtime
[1383329280] SERVICE DOWNTIME ALERT: nycnn-adbeame5;PING;STARTED; Service has entered a period of scheduled downtime
[1383329280] SERVICE DOWNTIME ALERT: nycnn-adbeame5;AME Job status;STARTED; Service has entered a period of scheduled downtime
[1383329719] HOST DOWNTIME ALERT: nycnn-adbeame5;STARTED; Host has entered a period of scheduled downtime
[1383329737] HOST DOWNTIME ALERT: nycnn-adbeame6;STARTED; Host has entered a period of scheduled downtime
[1383329756] HOST DOWNTIME ALERT: nycnn-adbeame7;STARTED; Host has entered a period of scheduled downtime
[1383329774] HOST DOWNTIME ALERT: nycnn-adbeame8;STARTED; Host has entered a period of scheduled downtime
[1383329795] HOST DOWNTIME ALERT: nycnn-adbeame9;STARTED; Host has entered a period of scheduled downtime
[1383331978] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383332008] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;SOFT;1;(Service Check Timed Out)
[1383332008] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1383332038] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383332078] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;HARD;2;(Service Check Timed Out)
[1383332108] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1383332148] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1383332158] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1383332278] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383332318] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1383332338] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383332418] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1383332488] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1383332488] HOST FLAPPING ALERT: nycnn-adbeame10;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1383332508] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1383332508] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1383332618] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;HARD;2;AME is idle:
[1383333528] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383333778] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1383333838] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1383333898] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1383335188] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383335228] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;1;(Host Check Timed Out)
[1383335228] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383335248] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383335258] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;1;(Host Check Timed Out)
[1383335278] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383335278] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;2;(Host Check Timed Out)
[1383335288] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383335298] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383335308] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1383335328] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383335338] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383335348] HOST ALERT: nycnn-adbeame5;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1383335358] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;2;(Host Check Timed Out)
[1383335358] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1383335378] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383335398] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1383335408] SERVICE ALERT: nycnn-adbeame3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1383335418] HOST ALERT: nycnn-adbeame3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383335448] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1383335458] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 16%, RTA = 0.37 ms
[1383335538] SERVICE ALERT: nycnn-adbeame5;AME Job status;OK;HARD;1;AME is idle:
[1383335588] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383335628] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;1;(Host Check Timed Out)
[1383335638] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383335648] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383335668] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;SOFT;1;(Service Check Timed Out)
[1383335668] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1383335698] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383335728] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;2;(Host Check Timed Out)
[1383335768] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1383335808] HOST ALERT: nycnn-adbeame3;UP;SOFT;3;PING OK - Packet loss = 73%, RTA = 0.54 ms
[1383335808] HOST FLAPPING ALERT: nycnn-adbeame3;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1383335818] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383335828] SERVICE ALERT: nycnn-adbeame3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1383335828] SERVICE FLAPPING ALERT: nycnn-adbeame3;PING;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1383335828] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;1;(Host Check Timed Out)
[1383335848] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 58%, RTA = 0.48 ms
[1383335868] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;HARD;1;AME is idle:
[1383335868] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1383335868] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1383335878] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383335928] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;2;(Host Check Timed Out)
[1383335958] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383335978] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;HARD;2;AME is idle:
[1383335978] SERVICE FLAPPING ALERT: nycnn-adbeame10;AME Job status;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1383336018] HOST ALERT: nycnn-adbeame1;UP;SOFT;3;PING WARNING - Packet loss = 83%, RTA = 0.47 ms
[1383336048] SERVICE ALERT: nycnn-adbeame1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1383336178] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383336208] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;1;(Host Check Timed Out)
[1383336238] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383336308] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;2;(Host Check Timed Out)
[1383336318] HOST DOWNTIME ALERT: nycnn-adbeame10;STOPPED; Host has exited from a period of scheduled downtime
[1383336370] HOST DOWNTIME ALERT: nycnn-adbeame1;STOPPED; Host has exited from a period of scheduled downtime
[1383336388] HOST ALERT: nycnn-adbeame1;UP;SOFT;3;PING OK - Packet loss = 61%, RTA = 0.59 ms
[1383336388] HOST FLAPPING ALERT: nycnn-adbeame1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1383336408] SERVICE ALERT: nycnn-adbeame1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383336408] SERVICE FLAPPING ALERT: nycnn-adbeame1;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1383336438] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1383336498] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1383336498] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;HARD;1;AME is idle:
[1383336608] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383336638] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;1;(Host Check Timed Out)
[1383336668] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383336738] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;2;(Host Check Timed Out)
[1383336768] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1383336788] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1383336798] HOST ALERT: nycnn-adbeame4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1383336838] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1383336848] SERVICE FLAPPING ALERT: nycnn-adbeame3;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383336879] HOST DOWNTIME ALERT: nycnn-ms2app4;STARTED; Host has entered a period of scheduled downtime
[1383336888] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1383336888] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridA/Temp/.gridamounted not found or not a regular file
[1383336888] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383336898] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1383336948] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;SOFT;2;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1383336948] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;SOFT;2;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1383336958] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1383337018] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.007 second response time
[1383337058] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;SOFT;1;Server is down!
[1383337088] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337128] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;1;(Host Check Timed Out)
[1383337138] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337148] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337158] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;(Service Check Timed Out)
[1383337178] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;1;(Host Check Timed Out)
[1383337178] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;HARD;2;(Service Check Timed Out)
[1383337198] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337218] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;HARD;2;(Service Check Timed Out)
[1383337218] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;2;(Host Check Timed Out)
[1383337278] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;2;(Host Check Timed Out)
[1383337288] HOST ALERT: nycnn-adbeame4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1383337288] HOST FLAPPING ALERT: nycnn-adbeame4;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1383337308] SERVICE ALERT: nycnn-adbeame1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1383337318] HOST ALERT: nycnn-adbeame1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1383337328] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1383337328] SERVICE FLAPPING ALERT: nycnn-adbeame4;PING;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1383337418] SERVICE ALERT: nycnn-adbeame4;AME Job status;OK;HARD;2;AME is idle:
[1383337438] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337478] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;1;(Host Check Timed Out)
[1383337498] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337568] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337578] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;2;(Host Check Timed Out)
[1383337588] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337598] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;1;(Host Check Timed Out)
[1383337608] SERVICE ALERT: nycnn-adbeame1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1383337618] HOST ALERT: nycnn-adbeame1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383337618] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;1;(Host Check Timed Out)
[1383337628] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337628] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337648] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337668] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;1;(Host Check Timed Out)
[1383337688] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337688] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;2;(Host Check Timed Out)
[1383337718] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;2;(Host Check Timed Out)
[1383337758] HOST ALERT: nycnn-adbeame2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383337768] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;HARD;2;AME is idle:
[1383337768] SERVICE FLAPPING ALERT: nycnn-adbeame1;AME Job status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1383337768] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;2;(Host Check Timed Out)
[1383337788] HOST ALERT: nycnn-adbeame5;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1383337788] HOST FLAPPING ALERT: nycnn-adbeame5;STARTED; Host appears to have started flapping (20.0% change > 20.0% threshold)
[1383337808] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383337808] SERVICE ALERT: nycnn-adbeame6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1383337808] HOST ALERT: nycnn-adbeame6;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1383337818] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1383337888] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337928] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;1;(Host Check Timed Out)
[1383337928] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337928] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337928] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337948] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337968] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1383337968] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1383337968] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;1;(Host Check Timed Out)
[1383337988] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383337988] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337988] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383337988] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338028] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;1;(Host Check Timed Out)
[1383338028] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;2;(Host Check Timed Out)
[1383338038] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1383338042] HOST DOWNTIME ALERT: nycnn-adbeame1;STARTED; Host has entered a period of scheduled downtime
[1383338048] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338055] HOST DOWNTIME ALERT: nycnn-adbeame10;STARTED; Host has entered a period of scheduled downtime
[1383338058] SERVICE ALERT: nycnn-adbeame5;PING;WARNING;HARD;2;PING WARNING - Packet loss = 37%, RTA = 0.48 ms
[1383338058] SERVICE FLAPPING ALERT: nycnn-adbeame5;PING;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1383338058] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;2;(Host Check Timed Out)
[1383338058] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;2;(Host Check Timed Out)
[1383338058] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;2;(Host Check Timed Out)
[1383338098] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;2;(Host Check Timed Out)
[1383338098] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1383338098] HOST ALERT: nycnn-adbeame5;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1383338118] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1383338138] HOST ALERT: nycnn-adbeame6;UP;SOFT;3;PING OK - Packet loss = 54%, RTA = 0.51 ms
[1383338138] HOST FLAPPING ALERT: nycnn-adbeame6;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1383338138] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Server is down!
[1383338148] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383338158] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;3;(Host Check Timed Out)
[1383338158] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;3;(Host Check Timed Out)
[1383338158] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;SOFT;2;AME is idle:
[1383338158] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1383338158] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383338158] SERVICE FLAPPING ALERT: nycnn-adbeame2;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1383338168] SERVICE ALERT: nycnn-adbeame6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1383338168] SERVICE FLAPPING ALERT: nycnn-adbeame6;PING;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1383338168] HOST ALERT: nycnn-adbeame2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1383338168] HOST FLAPPING ALERT: nycnn-adbeame2;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1383338258] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;4;(Host Check Timed Out)
[1383338258] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;4;(Host Check Timed Out)
[1383338298] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383338298] SERVICE ALERT: nycnn-adbeame5;AME Job status;OK;HARD;2;AME is idle:
[1383338298] SERVICE FLAPPING ALERT: nycnn-adbeame5;AME Job status;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1383338328] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;1;(Host Check Timed Out)
[1383338338] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383338348] SERVICE FLAPPING ALERT: nycnn-adbeame4;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383338358] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1383338358] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;5;(Host Check Timed Out)
[1383338358] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;5;(Host Check Timed Out)
[1383338358] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338388] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;HARD;1;AME is idle:
[1383338398] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338408] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383338418] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;2;(Host Check Timed Out)
[1383338438] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383338438] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;1;(Host Check Timed Out)
[1383338458] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1383338458] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;6;(Host Check Timed Out)
[1383338458] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;6;(Host Check Timed Out)
[1383338468] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338488] HOST ALERT: nycnn-adbeame9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1383338498] SERVICE ALERT: nycnn-adbeame7;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383338508] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1383338518] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1383338528] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383338538] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;2;(Host Check Timed Out)
[1383338558] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;7;(Host Check Timed Out)
[1383338558] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;7;(Host Check Timed Out)
[1383338588] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.44 ms
[1383338598] HOST ALERT: nycnn-adbeame7;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1383338598] HOST FLAPPING ALERT: nycnn-adbeame3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383338638] HOST ALERT: nycnn-ipvx27;UP;SOFT;8;PING OK - Packet loss = 70%, RTA = 0.49 ms
[1383338638] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;8;PING OK - Packet loss = 70%, RTA = 0.39 ms
[1383338648] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1383338648] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1383338658] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383338688] SERVICE FLAPPING ALERT: nycnn-adbeame1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383338688] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;1;(Host Check Timed Out)
[1383338698] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383338718] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338738] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;SOFT;1;(Service Check Timed Out)
[1383338738] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1383338758] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338758] SERVICE ALERT: nycnn-adbeame7;AME Job status;CRITICAL;SOFT;1;Server is down!
[1383338788] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;2;(Host Check Timed Out)
[1383338828] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383338833] SERVICE DOWNTIME ALERT: nycnn-ms2app4;PING;STARTED; Service has entered a period of scheduled downtime
[1383338833] SERVICE DOWNTIME ALERT: nycnn-ms2app4;GridA_Resources_Folder_Count;STARTED; Service has entered a period of scheduled downtime
[1383338833] SERVICE DOWNTIME ALERT: nycnn-ms2app4;GridA_Render_Folder_Count;STARTED; Service has entered a period of scheduled downtime
[1383338833] SERVICE DOWNTIME ALERT: nycnn-ms2app4;GridA_Projects_Folder_Count;STARTED; Service has entered a period of scheduled downtime
[1383338838] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;2;(Host Check Timed Out)
[1383338858] HOST ALERT: nycnn-adbeame9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1383338858] HOST FLAPPING ALERT: nycnn-adbeame9;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1383338858] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;1;(Host Check Timed Out)
[1383338863] SERVICE DOWNTIME ALERT: msny;MS2 - Grid A Mount;STARTED; Service has entered a period of scheduled downtime
[1383338888] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1383338888] SERVICE FLAPPING ALERT: nycnn-adbeame9;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1383338888] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383338908] HOST ALERT: nycnn-adbeame10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1383338928] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1383338928] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1383338938] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;2;(Host Check Timed Out)
[1383338998] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1383339008] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1383339008] SERVICE FLAPPING ALERT: nycnn-adbeame7;PING;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1383339008] HOST ALERT: nycnn-adbeame7;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1383339008] HOST FLAPPING ALERT: nycnn-adbeame7;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1383339018] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1383339048] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;HARD;2;AME is idle:
[1383339048] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1383339048] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;1;(Host Check Timed Out)
[1383339058] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1383339078] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383339098] SERVICE ALERT: nycnn-adbeame7;AME Job status;OK;HARD;2;AME is idle:
[1383339108] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridA/Temp/.gridamounted not found or not a regular file
[1383339118] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1383339128] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;2;(Host Check Timed Out)
[1383339168] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;HARD;1;(Service Check Timed Out)
[1383339168] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - File /Volumes/GridA/Temp/.gridamounted not found or not a regular file
[1383339178] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383339188] SERVICE FLAPPING ALERT: nycnn-adbeame6;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383339198] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383339208] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;1;(Host Check Timed Out)
[1383339218] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;3;(Host Check Timed Out)
[1383339228] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - File /Volumes/GridA/Temp/.gridamounted not found or not a regular file
[1383339228] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;1;(Host Check Timed Out)
[1383339238] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383339243] SERVICE DOWNTIME ALERT: msny;PING;STARTED; Service has entered a period of scheduled downtime
[1383339243] SERVICE DOWNTIME ALERT: msny;MS2 - Grid B Mount;STARTED; Service has entered a period of scheduled downtime
[1383339243] SERVICE DOWNTIME ALERT: msny;IMQ Depth Check - Internal GF Queues;STARTED; Service has entered a period of scheduled downtime
[1383339243] SERVICE DOWNTIME ALERT: msny;HTTP;STARTED; Service has entered a period of scheduled downtime
[1383339247] HOST DOWNTIME ALERT: nycnn-ms2app4;STOPPED; Host has exited from a period of scheduled downtime
[1383339248] SERVICE FLAPPING ALERT: nycnn-adbeame2;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383339258] SERVICE FLAPPING ALERT: nycnn-adbeame5;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1383339258] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383339288] HOST ALERT: nycnn-adbeame8;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1383339292] HOST DOWNTIME ALERT: nycnn-ms2fsd;STARTED; Host has entered a period of scheduled downtime
[1383339308] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;2;(Host Check Timed Out)
[1383339308] SERVICE ALERT: nycnn-adbeame8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1383339318] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;2;(Host Check Timed Out)
[1383339324] HOST DOWNTIME ALERT: nycnn-ms2fsd1;STARTED; Host has entered a period of scheduled downtime
[1383339398] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;3;(Host Check Timed Out)
[1383339408] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;HARD;1;AME is idle:
[1383339418] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;3;(Host Check Timed Out)
[1383339488] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;4;(Host Check Timed Out)
[1383339498] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383339518] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;4;(Host Check Timed Out)
[1383339528] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;1;(Host Check Timed Out)
[1383339558] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383339588] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;5;(Host Check Timed Out)
[1383339618] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;5;(Host Check Timed Out)
[1383339628] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;2;(Host Check Timed Out)
[1383339688] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;6;(Host Check Timed Out)
[1383339718] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;6;(Host Check Timed Out)
[1383339718] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;HARD;1;Server is down!
[1383339718] HOST ALERT: nycnn-adbeame8;UP;SOFT;3;PING WARNING - Packet loss = 80%, RTA = 0.53 ms
[1383339718] HOST FLAPPING ALERT: nycnn-adbeame8;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1383339728] SERVICE ALERT: nycnn-adbeame8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1383339728] SERVICE FLAPPING ALERT: nycnn-adbeame8;PING;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1383339788] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;7;(Host Check Timed Out)
[1383339818] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;7;(Host Check Timed Out)
[1383339888] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;8;(Host Check Timed Out)
[1383339908] SERVICE FLAPPING ALERT: nycnn-adbeame9;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383339918] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;8;(Host Check Timed Out)
[1383339948] SERVICE FLAPPING ALERT: nycnn-adbeame10;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383339948] SERVICE ALERT: nycnn-ms2fsd1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1383339958] HOST ALERT: nycnn-ms2fsd1;UP;SOFT;9;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1383340008] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;HARD;1;AME is idle:
[1383340008] SERVICE FLAPPING ALERT: nycnn-adbeame8;AME Job status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1383340018] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;9;(Host Check Timed Out)
[1383340046] HOST DOWNTIME ALERT: nycnn-adbeame4;STOPPED; Host has exited from a period of scheduled downtime
[1383340070] SERVICE DOWNTIME ALERT: nycnn-adbeame5;PING;STOPPED; Service has exited from a period of scheduled downtime
[1383340070] SERVICE DOWNTIME ALERT: nycnn-adbeame5;AME Job status;STOPPED; Service has exited from a period of scheduled downtime
[1383340078] HOST FLAPPING ALERT: nycnn-adbeame4;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383340088] SERVICE FLAPPING ALERT: nycnn-adbeame7;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383340098] HOST ALERT: nycnn-ms2fsd;UP;SOFT;10;PING OK - Packet loss = 61%, RTA = 0.24 ms
[1383340098] SERVICE ALERT: nycnn-ms2fsd;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1383340128] SERVICE ALERT: msny;MS2 - Grid A Mount;UNKNOWN;HARD;3;Remote command execution failed: ssh: connect to host 10.164.219.196 port 22: Connection refused
[1383340128] SERVICE ALERT: msny;MS2 - Grid B Mount;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.219.196 port 22: Connection refused
[1383340128] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.219.196 port 22: Connection refused
[1383340148] SERVICE ALERT: msny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383340158] HOST ALERT: msny;DOWN;SOFT;1;(Host Check Timed Out)
[1383340198] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1383340198] SERVICE FLAPPING ALERT: msny;MS2 - Grid A Mount;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1383340198] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1383340198] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1383340208] SERVICE ALERT: msny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383340258] HOST ALERT: msny;DOWN;SOFT;2;(Host Check Timed Out)
[1383340308] HOST FLAPPING ALERT: nycnn-adbeame1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383340348] HOST ALERT: msny;DOWN;SOFT;3;(Host Check Timed Out)
[1383340418] HOST ALERT: msny;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1383340428] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;2;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1383340428] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1383340428] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;2;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1383340438] SERVICE ALERT: msny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1383340558] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.019 second response time
[1383340748] SERVICE FLAPPING ALERT: nycnn-adbeame8;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1383340808] HOST FLAPPING ALERT: nycnn-adbeame6;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383340888] HOST FLAPPING ALERT: nycnn-adbeame5;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383340948] HOST FLAPPING ALERT: nycnn-adbeame2;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383341388] SERVICE FLAPPING ALERT: msny;MS2 - Grid A Mount;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1383341648] HOST FLAPPING ALERT: nycnn-adbeame9;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383341698] HOST FLAPPING ALERT: nycnn-adbeame10;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383341798] HOST FLAPPING ALERT: nycnn-adbeame7;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383342415] SERVICE DOWNTIME ALERT: nycnn-ms2app4;PING;STOPPED; Service has exited from a period of scheduled downtime
[1383342415] SERVICE DOWNTIME ALERT: nycnn-ms2app4;GridA_Resources_Folder_Count;STOPPED; Service has exited from a period of scheduled downtime
[1383342415] SERVICE DOWNTIME ALERT: nycnn-ms2app4;GridA_Render_Folder_Count;STOPPED; Service has exited from a period of scheduled downtime
[1383342415] SERVICE DOWNTIME ALERT: nycnn-ms2app4;GridA_Projects_Folder_Count;STOPPED; Service has exited from a period of scheduled downtime
[1383342498] SERVICE FLAPPING ALERT: nycnn-adbeame5;AME Job status;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1383342508] HOST FLAPPING ALERT: nycnn-adbeame8;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383342829] SERVICE DOWNTIME ALERT: msny;PING;STOPPED; Service has exited from a period of scheduled downtime
[1383342829] SERVICE DOWNTIME ALERT: msny;MS2 - Grid B Mount;STOPPED; Service has exited from a period of scheduled downtime
[1383342829] SERVICE DOWNTIME ALERT: msny;MS2 - Grid A Mount;STOPPED; Service has exited from a period of scheduled downtime
[1383342829] SERVICE DOWNTIME ALERT: msny;IMQ Depth Check - Internal GF Queues;STOPPED; Service has exited from a period of scheduled downtime
[1383342829] SERVICE DOWNTIME ALERT: msny;HTTP;STOPPED; Service has exited from a period of scheduled downtime
[1383343168] SERVICE FLAPPING ALERT: nycnn-adbeame1;AME Job status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383344148] SERVICE ALERT: rt-nycchdagg-b;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 108.20 ms
[1383344208] SERVICE ALERT: rt-nycchdagg-b;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1383344448] SERVICE FLAPPING ALERT: nycnn-adbeame10;AME Job status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1383345223] HOST DOWNTIME ALERT: nycnn-adbeame1;STOPPED; Host has exited from a period of scheduled downtime
[1383345250] HOST DOWNTIME ALERT: nycnn-adbeame10;STOPPED; Host has exited from a period of scheduled downtime
[1383345262] HOST DOWNTIME ALERT: nycnn-adbeame2;STOPPED; Host has exited from a period of scheduled downtime
[1383345276] HOST DOWNTIME ALERT: nycnn-adbeame3;STOPPED; Host has exited from a period of scheduled downtime
[1383345290] HOST DOWNTIME ALERT: nycnn-adbeame5;STOPPED; Host has exited from a period of scheduled downtime
[1383345300] HOST DOWNTIME ALERT: nycnn-adbeame6;STOPPED; Host has exited from a period of scheduled downtime
[1383345321] HOST DOWNTIME ALERT: nycnn-adbeame7;STOPPED; Host has exited from a period of scheduled downtime
[1383345331] HOST DOWNTIME ALERT: nycnn-adbeame8;STOPPED; Host has exited from a period of scheduled downtime
[1383345340] HOST DOWNTIME ALERT: nycnn-adbeame9;STOPPED; Host has exited from a period of scheduled downtime
[1383345708] SERVICE FLAPPING ALERT: nycnn-adbeame8;AME Job status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1383346486] HOST DOWNTIME ALERT: nycnn-ms2fsd;STOPPED; Host has exited from a period of scheduled downtime
[1383346515] HOST DOWNTIME ALERT: nycnn-ms2fsd1;STOPPED; Host has exited from a period of scheduled downtime
[1383347568] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1383347628] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1383351828] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;HARD;2;AME is idle: