[1342811907] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342811947] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1342811967] SERVICE ALERT: wdcps001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812017] HOST ALERT: wdcps001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.50 ms
[1342812027] SERVICE ALERT: wdcps001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342812627] SERVICE ALERT: omnldap2-ny;LDAP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1342812627] SERVICE ALERT: omnldap1-ny;LDAP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1342812627] SERVICE ALERT: msny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nyarxlator3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nyarxlator4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nyarxlator1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nyarxlator2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbectrl1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbectrl2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-arapp1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvproxy1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-arxlator1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvproxy2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvsphn2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvsphn3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdir03e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk10-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk11-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk09-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk14-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk12-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk15-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk13-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk16-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk18-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk19-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk17-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox03e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox04e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox05e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnprox05-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk19h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk18h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk10h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk16h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk14h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk11h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk13h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk12h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1342812627] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-mamdb2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-mamdb1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-mamkey1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-mamkey2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ms2app2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdir03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk08-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-gridpushpull1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ms2app1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-gridpushpull2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk17h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk08h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk15h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk07h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: omnmdk09h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcps002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcod01;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcod02;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcod002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcpf002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcps003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcpf001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcps004;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcpf03;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcod001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcms002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812627] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nycnn-ms2app4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nycnn-ms2db1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nycnn-ms2db2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nycnn-ms2db3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nymamkey;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nyproxygenhd;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: nyhdscheduler;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnclda02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnclda02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnclda03e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnclda03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba05-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1342812637] SERVICE ALERT: omnhbcba08-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba11-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnmas1app-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812637] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812647] SERVICE ALERT: nycnn-mamfog;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812647] SERVICE ALERT: nycnn-splunk;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812647] SERVICE ALERT: nycnn-kickstart;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812647] SERVICE ALERT: spider;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812657] HOST ALERT: nycnn-ardb2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-arxlator1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-arxlator2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-arxlator3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-arxlator4;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812657] HOST ALERT: omnmdk16-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnldap2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnldap1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: msny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nyarxlator3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nyarxlator4;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nyarxlator1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nyarxlator2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame10;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame5;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame8;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame9;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame7;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame4;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbectrl1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbectrl2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-arapp1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvproxy1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvproxy2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvsphn2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvsphn3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-adbeame6;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvsphn5;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdir03e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdir04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx18e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk10-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk11-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx18;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk09-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk14-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk12-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk15-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk13-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk18-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk19-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk17-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox03e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox04e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox05e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnspsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnprox05-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk19h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk18h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk10h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk16h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk14h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk11h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk13h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk12h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx26;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-mamdb2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-mamkey1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-mamkey2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ms2app2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdir03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk08-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-gridpushpull1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ms2app1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx26e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-gridpushpull2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk17h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk08h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk15h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk07h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: omnmdk09h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcps002;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcod01;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcod02;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcod002;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcpf002;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcps003;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcpf001;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcps004;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcpf03;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcod001;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcms002;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1342812667] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1342812667] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;1;CRITICAL - Plugin timed out after 10 seconds
[1342812667] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas1db1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas1db-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas2app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas1db2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas2app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas2db-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas2db1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas2db2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdir01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdir01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdir02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdir02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk05-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk06-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk07-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk01h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk03h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk04h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk06h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk02h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812667] SERVICE ALERT: omnmdk05h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812677] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: nymamkey;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: nyproxygenhd;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: nyhdscheduler;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnclda01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnclda01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnclda02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnclda03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba05-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba08-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba07-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba11-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnmas1app-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] HOST ALERT: spider;DOWN;SOFT;1;(Host Check Timed Out)
[1342812677] SERVICE ALERT: nyvideodrop;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812687] HOST ALERT: nycnn-mamfog;DOWN;SOFT;1;(Host Check Timed Out)
[1342812687] HOST ALERT: nycnn-splunk;DOWN;SOFT;1;(Host Check Timed Out)
[1342812687] HOST ALERT: nycnn-kickstart;DOWN;SOFT;1;(Host Check Timed Out)
[1342812687] SERVICE ALERT: nyhdscheduler;HTTP;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1342812687] SERVICE ALERT: omnldap2-ny;LDAP;CRITICAL;HARD;2;CRITICAL - Socket timeout after 10 seconds
[1342812687] SERVICE ALERT: omnldap1-ny;LDAP;CRITICAL;HARD;2;CRITICAL - Socket timeout after 10 seconds
[1342812687] SERVICE ALERT: msny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nyarxlator3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nyarxlator4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nyarxlator1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nyarxlator2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbectrl1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbectrl2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-arapp1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvproxy1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-arxlator1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvproxy2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvsphn2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvsphn3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdir03e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk10-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk11-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk09-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk14-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk12-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk15-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk13-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk16-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk18-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk19-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk17-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox01e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox02e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox03e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox04e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox05e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnprox05-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk19h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk18h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk10h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk16h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk14h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk11h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk13h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk12h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1342812687] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-mamdb2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-mamdb1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-mamkey1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-mamkey2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ms2app2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdir03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk08-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-gridpushpull1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ms2app1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-gridpushpull2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk17h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk08h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk15h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk07h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: omnmdk09h0-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcps002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcod01;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcod02;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcod002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcms001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcpf002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcps003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcpf001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcps004;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcpf03;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcod001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcms002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812687] SERVICE ALERT: wdcps001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] HOST ALERT: nyvideodrop;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] SERVICE ALERT: nycnn-ms2app4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: nycnn-ms2db1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: nycnn-ms2db2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: nycnn-ms2db3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: nymamkey;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: nyproxygenhd;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: nyhdscheduler;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnclda02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnclda02e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnclda03e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnclda03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba05-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1342812697] SERVICE ALERT: omnhbcba08-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba11-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnmas1app-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342812697] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas1db1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas1db-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas2app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas2app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas2db-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas2db1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas2db2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdir01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdir01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] SERVICE ALERT: nycnn-mamfog;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 90%, RTA = 0.14 ms
[1342812697] HOST ALERT: omnmdir02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] SERVICE ALERT: nycnn-splunk;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 90%, RTA = 0.26 ms
[1342812697] HOST ALERT: omnmdir02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] SERVICE ALERT: nycnn-kickstart;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 90%, RTA = 0.22 ms
[1342812697] HOST ALERT: omnmdir04e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdk01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdk04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdk02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdk03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdk05-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdk06-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: omnmdk07-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812697] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812707] HOST ALERT: omnmdk01h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812707] HOST ALERT: omnmdk03h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812707] HOST ALERT: omnmdk04h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812707] HOST ALERT: omnmdk06h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812707] HOST ALERT: omnmdk02h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812707] HOST ALERT: omnmdk05h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812707] SERVICE ALERT: spider;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 80%, RTA = 0.25 ms
[1342812707] HOST ALERT: nycnn-adbeame5;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.36 ms
[1342812707] HOST ALERT: msny;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.19 ms
[1342812707] HOST ALERT: nycnn-ipvsphn3;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.20 ms
[1342812707] HOST ALERT: nycnn-ms2app2;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.45 ms
[1342812707] HOST ALERT: nycnn-ipvx18;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.19 ms
[1342812707] HOST ALERT: nycnn-ipvx18e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.19 ms
[1342812707] HOST ALERT: nycnn-ipvx26;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.18 ms
[1342812707] HOST ALERT: nycnn-ipvx26e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.18 ms
[1342812707] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.20 ms
[1342812707] HOST ALERT: nycnn-adbeame4;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.32 ms
[1342812707] HOST ALERT: nycnn-adbeame6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342812707] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342812707] HOST ALERT: nycnn-ipvx13;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812707] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812707] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812707] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812707] HOST ALERT: nycnn-ms2app3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342812707] HOST ALERT: nycnn-ms2app4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812717] SERVICE ALERT: msny;HTTP;OK;HARD;1;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.002 second response time
[1342812717] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;1;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1342812717] HOST ALERT: nycnn-ipvx29;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812717] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812717] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342812717] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812717] HOST ALERT: omnspsm-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812717] SERVICE ALERT: nycnn-ipvx3;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812717] SERVICE ALERT: omnmas1db1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812717] SERVICE ALERT: omnmas1db-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812717] SERVICE ALERT: omnmas2app1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812717] SERVICE ALERT: omnmas1db2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812717] SERVICE ALERT: omnmas2app2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342812717] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812717] SERVICE ALERT: omnmas2db-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342812717] SERVICE ALERT: omnmas2db1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812717] SERVICE ALERT: omnmas2db2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342812717] SERVICE ALERT: omnmdir01-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812717] SERVICE ALERT: omnmdir01e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1342812717] SERVICE ALERT: omnmdir02-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812717] SERVICE ALERT: omnmdir02e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342812717] SERVICE ALERT: omnmdir04e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342812717] SERVICE ALERT: omnmdk01-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1342812717] SERVICE ALERT: omnmdk04-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342812717] SERVICE ALERT: omnmdk02-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342812717] SERVICE ALERT: omnmdk03-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1342812717] SERVICE ALERT: omnmdk05-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1342812717] SERVICE ALERT: omnmdk06-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1342812717] SERVICE ALERT: omnmas1app2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812717] SERVICE ALERT: omnmdk07-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342812717] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812717] SERVICE ALERT: omnmdk01h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812717] SERVICE ALERT: omnmdk03h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1342812717] SERVICE ALERT: omnmdk04h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.25 ms
[1342812717] SERVICE ALERT: omnmdk06h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342812717] SERVICE ALERT: omnmdk02h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1342812717] SERVICE ALERT: omnmdk05h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 4.18 ms
[1342812727] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812727] HOST ALERT: omnmas1db1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812727] HOST ALERT: omnmas1db-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812727] HOST ALERT: omnmas2app1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342812727] HOST ALERT: omnmas1db2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342812727] HOST ALERT: omnmas2app2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812727] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812727] HOST ALERT: omnmas2db-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812727] HOST ALERT: omnmas2db1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812727] HOST ALERT: omnmas2db2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342812727] HOST ALERT: omnmdir01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342812727] HOST ALERT: omnmdir01e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1342812727] HOST ALERT: omnmdir02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812727] HOST ALERT: omnmdir02e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342812727] HOST ALERT: omnmdir04e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1342812727] HOST ALERT: omnmdk01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1342812727] HOST ALERT: omnmdk04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342812727] HOST ALERT: omnmdk02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342812727] HOST ALERT: omnmdk03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342812727] HOST ALERT: omnmdk05-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342812727] HOST ALERT: omnmdk06-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342812727] HOST ALERT: omnmas1app2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812727] HOST ALERT: omnmdk07-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342812727] HOST ALERT: omnmdk01h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342812727] HOST ALERT: omnmdk03h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342812727] HOST ALERT: omnmdk04h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1342812727] HOST ALERT: omnmdk06h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1342812727] HOST ALERT: omnmdk02h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342812727] HOST ALERT: omnmdk05h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 4.30 ms
[1342812727] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812727] HOST ALERT: nycnn-kickstart;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342812727] HOST ALERT: nycnn-mamdb1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812727] HOST ALERT: nycnn-mamdb2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812727] HOST ALERT: nycnn-mamfog;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342812727] HOST ALERT: nycnn-mamkey1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812727] HOST ALERT: nycnn-mamkey2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812727] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812727] HOST ALERT: nycnn-ardb2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812727] HOST ALERT: wdcps003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.47 ms
[1342812727] HOST ALERT: nycnn-adbeame3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812727] HOST ALERT: nycnn-arvtrmops;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812727] HOST ALERT: nycnn-arxlator1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812727] HOST ALERT: nycnn-arvtrtape;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812727] HOST ALERT: nycnn-arxlator2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342812727] HOST ALERT: nycnn-arxlator3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812727] HOST ALERT: nycnn-arxlator4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342812727] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812727] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812727] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812727] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812727] HOST ALERT: omnmdk16-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342812727] SERVICE ALERT: nyvideodrop;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nyvideodrop;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: omnldap2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342812737] HOST ALERT: omnldap1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342812737] HOST ALERT: nyarxlator3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342812737] HOST ALERT: nyarxlator4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342812737] HOST ALERT: nyarxlator1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nyarxlator2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812737] HOST ALERT: nycnn-adbeame10;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342812737] HOST ALERT: nycnn-adbeame1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342812737] HOST ALERT: nycnn-adbeame2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342812737] HOST ALERT: nycnn-adbeame8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342812737] HOST ALERT: nycnn-adbeame9;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342812737] HOST ALERT: nycnn-adbeame7;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342812737] HOST ALERT: nycnn-adbectrl1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342812737] HOST ALERT: nycnn-adbectrl2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812737] HOST ALERT: nycnn-arapp1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812737] HOST ALERT: nycnn-ipvproxy1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342812737] HOST ALERT: nycnn-ipvproxy2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812737] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvsphn2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvapilic;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: nycnn-ipvsphn4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: nycnn-ipvsphn5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342812737] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: omnmdir03e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812737] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: nycnn-ipvx10;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: omnmdir04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: omnmdk10-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342812737] HOST ALERT: omnmdk11-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342812737] HOST ALERT: omnmdk09-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342812737] HOST ALERT: omnmdk14-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: nycnn-ipvx15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: omnmdk12-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342812737] HOST ALERT: omnmdk15-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: omnmdk13-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: omnmdk18-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342812737] HOST ALERT: omnmdk19-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1342812737] HOST ALERT: omnprox02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: omnmdk17-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342812737] HOST ALERT: omnprox01e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: omnprox02e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: omnprox01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: omnprox03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: omnprox03e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: omnprox04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: omnprox04e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: omnprox05e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: omnprox05-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: omnmgsm-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: omnmdk19h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: omnmdk18h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: omnmdk10h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: omnmdk16h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342812737] HOST ALERT: omnmdk14h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1342812737] HOST ALERT: omnmdk11h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342812737] HOST ALERT: omnmdk13h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: omnmdk12h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342812737] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342812737] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ms2fsd;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: nycnn-ms2fsd2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1342812737] HOST ALERT: omnmdir03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: omnmdk08-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812737] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx19;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx28;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvx2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: nycnn-gridpushpull1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812737] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812737] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: nycnn-ipvx25;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342812737] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: nycnn-ipvx31;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342812737] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342812737] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342812737] HOST ALERT: nycnn-ms2app1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342812737] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342812737] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342812737] HOST ALERT: nycnn-gridpushpull2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342812737] HOST ALERT: omnmdk17h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.48 ms
[1342812737] HOST ALERT: omnmdk08h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1342812737] HOST ALERT: omnmdk15h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342812737] HOST ALERT: omnmdk07h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: omnmdk09h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342812737] HOST ALERT: nycnn-ms2fsd1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342812737] HOST ALERT: wdcps002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.53 ms
[1342812737] HOST ALERT: wdcod01;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.51 ms
[1342812737] HOST ALERT: wdcod02;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.50 ms
[1342812737] HOST ALERT: wdcod002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.51 ms
[1342812737] HOST ALERT: wdcms001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1342812737] HOST ALERT: wdcpf002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1342812737] HOST ALERT: wdcpf001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1342812737] HOST ALERT: wdcps004;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342812737] HOST ALERT: wdcpf03;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.36 ms
[1342812737] HOST ALERT: wdcod001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.73 ms
[1342812737] HOST ALERT: wdcms003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342812737] HOST ALERT: wdcms002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1342812737] HOST ALERT: wdcps001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1342812737] HOST FLAPPING ALERT: wdcps001;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1342812747] SERVICE ALERT: omnldap2-ny;LDAP;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1342812747] SERVICE ALERT: omnldap1-ny;LDAP;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1342812747] SERVICE ALERT: msny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nyarxlator3;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nyarxlator4;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nyarxlator1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nyarxlator2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame10;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame5;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame8;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame9;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame7;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame4;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbectrl1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbectrl2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-arapp1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvproxy1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-arxlator1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvproxy2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvsphn2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvsphn3;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-adbeame6;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdir03e1-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx18e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk10-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk11-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx18;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk09-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk14-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk12-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk15-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk13-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk16-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk18-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk19-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox02-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk17-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox01e1-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox02e1-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox01-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox03-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox03e1-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox04-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox04e1-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox05e1-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnprox05-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk19h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk18h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk10h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk16h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk14h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk11h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk13h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk12h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1342812747] SERVICE ALERT: nycnn-ipvx26;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-mamdb2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-mamdb1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-mamkey1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-mamkey2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ms2app2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdir03-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk08-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-gridpushpull1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ms2app1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx26e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-gridpushpull2;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk17h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk08h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk15h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk07h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: omnmdk09h0-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcps002;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcod01;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcod02;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcod002;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcpf002;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcps003;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcpf001;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcps004;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcpf03;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcod001;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcms002;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812747] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812757] SERVICE ALERT: nycnn-ms2app4;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812757] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812757] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1342812767] SERVICE ALERT: nycnn-mamfog;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812767] SERVICE ALERT: nycnn-kickstart;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1342812777] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: nycnn-splunk;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: nymamkey;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: nyproxygenhd;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: nyhdscheduler;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnclda01e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnclda01-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnclda02-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnclda03-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba05-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba02-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba08-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba07-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba11-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnmas1app-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1342812777] HOST ALERT: spider;DOWN;SOFT;2;(Host Check Timed Out)
[1342812787] HOST ALERT: nymq;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: rt-nycc3750tgr-ilo;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd003;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd010;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd015;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd011;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd026;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd013;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd028;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd027;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1342812787] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1342812787] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas1db1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas1db-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas2app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas1db2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas2app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas2db-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas2db1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas2db2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdir01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdir01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdir02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdir02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk05-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk06-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk07-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk01h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk03h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk04h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk06h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk02h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] SERVICE ALERT: omnmdk05h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342812787] HOST ALERT: wdchd012;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd014;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd016;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd017;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: wdchd024;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: nyvideodrop;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: omnmas1db1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: omnmas1db-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812787] HOST ALERT: omnmas2app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmas2app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmas2db-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmas2db1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmas2db2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdir01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdir01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdir02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdir02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdir04e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdk01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdk04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdk02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdk03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdk05-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdk06-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] HOST ALERT: omnmdk07-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812797] SERVICE ALERT: nyvideodrop;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1342812807] HOST ALERT: wdchd002;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd005;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd006;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd007;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd008;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd020;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd019;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd030;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdchd031;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdcmd002;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: wdcmd001;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: omnmdk01h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: omnmdk03h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: omnmdk04h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: omnmdk02h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342812807] HOST ALERT: omnmdk05h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342813019] HOST ALERT: wdchd004;DOWN;SOFT;1;(Host Check Timed Out)
[1342813019] HOST ALERT: wdchd009;DOWN;SOFT;1;(Host Check Timed Out)
[1342813019] HOST ALERT: wdchd021;DOWN;SOFT;1;(Host Check Timed Out)
[1342813019] HOST ALERT: wdchd018;DOWN;SOFT;1;(Host Check Timed Out)
[1342813019] HOST ALERT: wdchd022;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: msny;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: omnmdk06h0-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: wdchd001;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: wdchd029;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: wdchd023;DOWN;SOFT;1;(Host Check Timed Out)
[1342813021] HOST ALERT: wdchd025;DOWN;SOFT;1;(Host Check Timed Out)
[1342813025] SERVICE ALERT: msny;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.002 second response time
[1342813025] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;SOFT;2;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1342813025] SERVICE ALERT: omnldap2-ny;LDAP;OK;HARD;3;LDAP OK - 0.002 seconds response time
[1342813025] SERVICE ALERT: omnldap1-ny;LDAP;OK;HARD;3;LDAP OK - 0.002 seconds response time
[1342813026] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1342813027] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;3;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1342813035] SERVICE ALERT: nycnn-mamfog;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813035] SERVICE ALERT: nycnn-splunk;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1342813035] SERVICE ALERT: nycnn-kickstart;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813035] SERVICE ALERT: spider;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813035] SERVICE ALERT: nycnn-ipvx3;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813035] SERVICE ALERT: omnmas1db1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813035] SERVICE ALERT: omnmas1db-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813035] SERVICE ALERT: omnmas2app1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1342813035] SERVICE ALERT: omnmas1db2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813035] SERVICE ALERT: omnmas2app2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342813035] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813035] SERVICE ALERT: omnmas2db-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813035] SERVICE ALERT: omnmas2db1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813035] SERVICE ALERT: omnmas2db2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813035] SERVICE ALERT: omnmdir01-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813035] SERVICE ALERT: omnmdir01e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813035] SERVICE ALERT: omnmdir02-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813035] SERVICE ALERT: omnmdir02e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813035] SERVICE ALERT: omnmdir04e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813035] SERVICE ALERT: omnmdk01-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813035] SERVICE ALERT: omnmdk04-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1342813035] SERVICE ALERT: omnmdk02-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1342813035] SERVICE ALERT: omnmdk03-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342813035] SERVICE ALERT: omnmdk05-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1342813035] SERVICE ALERT: omnmdk06-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813035] SERVICE ALERT: omnmas1app2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813035] SERVICE ALERT: omnmdk07-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1342813035] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813035] SERVICE ALERT: omnmdk01h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1342813035] SERVICE ALERT: omnmdk03h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1342813035] SERVICE ALERT: omnmdk04h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813035] SERVICE ALERT: omnmdk06h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1342813035] SERVICE ALERT: omnmdk02h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813035] SERVICE ALERT: omnmdk05h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1342813035] HOST ALERT: nycnn-ms2db2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813035] HOST ALERT: nycnn-ms2db3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813035] HOST ALERT: nycnn-splunk;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813035] HOST ALERT: nycnn-ms2db1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813035] HOST ALERT: nymamkey;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342813035] HOST ALERT: nyproxygenhd;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1342813035] HOST ALERT: nyhdscheduler;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813035] HOST ALERT: omnclda01e1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813035] HOST ALERT: omnclda01-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813035] HOST ALERT: omnclda02-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813035] HOST ALERT: omnclda02e1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813035] HOST ALERT: omnclda03e1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1342813035] HOST ALERT: omnclda03-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1342813035] HOST ALERT: omnhbcba05-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342813035] HOST ALERT: omnhbcba03-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813035] HOST ALERT: omnhbcba01-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1342813035] HOST ALERT: omnhbcba04-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1342813035] HOST ALERT: omnhbcba02-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1342813035] HOST ALERT: omnhbcba08-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813035] HOST ALERT: omnhbcba07-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813035] HOST ALERT: omnhbcba06-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813035] HOST ALERT: omnhbcba09-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342813035] HOST ALERT: omnhbcba11-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342813035] HOST ALERT: omnhbcba10-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813035] HOST ALERT: omnmas1app-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813035] HOST ALERT: omnmas1app1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813035] HOST ALERT: omnhbcba12-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813035] HOST ALERT: spider;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1342813035] SERVICE ALERT: nyvideodrop;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813035] HOST ALERT: nymq;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1342813035] HOST ALERT: rt-nycc3750tgr-ilo;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.61 ms
[1342813035] HOST ALERT: wdchd003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.48 ms
[1342813035] HOST ALERT: wdchd010;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.70 ms
[1342813035] HOST ALERT: wdchd015;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.70 ms
[1342813035] HOST ALERT: wdchd011;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.72 ms
[1342813035] HOST ALERT: wdchd026;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1342813035] HOST ALERT: wdchd013;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.54 ms
[1342813035] HOST ALERT: wdchd028;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.71 ms
[1342813035] HOST ALERT: wdchd027;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.70 ms
[1342813035] HOST ALERT: wdchd012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.74 ms
[1342813035] HOST ALERT: wdchd014;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1342813035] HOST ALERT: wdchd016;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342813035] HOST ALERT: wdchd017;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.48 ms
[1342813035] HOST ALERT: nyvideodrop;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1342813035] HOST FLAPPING ALERT: nyvideodrop;STARTED; Host appears to have started flapping (23.8% change > 20.0% threshold)
[1342813035] HOST ALERT: wdchd024;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1342813035] HOST ALERT: omnmas1db1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813035] HOST FLAPPING ALERT: omnmas1db1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813035] HOST ALERT: omnmas1db-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813035] HOST FLAPPING ALERT: omnmas1db-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813035] HOST ALERT: omnmas2app1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813035] HOST FLAPPING ALERT: omnmas2app1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813035] SERVICE ALERT: msny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813035] SERVICE ALERT: nyarxlator3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813036] SERVICE ALERT: nyarxlator4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1342813036] SERVICE ALERT: nyarxlator1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813036] SERVICE ALERT: nyarxlator2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1342813037] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1342813037] SERVICE ALERT: nycnn-adbeame1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1342813038] SERVICE ALERT: nycnn-adbeame3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813039] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813039] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813040] SERVICE ALERT: nycnn-adbeame8;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813040] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813041] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813042] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342813042] SERVICE ALERT: nycnn-adbectrl1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813043] SERVICE ALERT: nycnn-adbectrl2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813044] SERVICE ALERT: nycnn-arapp1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813044] HOST ALERT: omnmas1db2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342813044] HOST FLAPPING ALERT: omnmas1db2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmas2app2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813044] HOST FLAPPING ALERT: omnmas2app2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813044] HOST FLAPPING ALERT: nycnn-ipvx2e1;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmas2db-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813044] HOST FLAPPING ALERT: omnmas2db-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmas2db1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813044] HOST FLAPPING ALERT: omnmas2db1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmas2db2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813044] HOST FLAPPING ALERT: omnmas2db2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdir01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813044] HOST FLAPPING ALERT: omnmdir01-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdir01e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813044] HOST FLAPPING ALERT: omnmdir01e1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdir02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813044] HOST FLAPPING ALERT: omnmdir02-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdir02e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813044] HOST FLAPPING ALERT: omnmdir02e1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdir04e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813044] HOST FLAPPING ALERT: omnmdir04e1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdk01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813044] HOST FLAPPING ALERT: omnmdk01-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdk04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1342813044] HOST FLAPPING ALERT: omnmdk04-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdk02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813044] HOST FLAPPING ALERT: omnmdk02-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdk03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1342813044] HOST FLAPPING ALERT: omnmdk03-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdk05-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342813044] HOST FLAPPING ALERT: omnmdk05-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdk06-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813044] HOST FLAPPING ALERT: omnmdk06-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmas1app2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813044] HOST FLAPPING ALERT: omnmas1app2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] HOST ALERT: omnmdk07-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1342813044] HOST FLAPPING ALERT: omnmdk07-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813044] SERVICE ALERT: nycnn-ipvproxy1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813045] SERVICE ALERT: nycnn-arvtrtape;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1342813045] SERVICE ALERT: nycnn-arxlator1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813045] SERVICE ALERT: nycnn-ipvproxy2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813046] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813047] SERVICE ALERT: nycnn-ipvsphn2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813047] SERVICE ALERT: nycnn-ipvsphn3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813048] SERVICE ALERT: nycnn-adbeame6;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813048] SERVICE ALERT: nycnn-ipvapilic;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813049] SERVICE ALERT: nycnn-ipvsphn4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813050] SERVICE ALERT: nycnn-ipvsphn5;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813050] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813051] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813051] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813052] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813053] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1342813053] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813054] HOST ALERT: wdchd002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.73 ms
[1342813054] HOST ALERT: wdchd005;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.68 ms
[1342813054] HOST ALERT: wdchd006;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.68 ms
[1342813054] HOST ALERT: wdchd007;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.47 ms
[1342813054] HOST ALERT: wdchd008;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342813054] HOST ALERT: wdchd020;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.76 ms
[1342813054] HOST ALERT: wdchd019;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.87 ms
[1342813054] HOST ALERT: wdchd030;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.80 ms
[1342813054] HOST ALERT: wdchd031;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.77 ms
[1342813054] HOST ALERT: wdcmd002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.72 ms
[1342813054] HOST ALERT: wdcmd001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1342813054] HOST ALERT: omnmdk01h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1342813054] HOST FLAPPING ALERT: omnmdk01h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813054] HOST ALERT: omnmdk03h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1342813054] HOST FLAPPING ALERT: omnmdk03h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813054] HOST ALERT: omnmdk04h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1342813054] HOST FLAPPING ALERT: omnmdk04h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813054] HOST ALERT: omnmdk02h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813054] HOST FLAPPING ALERT: omnmdk02h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813054] HOST ALERT: omnmdk05h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.34 ms
[1342813054] HOST FLAPPING ALERT: omnmdk05h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813054] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813054] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813055] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813056] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813056] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813057] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813058] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813058] SERVICE ALERT: omnmdir03e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813059] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1342813090] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813093] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342813093] SERVICE ALERT: omnmdir04-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813094] SERVICE ALERT: nycnn-ipvx18e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813095] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1342813095] SERVICE ALERT: omnmdk10-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1342813096] SERVICE ALERT: omnmdk11-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1342813096] SERVICE ALERT: nycnn-ipvx18;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813097] SERVICE ALERT: omnmdk09-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1342813098] SERVICE ALERT: omnmdk14-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813098] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813099] SERVICE ALERT: omnmdk12-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1342813099] SERVICE ALERT: omnmdk15-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813100] SERVICE ALERT: omnmdk13-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1342813101] SERVICE ALERT: omnmdk16-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1342813101] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813102] SERVICE ALERT: omnmdk18-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1342813102] SERVICE ALERT: omnmdk19-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1342813103] SERVICE ALERT: omnprox02-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813104] SERVICE ALERT: omnmdk17-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813104] SERVICE ALERT: omnprox01e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813105] SERVICE ALERT: omnprox02e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813106] SERVICE ALERT: omnprox01-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813106] SERVICE ALERT: omnprox03-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813107] SERVICE ALERT: omnprox03e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813107] SERVICE ALERT: omnprox04-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813108] SERVICE ALERT: omnprox04e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813109] SERVICE ALERT: omnprox05e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813109] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813110] SERVICE ALERT: omnprox05-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813110] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1342813111] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813112] SERVICE ALERT: omnmdk19h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.29 ms
[1342813115] SERVICE ALERT: omnmdk18h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1342813116] SERVICE ALERT: omnmdk10h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1342813116] SERVICE ALERT: omnmdk16h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1342813117] SERVICE ALERT: omnmdk14h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1342813118] SERVICE ALERT: omnmdk11h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1342813118] SERVICE ALERT: omnmdk13h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1342813119] SERVICE ALERT: omnmdk12h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813119] SERVICE ALERT: nycnn-ipvx26;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813120] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813121] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813121] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813122] SERVICE ALERT: nycnn-mamdb2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813123] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1342813123] SERVICE ALERT: nycnn-mamdb1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813124] SERVICE ALERT: nycnn-mamkey1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813124] SERVICE ALERT: nycnn-mamkey2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1342813125] SERVICE ALERT: nycnn-ms2fsd;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1342813126] SERVICE ALERT: nycnn-ms2app2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1342813126] SERVICE ALERT: nycnn-ms2fsd2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813127] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813127] SERVICE ALERT: omnmdir03-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813128] SERVICE ALERT: omnmdk08-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813129] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813129] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813130] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813130] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813131] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813132] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813132] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1342813133] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813133] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813134] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813135] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1342813135] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813164] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813165] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813165] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1342813166] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813166] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1342813167] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813168] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813168] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813169] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813169] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1342813170] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1342813171] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813171] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813172] SERVICE ALERT: nycnn-gridpushpull1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813172] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1342813173] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1342813174] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813174] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813175] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813175] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813176] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1342813177] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813347] HOST FLAPPING ALERT: wdcps001;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1342813427] SERVICE ALERT: msny;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.002 second response time
[1342813427] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;3;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1342813427] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;SOFT;2;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1342813427] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1342813427] SERVICE ALERT: msny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813437] SERVICE ALERT: nyarxlator1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813437] SERVICE ALERT: nyarxlator2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813437] SERVICE ALERT: nyarxlator3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813438] SERVICE ALERT: nyarxlator4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813438] SERVICE ALERT: nycnn-adbeame1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813439] SERVICE ALERT: nycnn-adbeame10;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342813439] SERVICE ALERT: nycnn-adbeame2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813440] SERVICE ALERT: nycnn-adbeame3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813441] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813441] SERVICE ALERT: nycnn-adbeame5;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342813442] SERVICE ALERT: nycnn-adbeame6;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342813442] SERVICE ALERT: nycnn-adbeame7;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813443] SERVICE ALERT: nycnn-adbeame8;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1342813444] SERVICE ALERT: nycnn-adbeame9;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813444] SERVICE ALERT: nycnn-adbectrl1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813445] SERVICE ALERT: nycnn-adbectrl2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813445] SERVICE ALERT: nycnn-arapp1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813446] SERVICE ALERT: nycnn-arvtrtape;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813446] SERVICE ALERT: nycnn-arxlator1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813447] SERVICE ALERT: nycnn-gridpushpull1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813447] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813447] SERVICE ALERT: nycnn-ipvapilic;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813447] SERVICE ALERT: nycnn-ipvproxy1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813448] SERVICE ALERT: nycnn-ipvproxy2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813448] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813449] SERVICE ALERT: nycnn-ipvsphn2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813450] SERVICE ALERT: nycnn-ipvsphn3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813450] SERVICE ALERT: nycnn-ipvsphn4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813451] SERVICE ALERT: nycnn-ipvsphn5;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813451] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813452] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813452] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813453] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813453] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813454] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813454] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813455] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813456] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342813456] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813456] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813457] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813471] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813471] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813472] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813473] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1342813473] HOST FLAPPING ALERT: nycnn-ipvx2e1;STARTED; Host appears to have started flapping (23.9% change > 20.0% threshold)
[1342813473] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813473] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813473] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813473] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813474] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813475] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813475] SERVICE ALERT: nycnn-ipvx18;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx18e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813476] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813476] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813476] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813476] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813476] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813476] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1342813476] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813476] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813476] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx26;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1342813476] SERVICE ALERT: nycnn-ipvx26e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813476] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813476] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813476] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813476] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813476] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813476] SERVICE ALERT: nycnn-ipvx3;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813476] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813476] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813476] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813476] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813476] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813476] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813476] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813476] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813476] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813476] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813476] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813476] SERVICE ALERT: nycnn-kickstart;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813476] SERVICE ALERT: nycnn-mamdb1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813476] SERVICE ALERT: nycnn-mamdb2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813476] SERVICE ALERT: nycnn-mamfog;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342813476] SERVICE ALERT: nycnn-mamkey1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813476] SERVICE ALERT: nycnn-mamkey2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813476] SERVICE ALERT: nycnn-ms2app1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1342813476] SERVICE ALERT: nycnn-ms2app2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813476] SERVICE ALERT: nycnn-ms2app3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1342813476] SERVICE ALERT: nycnn-ms2app4;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813476] SERVICE ALERT: nycnn-ms2db1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813476] SERVICE ALERT: nycnn-ms2db2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813476] SERVICE ALERT: nycnn-ms2db3;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813476] SERVICE ALERT: nycnn-ms2fsd;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342813476] SERVICE ALERT: nycnn-ms2fsd1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813476] SERVICE ALERT: nycnn-ms2fsd2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1342813476] SERVICE ALERT: nycnn-splunk;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342813476] SERVICE ALERT: nyhdscheduler;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813476] SERVICE ALERT: nymamkey;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813476] SERVICE ALERT: nyproxygenhd;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813476] SERVICE ALERT: nyvideodrop;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813476] SERVICE ALERT: omnclda01-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813476] SERVICE ALERT: omnclda01e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813476] SERVICE ALERT: omnclda02-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813476] SERVICE ALERT: omnclda02e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813476] SERVICE ALERT: omnldap1-ny;LDAP;OK;HARD;3;LDAP OK - 0.002 seconds response time
[1342813477] SERVICE ALERT: omnldap2-ny;LDAP;OK;HARD;3;LDAP OK - 0.004 seconds response time
[1342813478] SERVICE ALERT: omnclda03-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813478] SERVICE ALERT: omnclda03e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813478] SERVICE ALERT: omnhbcba01-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813478] SERVICE ALERT: omnhbcba02-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813478] SERVICE ALERT: omnhbcba03-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813478] SERVICE ALERT: omnhbcba04-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813478] SERVICE ALERT: omnhbcba05-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813478] SERVICE ALERT: omnhbcba06-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813478] SERVICE ALERT: omnhbcba07-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813478] SERVICE ALERT: omnhbcba08-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813478] SERVICE ALERT: omnhbcba09-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813478] SERVICE ALERT: omnhbcba10-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813478] SERVICE ALERT: omnhbcba11-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813478] SERVICE ALERT: omnhbcba12-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813478] SERVICE ALERT: omnmas1app-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813481] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813481] SERVICE ALERT: omnmas1app2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813481] SERVICE ALERT: omnmas1db-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813481] SERVICE ALERT: omnmas1db1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813481] SERVICE ALERT: omnmas1db2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813481] SERVICE ALERT: omnmas2app1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813481] SERVICE ALERT: omnmas2app2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813481] SERVICE ALERT: omnmas2db-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813481] SERVICE ALERT: omnmas2db1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813481] HOST ALERT: nycnn-ms2db1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1342813481] HOST ALERT: nycnn-ms2db2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813481] HOST ALERT: nycnn-ms2db3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813481] HOST ALERT: nycnn-splunk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813481] HOST ALERT: nyhdscheduler;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813481] HOST ALERT: nymamkey;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813481] HOST ALERT: nyproxygenhd;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813481] HOST ALERT: nyvideodrop;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342813481] HOST FLAPPING ALERT: nyvideodrop;STARTED; Host appears to have started flapping (23.8% change > 20.0% threshold)
[1342813481] HOST ALERT: omnclda01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813481] HOST ALERT: omnclda01e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813481] HOST ALERT: omnclda02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1342813481] HOST ALERT: omnclda02e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813491] HOST ALERT: omnclda03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813491] HOST ALERT: omnclda03e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813491] HOST ALERT: omnhbcba01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813491] HOST ALERT: omnhbcba02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813491] HOST ALERT: omnhbcba03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813491] HOST ALERT: omnhbcba04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813491] HOST ALERT: omnhbcba05-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813491] HOST ALERT: omnhbcba06-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813491] HOST ALERT: omnhbcba07-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813491] HOST ALERT: omnhbcba08-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813491] HOST ALERT: omnhbcba09-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813491] HOST ALERT: omnhbcba10-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813491] HOST ALERT: omnhbcba11-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813491] HOST ALERT: omnhbcba12-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1342813491] HOST ALERT: omnmas1app-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813491] SERVICE ALERT: omnmas2db2-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1342813491] SERVICE ALERT: omnmdir01-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813491] SERVICE ALERT: omnmdir01e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813491] SERVICE ALERT: omnmdir02-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813491] SERVICE ALERT: omnmdir02e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813491] SERVICE ALERT: omnmdir03-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813491] SERVICE ALERT: omnmdir03e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813491] SERVICE ALERT: omnmdir04-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813492] SERVICE ALERT: omnmdir04e1-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813492] SERVICE ALERT: omnmdk01-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813492] SERVICE ALERT: omnmdk01h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813492] SERVICE ALERT: omnmdk02-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1342813492] SERVICE ALERT: omnmdk02h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342813492] SERVICE ALERT: omnmdk03-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1342813492] SERVICE ALERT: omnmdk03h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.12 ms
[1342813492] SERVICE ALERT: omnmdk04-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813492] SERVICE ALERT: omnmdk04h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813492] SERVICE ALERT: omnmdk05-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813492] SERVICE ALERT: omnmdk05h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342813492] SERVICE ALERT: omnmdk06-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342813492] SERVICE ALERT: omnmdk06h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1342813492] SERVICE ALERT: omnmdk07-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813492] SERVICE ALERT: omnmdk07h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1342813492] SERVICE ALERT: omnmdk08-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813492] SERVICE ALERT: omnmdk08h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342813492] SERVICE ALERT: omnmdk09-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342813493] SERVICE ALERT: omnmdk09h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342813493] SERVICE ALERT: omnmdk10-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1342813493] SERVICE ALERT: omnmdk10h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1342813493] SERVICE ALERT: omnmdk11-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813494] SERVICE ALERT: omnmdk11h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1342813494] SERVICE ALERT: omnmdk12-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1342813494] SERVICE ALERT: omnmdk12h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813494] SERVICE ALERT: omnmdk13-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813494] HOST ALERT: omnmas1app1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813494] HOST ALERT: omnmas1app2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813494] HOST FLAPPING ALERT: omnmas1app2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] HOST ALERT: omnmas1db-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1342813494] HOST FLAPPING ALERT: omnmas1db-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] HOST ALERT: omnmas1db1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813494] HOST FLAPPING ALERT: omnmas1db1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] HOST ALERT: omnmas1db2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1342813494] HOST FLAPPING ALERT: omnmas1db2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] HOST ALERT: omnmas2app1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813494] HOST FLAPPING ALERT: omnmas2app1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] HOST ALERT: omnmas2app2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813494] HOST FLAPPING ALERT: omnmas2app2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] HOST ALERT: omnmas2db-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1342813494] HOST FLAPPING ALERT: omnmas2db-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] HOST ALERT: omnmas2db1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813494] HOST FLAPPING ALERT: omnmas2db1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813494] SERVICE ALERT: omnmdk13h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813494] SERVICE ALERT: omnmdk14-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813494] SERVICE ALERT: omnmdk14h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813494] SERVICE ALERT: omnmdk15-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813494] SERVICE ALERT: omnmdk15h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813494] SERVICE ALERT: omnmdk16-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1342813494] SERVICE ALERT: omnmdk16h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813494] SERVICE ALERT: omnmdk17-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813494] SERVICE ALERT: omnmdk17h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.02 ms
[1342813494] SERVICE ALERT: omnmdk18-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1342813494] SERVICE ALERT: omnmdk18h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342813494] SERVICE ALERT: omnmdk19-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1342813494] SERVICE ALERT: omnmdk19h0-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813494] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813494] SERVICE ALERT: omnprox01-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813494] SERVICE ALERT: omnprox01e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813494] SERVICE ALERT: omnprox02-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813494] SERVICE ALERT: omnprox02e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813494] SERVICE ALERT: omnprox03-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342813494] SERVICE ALERT: omnprox03e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813494] SERVICE ALERT: omnprox04-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813494] SERVICE ALERT: omnprox04e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342813494] SERVICE ALERT: omnprox05-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1342813494] SERVICE ALERT: omnprox05e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813501] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1342813501] SERVICE ALERT: spider;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1342813501] HOST ALERT: nymq;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1342813501] SERVICE ALERT: wdcms001;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342813501] SERVICE ALERT: wdcms002;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.72 ms
[1342813501] SERVICE ALERT: wdcms003;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1342813501] SERVICE ALERT: wdcod001;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1342813501] SERVICE ALERT: wdcod002;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1342813501] SERVICE ALERT: wdcod01;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.45 ms
[1342813501] SERVICE ALERT: wdcod02;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.45 ms
[1342813501] SERVICE ALERT: wdcpf001;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342813501] HOST ALERT: omnmas2db2-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1342813501] HOST FLAPPING ALERT: omnmas2db2-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdir01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1342813501] HOST FLAPPING ALERT: omnmdir01-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdir01e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342813501] HOST FLAPPING ALERT: omnmdir01e1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdir02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1342813501] HOST FLAPPING ALERT: omnmdir02-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdir02e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1342813501] HOST FLAPPING ALERT: omnmdir02e1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdir04e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1342813501] HOST FLAPPING ALERT: omnmdir04e1-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1342813501] HOST FLAPPING ALERT: omnmdk01-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk01h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1342813501] HOST FLAPPING ALERT: omnmdk01h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813501] HOST FLAPPING ALERT: omnmdk02-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk02h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.09 ms
[1342813501] HOST FLAPPING ALERT: omnmdk02h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1342813501] HOST FLAPPING ALERT: omnmdk03-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk03h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1342813501] HOST FLAPPING ALERT: omnmdk03h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1342813501] HOST FLAPPING ALERT: omnmdk04-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk04h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1342813501] HOST FLAPPING ALERT: omnmdk04h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk05-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1342813501] HOST FLAPPING ALERT: omnmdk05-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk05h0-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.38 ms
[1342813501] HOST FLAPPING ALERT: omnmdk05h0-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk06-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1342813501] HOST FLAPPING ALERT: omnmdk06-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] HOST ALERT: omnmdk07-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1342813501] HOST FLAPPING ALERT: omnmdk07-ny;STARTED; Host appears to have started flapping (23.6% change > 20.0% threshold)
[1342813501] SERVICE ALERT: wdcpf03;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.36 ms
[1342813501] SERVICE ALERT: wdcpf002;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.72 ms
[1342813501] SERVICE ALERT: wdcps001;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.54 ms
[1342813501] SERVICE FLAPPING ALERT: wdcps001;PING;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1342813501] SERVICE ALERT: wdcps002;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342813501] SERVICE ALERT: wdcps003;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342813501] SERVICE ALERT: wdcps004;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.53 ms
[1342813511] HOST ALERT: spider;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1342813581] SERVICE ALERT: nyhdscheduler;HTTP;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 400 Bad Request - 480 bytes in 0.001 second response time
[1342813591] HOST ALERT: rt-nycc3750tgr-ilo;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1342813611] HOST ALERT: wdchd002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.50 ms
[1342813611] HOST ALERT: wdchd003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.76 ms
[1342813611] HOST ALERT: wdchd005;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.54 ms
[1342813611] HOST ALERT: wdchd006;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342813611] HOST ALERT: wdchd007;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.72 ms
[1342813611] HOST ALERT: wdchd008;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342813611] HOST ALERT: wdchd010;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.53 ms
[1342813611] HOST ALERT: wdchd011;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1342813621] HOST ALERT: wdchd012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342813621] HOST ALERT: wdchd013;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342813621] HOST ALERT: wdchd014;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342813621] HOST ALERT: wdchd015;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.70 ms
[1342813621] HOST ALERT: wdchd016;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.47 ms
[1342813621] HOST ALERT: wdchd017;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.49 ms
[1342813621] HOST ALERT: wdchd019;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342813621] HOST ALERT: wdchd020;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342813631] HOST ALERT: wdchd024;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342813631] HOST ALERT: wdchd026;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.43 ms
[1342813631] HOST ALERT: wdchd027;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.74 ms
[1342813631] HOST ALERT: wdchd028;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1342813631] HOST ALERT: wdchd030;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342813641] HOST ALERT: wdchd031;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1342813641] HOST ALERT: wdcmd001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342813641] HOST ALERT: wdcmd002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1342814451] HOST FLAPPING ALERT: omnmas1app2-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814451] HOST FLAPPING ALERT: omnmas1db-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814451] HOST FLAPPING ALERT: omnmas1db1-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmas1db2-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmas2app1-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmas2app2-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmas2db-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmas2db1-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmas2db2-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmdir01-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmdir01e1-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814461] HOST FLAPPING ALERT: omnmdir02-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814471] HOST FLAPPING ALERT: omnmdir02e1-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814471] HOST FLAPPING ALERT: omnmdir04e1-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814471] HOST FLAPPING ALERT: omnmdk01-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814471] HOST FLAPPING ALERT: omnmdk01h0-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814471] HOST FLAPPING ALERT: omnmdk02-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814471] HOST FLAPPING ALERT: omnmdk02h0-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814551] HOST FLAPPING ALERT: nyvideodrop;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk03-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk04-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk03h0-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk04h0-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk05-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk06-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk05h0-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814571] HOST FLAPPING ALERT: omnmdk07-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1342814581] HOST FLAPPING ALERT: wdcps001;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1342814641] SERVICE FLAPPING ALERT: wdcps001;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1342814671] HOST FLAPPING ALERT: nycnn-ipvx2e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1342828211] HOST ALERT: wdchd001;DOWN;SOFT;1;(Host Check Timed Out)
[1342828211] HOST ALERT: wdchd003;DOWN;SOFT;1;(Host Check Timed Out)
[1342828211] HOST ALERT: wdchd002;DOWN;SOFT;1;(Host Check Timed Out)
[1342828211] HOST ALERT: wdchd004;DOWN;SOFT;1;(Host Check Timed Out)
[1342828211] HOST ALERT: wdchd006;DOWN;SOFT;1;(Host Check Timed Out)
[1342828301] HOST ALERT: wdchd001;DOWN;SOFT;2;(Host Check Timed Out)
[1342828301] HOST ALERT: wdchd003;DOWN;SOFT;2;(Host Check Timed Out)
[1342828301] HOST ALERT: wdchd002;DOWN;SOFT;2;(Host Check Timed Out)
[1342828301] HOST ALERT: wdchd004;DOWN;SOFT;2;(Host Check Timed Out)
[1342828301] HOST ALERT: wdchd006;DOWN;SOFT;2;(Host Check Timed Out)
[1342828401] HOST ALERT: wdchd001;DOWN;SOFT;3;(Host Check Timed Out)
[1342828401] HOST ALERT: wdchd003;DOWN;SOFT;3;(Host Check Timed Out)
[1342828401] HOST ALERT: wdchd002;DOWN;SOFT;3;(Host Check Timed Out)
[1342828401] HOST ALERT: wdchd004;DOWN;SOFT;3;(Host Check Timed Out)
[1342828401] HOST ALERT: wdchd006;DOWN;SOFT;3;(Host Check Timed Out)
[1342828501] HOST ALERT: wdchd001;DOWN;SOFT;4;(Host Check Timed Out)
[1342828501] HOST ALERT: wdchd003;DOWN;SOFT;4;(Host Check Timed Out)
[1342828501] HOST ALERT: wdchd002;DOWN;SOFT;4;(Host Check Timed Out)
[1342828501] HOST ALERT: wdchd004;DOWN;SOFT;4;(Host Check Timed Out)
[1342828501] HOST ALERT: wdchd006;DOWN;SOFT;4;(Host Check Timed Out)
[1342828511] HOST ALERT: wdchd007;DOWN;SOFT;1;(Host Check Timed Out)
[1342828511] HOST ALERT: wdchd008;DOWN;SOFT;1;(Host Check Timed Out)
[1342828511] HOST ALERT: wdchd009;DOWN;SOFT;1;(Host Check Timed Out)
[1342828511] HOST ALERT: wdchd010;DOWN;SOFT;1;(Host Check Timed Out)
[1342828511] HOST ALERT: wdchd011;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd012;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd013;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd014;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd015;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd016;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd017;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd018;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd019;DOWN;SOFT;1;(Host Check Timed Out)
[1342828531] HOST ALERT: wdchd020;DOWN;SOFT;1;(Host Check Timed Out)
[1342828541] HOST ALERT: wdchd021;DOWN;SOFT;1;(Host Check Timed Out)
[1342828541] HOST ALERT: wdchd022;DOWN;SOFT;1;(Host Check Timed Out)
[1342828541] HOST ALERT: wdchd023;DOWN;SOFT;1;(Host Check Timed Out)
[1342828541] HOST ALERT: wdchd024;DOWN;SOFT;1;(Host Check Timed Out)
[1342828541] HOST ALERT: wdchd025;DOWN;SOFT;1;(Host Check Timed Out)
[1342828541] HOST ALERT: wdchd026;DOWN;SOFT;1;(Host Check Timed Out)
[1342828541] HOST ALERT: wdchd027;DOWN;SOFT;1;(Host Check Timed Out)
[1342828601] HOST ALERT: wdchd001;DOWN;SOFT;5;(Host Check Timed Out)
[1342828601] HOST ALERT: wdchd003;DOWN;SOFT;5;(Host Check Timed Out)
[1342828601] HOST ALERT: wdchd002;DOWN;SOFT;5;(Host Check Timed Out)
[1342828601] HOST ALERT: wdchd004;DOWN;SOFT;5;(Host Check Timed Out)
[1342828601] HOST ALERT: wdchd006;DOWN;SOFT;5;(Host Check Timed Out)
[1342828611] HOST ALERT: wdchd007;DOWN;SOFT;2;(Host Check Timed Out)
[1342828611] HOST ALERT: wdchd008;DOWN;SOFT;2;(Host Check Timed Out)
[1342828611] HOST ALERT: wdchd009;DOWN;SOFT;2;(Host Check Timed Out)
[1342828611] HOST ALERT: wdchd010;DOWN;SOFT;2;(Host Check Timed Out)
[1342828611] HOST ALERT: wdchd011;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd012;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd013;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd014;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd015;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd016;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd017;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd018;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd019;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd020;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd021;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd022;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd023;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd024;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd025;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd026;DOWN;SOFT;2;(Host Check Timed Out)
[1342828631] HOST ALERT: wdchd027;DOWN;SOFT;2;(Host Check Timed Out)
[1342828701] HOST ALERT: wdchd001;DOWN;SOFT;6;(Host Check Timed Out)
[1342828701] HOST ALERT: wdchd003;DOWN;SOFT;6;(Host Check Timed Out)
[1342828701] HOST ALERT: wdchd002;DOWN;SOFT;6;(Host Check Timed Out)
[1342828701] HOST ALERT: wdchd004;DOWN;SOFT;6;(Host Check Timed Out)
[1342828701] HOST ALERT: wdchd006;DOWN;SOFT;6;(Host Check Timed Out)
[1342828711] HOST ALERT: wdchd007;DOWN;SOFT;3;(Host Check Timed Out)
[1342828711] HOST ALERT: wdchd008;DOWN;SOFT;3;(Host Check Timed Out)
[1342828711] HOST ALERT: wdchd009;DOWN;SOFT;3;(Host Check Timed Out)
[1342828711] HOST ALERT: wdchd010;DOWN;SOFT;3;(Host Check Timed Out)
[1342828711] HOST ALERT: wdchd011;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd012;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd013;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd014;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd015;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd016;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd017;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd018;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd019;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd020;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd021;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd022;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd023;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd024;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd025;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd026;DOWN;SOFT;3;(Host Check Timed Out)
[1342828731] HOST ALERT: wdchd027;DOWN;SOFT;3;(Host Check Timed Out)
[1342828751] SERVICE ALERT: wdcps003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342828751] SERVICE ALERT: wdcps004;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342828791] HOST ALERT: wdcps003;DOWN;SOFT;1;(Host Check Timed Out)
[1342828791] HOST ALERT: wdcps004;DOWN;SOFT;1;(Host Check Timed Out)
[1342828801] HOST ALERT: wdchd001;DOWN;SOFT;7;(Host Check Timed Out)
[1342828801] HOST ALERT: wdchd003;DOWN;SOFT;7;(Host Check Timed Out)
[1342828801] HOST ALERT: wdchd002;DOWN;SOFT;7;(Host Check Timed Out)
[1342828801] HOST ALERT: wdchd004;DOWN;SOFT;7;(Host Check Timed Out)
[1342828801] HOST ALERT: wdchd006;DOWN;SOFT;7;(Host Check Timed Out)
[1342828811] HOST ALERT: wdchd007;DOWN;SOFT;4;(Host Check Timed Out)
[1342828811] HOST ALERT: wdchd008;DOWN;SOFT;4;(Host Check Timed Out)
[1342828811] HOST ALERT: wdchd009;DOWN;SOFT;4;(Host Check Timed Out)
[1342828811] HOST ALERT: wdchd010;DOWN;SOFT;4;(Host Check Timed Out)
[1342828811] HOST ALERT: wdchd011;DOWN;SOFT;4;(Host Check Timed Out)
[1342828811] SERVICE ALERT: wdcps003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342828811] SERVICE ALERT: wdcps004;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342828831] HOST ALERT: wdchd012;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd013;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd014;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd015;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd016;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd017;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd018;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd019;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd020;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd021;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd022;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd023;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd024;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd025;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd026;DOWN;SOFT;4;(Host Check Timed Out)
[1342828831] HOST ALERT: wdchd027;DOWN;SOFT;4;(Host Check Timed Out)
[1342828841] HOST ALERT: wdchd028;DOWN;SOFT;1;(Host Check Timed Out)
[1342828841] HOST ALERT: wdchd029;DOWN;SOFT;1;(Host Check Timed Out)
[1342828841] HOST ALERT: wdchd030;DOWN;SOFT;1;(Host Check Timed Out)
[1342828861] HOST ALERT: wdchd031;DOWN;SOFT;1;(Host Check Timed Out)
[1342828871] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1342828871] HOST ALERT: wdcps003;DOWN;SOFT;2;(Host Check Timed Out)
[1342828871] HOST ALERT: wdcps004;DOWN;SOFT;2;(Host Check Timed Out)
[1342828871] SERVICE ALERT: wdcps001;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1342828901] HOST ALERT: wdchd001;DOWN;SOFT;8;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd003;DOWN;SOFT;8;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd002;DOWN;SOFT;8;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd004;DOWN;SOFT;8;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd006;DOWN;SOFT;8;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd007;DOWN;SOFT;5;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd008;DOWN;SOFT;5;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd009;DOWN;SOFT;5;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd010;DOWN;SOFT;5;(Host Check Timed Out)
[1342828901] HOST ALERT: wdchd011;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd012;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd013;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd014;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd015;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd016;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd017;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd018;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd019;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd020;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd021;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd022;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd023;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd024;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd025;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd026;DOWN;SOFT;5;(Host Check Timed Out)
[1342828931] HOST ALERT: wdchd027;DOWN;SOFT;5;(Host Check Timed Out)
[1342828941] HOST ALERT: wdchd028;DOWN;SOFT;2;(Host Check Timed Out)
[1342828941] HOST ALERT: wdchd029;DOWN;SOFT;2;(Host Check Timed Out)
[1342828941] HOST ALERT: wdchd030;DOWN;SOFT;2;(Host Check Timed Out)
[1342828961] HOST ALERT: wdchd031;DOWN;SOFT;2;(Host Check Timed Out)
[1342828961] HOST ALERT: wdcps001;DOWN;SOFT;2;(Host Check Timed Out)
[1342828961] HOST ALERT: wdcps003;DOWN;SOFT;3;(Host Check Timed Out)
[1342828961] HOST ALERT: wdcps004;DOWN;SOFT;3;(Host Check Timed Out)
[1342828981] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342828981] SERVICE ALERT: wdcpf001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342828991] SERVICE ALERT: wdcpf002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342829001] HOST ALERT: wdchd001;DOWN;SOFT;9;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd003;DOWN;SOFT;9;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd002;DOWN;SOFT;9;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd004;DOWN;SOFT;9;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd006;DOWN;SOFT;9;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd007;DOWN;SOFT;6;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd008;DOWN;SOFT;6;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd009;DOWN;SOFT;6;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd010;DOWN;SOFT;6;(Host Check Timed Out)
[1342829001] HOST ALERT: wdchd011;DOWN;SOFT;6;(Host Check Timed Out)
[1342829011] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1342829011] HOST ALERT: wdcpf001;DOWN;SOFT;1;(Host Check Timed Out)
[1342829031] HOST ALERT: wdcpf002;DOWN;SOFT;1;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd012;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd013;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd014;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd015;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd016;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd017;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd018;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd019;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd020;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd021;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd022;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd023;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd024;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd025;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd026;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd027;DOWN;SOFT;6;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd028;DOWN;SOFT;3;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd029;DOWN;SOFT;3;(Host Check Timed Out)
[1342829031] HOST ALERT: wdchd030;DOWN;SOFT;3;(Host Check Timed Out)
[1342829041] SERVICE ALERT: wdcms001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829041] SERVICE ALERT: wdcms002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342829041] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342829041] SERVICE ALERT: wdcod001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342829041] SERVICE ALERT: wdcod002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342829041] SERVICE ALERT: wdcpf001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829051] SERVICE ALERT: wdcpf002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829061] HOST ALERT: wdchd031;DOWN;SOFT;3;(Host Check Timed Out)
[1342829061] HOST ALERT: wdcps001;DOWN;SOFT;3;(Host Check Timed Out)
[1342829061] HOST ALERT: wdcps003;DOWN;SOFT;4;(Host Check Timed Out)
[1342829061] HOST ALERT: wdcps004;DOWN;SOFT;4;(Host Check Timed Out)
[1342829071] HOST ALERT: wdcms002;DOWN;SOFT;1;(Host Check Timed Out)
[1342829071] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1342829071] HOST ALERT: wdcod001;DOWN;SOFT;1;(Host Check Timed Out)
[1342829081] HOST ALERT: wdcod002;DOWN;SOFT;1;(Host Check Timed Out)
[1342829091] HOST ALERT: wdchd001;DOWN;HARD;10;(Host Check Timed Out)
[1342829091] HOST ALERT: wdchd003;DOWN;HARD;10;(Host Check Timed Out)
[1342829091] HOST ALERT: wdchd002;DOWN;HARD;10;(Host Check Timed Out)
[1342829101] HOST ALERT: wdchd004;DOWN;HARD;10;(Host Check Timed Out)
[1342829101] HOST ALERT: wdchd006;DOWN;HARD;10;(Host Check Timed Out)
[1342829101] HOST ALERT: wdchd007;DOWN;SOFT;7;(Host Check Timed Out)
[1342829101] HOST ALERT: wdchd008;DOWN;SOFT;7;(Host Check Timed Out)
[1342829101] HOST ALERT: wdchd009;DOWN;SOFT;7;(Host Check Timed Out)
[1342829101] HOST ALERT: wdchd010;DOWN;SOFT;7;(Host Check Timed Out)
[1342829101] HOST ALERT: wdchd011;DOWN;SOFT;7;(Host Check Timed Out)
[1342829101] SERVICE ALERT: wdcms002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829101] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829101] SERVICE ALERT: wdcod001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829101] SERVICE ALERT: wdcod002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829101] HOST ALERT: wdcms001;DOWN;SOFT;2;(Host Check Timed Out)
[1342829101] HOST ALERT: wdcpf001;DOWN;SOFT;2;(Host Check Timed Out)
[1342829131] HOST ALERT: wdcpf002;DOWN;SOFT;2;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd012;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd013;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd014;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd015;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd016;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd017;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd018;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd019;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd020;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd021;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd022;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd023;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd024;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd025;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd026;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd027;DOWN;SOFT;7;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd028;DOWN;SOFT;4;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd029;DOWN;SOFT;4;(Host Check Timed Out)
[1342829131] HOST ALERT: wdchd030;DOWN;SOFT;4;(Host Check Timed Out)
[1342829161] HOST ALERT: wdchd031;DOWN;SOFT;4;(Host Check Timed Out)
[1342829161] HOST ALERT: wdcps001;DOWN;SOFT;4;(Host Check Timed Out)
[1342829161] HOST ALERT: wdcps003;DOWN;SOFT;5;(Host Check Timed Out)
[1342829161] HOST ALERT: wdcps004;DOWN;SOFT;5;(Host Check Timed Out)
[1342829171] HOST ALERT: wdcms003;DOWN;SOFT;2;(Host Check Timed Out)
[1342829171] HOST ALERT: wdcod001;DOWN;SOFT;2;(Host Check Timed Out)
[1342829171] HOST ALERT: wdcms002;DOWN;SOFT;2;(Host Check Timed Out)
[1342829171] HOST ALERT: wdcod002;DOWN;SOFT;2;(Host Check Timed Out)
[1342829201] HOST ALERT: wdchd007;DOWN;SOFT;8;(Host Check Timed Out)
[1342829201] HOST ALERT: wdchd008;DOWN;SOFT;8;(Host Check Timed Out)
[1342829201] HOST ALERT: wdchd009;DOWN;SOFT;8;(Host Check Timed Out)
[1342829201] HOST ALERT: wdchd010;DOWN;SOFT;8;(Host Check Timed Out)
[1342829201] HOST ALERT: wdchd011;DOWN;SOFT;8;(Host Check Timed Out)
[1342829201] HOST ALERT: wdcms001;DOWN;SOFT;3;(Host Check Timed Out)
[1342829201] HOST ALERT: wdcpf001;DOWN;SOFT;3;(Host Check Timed Out)
[1342829221] HOST ALERT: wdcpf002;DOWN;SOFT;3;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd012;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd013;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd014;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd015;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd016;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd017;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd018;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd019;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd020;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd021;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd022;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd023;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd024;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd025;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd026;DOWN;SOFT;8;(Host Check Timed Out)
[1342829221] HOST ALERT: wdchd027;DOWN;SOFT;8;(Host Check Timed Out)
[1342829231] HOST ALERT: wdchd028;DOWN;SOFT;5;(Host Check Timed Out)
[1342829231] HOST ALERT: wdchd029;DOWN;SOFT;5;(Host Check Timed Out)
[1342829231] HOST ALERT: wdchd030;DOWN;SOFT;5;(Host Check Timed Out)
[1342829261] HOST ALERT: wdchd031;DOWN;SOFT;5;(Host Check Timed Out)
[1342829261] HOST ALERT: wdcps001;DOWN;SOFT;5;(Host Check Timed Out)
[1342829261] HOST ALERT: wdcps003;DOWN;SOFT;6;(Host Check Timed Out)
[1342829261] HOST ALERT: wdcps004;DOWN;SOFT;6;(Host Check Timed Out)
[1342829271] HOST ALERT: wdcms003;DOWN;SOFT;3;(Host Check Timed Out)
[1342829271] HOST ALERT: wdcod001;DOWN;SOFT;3;(Host Check Timed Out)
[1342829271] HOST ALERT: wdcms002;DOWN;SOFT;3;(Host Check Timed Out)
[1342829271] HOST ALERT: wdcod002;DOWN;SOFT;3;(Host Check Timed Out)
[1342829301] HOST ALERT: wdchd007;DOWN;SOFT;9;(Host Check Timed Out)
[1342829301] HOST ALERT: wdchd008;DOWN;SOFT;9;(Host Check Timed Out)
[1342829301] HOST ALERT: wdchd009;DOWN;SOFT;9;(Host Check Timed Out)
[1342829301] HOST ALERT: wdchd010;DOWN;SOFT;9;(Host Check Timed Out)
[1342829301] HOST ALERT: wdchd011;DOWN;SOFT;9;(Host Check Timed Out)
[1342829301] HOST ALERT: wdcms001;DOWN;SOFT;4;(Host Check Timed Out)
[1342829301] HOST ALERT: wdcpf001;DOWN;SOFT;4;(Host Check Timed Out)
[1342829311] HOST ALERT: wdcpf002;DOWN;SOFT;4;(Host Check Timed Out)
[1342829311] HOST ALERT: wdchd012;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd013;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd014;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd015;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd016;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd017;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd018;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd019;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd020;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd021;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd022;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd023;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd024;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd025;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd026;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd027;DOWN;SOFT;9;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd028;DOWN;SOFT;6;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd029;DOWN;SOFT;6;(Host Check Timed Out)
[1342829321] HOST ALERT: wdchd030;DOWN;SOFT;6;(Host Check Timed Out)
[1342829361] HOST ALERT: wdchd031;DOWN;SOFT;6;(Host Check Timed Out)
[1342829361] HOST ALERT: wdcps001;DOWN;SOFT;6;(Host Check Timed Out)
[1342829361] HOST ALERT: wdcps003;DOWN;SOFT;7;(Host Check Timed Out)
[1342829361] HOST ALERT: wdcps004;DOWN;SOFT;7;(Host Check Timed Out)
[1342829361] HOST ALERT: wdcms003;DOWN;SOFT;4;(Host Check Timed Out)
[1342829361] HOST ALERT: wdcod001;DOWN;SOFT;4;(Host Check Timed Out)
[1342829361] HOST ALERT: wdcms002;DOWN;SOFT;4;(Host Check Timed Out)
[1342829361] HOST ALERT: wdcod002;DOWN;SOFT;4;(Host Check Timed Out)
[1342829401] HOST ALERT: wdchd007;DOWN;HARD;10;(Host Check Timed Out)
[1342829401] HOST ALERT: wdchd008;DOWN;HARD;10;(Host Check Timed Out)
[1342829401] HOST ALERT: wdchd009;DOWN;HARD;10;(Host Check Timed Out)
[1342829401] HOST ALERT: wdchd010;DOWN;HARD;10;(Host Check Timed Out)
[1342829401] HOST ALERT: wdchd011;DOWN;HARD;10;(Host Check Timed Out)
[1342829401] HOST ALERT: wdcms001;DOWN;SOFT;5;(Host Check Timed Out)
[1342829401] HOST ALERT: wdcpf001;DOWN;SOFT;5;(Host Check Timed Out)
[1342829411] HOST ALERT: wdcpf002;DOWN;SOFT;5;(Host Check Timed Out)
[1342829411] HOST ALERT: wdchd012;DOWN;HARD;10;(Host Check Timed Out)
[1342829411] SERVICE ALERT: wdcps002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342829421] HOST ALERT: wdchd013;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd014;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd015;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd016;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd017;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd018;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd019;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd020;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd021;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd022;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd023;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd024;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd025;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd026;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd027;DOWN;HARD;10;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd028;DOWN;SOFT;7;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd029;DOWN;SOFT;7;(Host Check Timed Out)
[1342829421] HOST ALERT: wdchd030;DOWN;SOFT;7;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcps002;DOWN;SOFT;1;(Host Check Timed Out)
[1342829451] HOST ALERT: wdchd031;DOWN;SOFT;7;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcps001;DOWN;SOFT;7;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcps003;DOWN;SOFT;8;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcps004;DOWN;SOFT;8;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcms003;DOWN;SOFT;5;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcod001;DOWN;SOFT;5;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcms002;DOWN;SOFT;5;(Host Check Timed Out)
[1342829451] HOST ALERT: wdcod002;DOWN;SOFT;5;(Host Check Timed Out)
[1342829471] SERVICE ALERT: wdcps002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342829481] HOST ALERT: wdcmd001;DOWN;SOFT;1;(Host Check Timed Out)
[1342829481] HOST ALERT: wdcmd002;DOWN;SOFT;1;(Host Check Timed Out)
[1342829491] HOST ALERT: wdcps002;DOWN;SOFT;2;(Host Check Timed Out)
[1342829501] HOST ALERT: wdcms001;DOWN;SOFT;6;(Host Check Timed Out)
[1342829501] HOST ALERT: wdcpf001;DOWN;SOFT;6;(Host Check Timed Out)
[1342829511] HOST ALERT: wdcpf002;DOWN;SOFT;6;(Host Check Timed Out)
[1342829521] HOST ALERT: wdchd028;DOWN;SOFT;8;(Host Check Timed Out)
[1342829521] HOST ALERT: wdchd029;DOWN;SOFT;8;(Host Check Timed Out)
[1342829521] HOST ALERT: wdchd030;DOWN;SOFT;8;(Host Check Timed Out)
[1342829551] HOST ALERT: wdchd031;DOWN;SOFT;8;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcps001;DOWN;SOFT;8;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcps003;DOWN;SOFT;9;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcps004;DOWN;SOFT;9;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcms003;DOWN;SOFT;6;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcod001;DOWN;SOFT;6;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcms002;DOWN;SOFT;6;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcod002;DOWN;SOFT;6;(Host Check Timed Out)
[1342829551] HOST ALERT: wdcmd001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342829551] HOST ALERT: wdcmd002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1342829581] HOST ALERT: wdcps002;DOWN;SOFT;3;(Host Check Timed Out)
[1342829601] HOST ALERT: wdcms001;DOWN;SOFT;7;(Host Check Timed Out)
[1342829601] HOST ALERT: wdcpf001;DOWN;SOFT;7;(Host Check Timed Out)
[1342829611] HOST ALERT: wdcpf002;DOWN;SOFT;7;(Host Check Timed Out)
[1342829621] HOST ALERT: wdchd028;DOWN;SOFT;9;(Host Check Timed Out)
[1342829621] HOST ALERT: wdchd029;DOWN;SOFT;9;(Host Check Timed Out)
[1342829621] HOST ALERT: wdchd030;DOWN;SOFT;9;(Host Check Timed Out)
[1342829641] HOST ALERT: wdchd031;DOWN;SOFT;9;(Host Check Timed Out)
[1342829641] HOST ALERT: wdcps001;DOWN;SOFT;9;(Host Check Timed Out)
[1342829641] HOST ALERT: wdcps003;DOWN;HARD;10;(Host Check Timed Out)
[1342829641] HOST ALERT: wdcps004;DOWN;HARD;10;(Host Check Timed Out)
[1342829641] HOST ALERT: wdcms003;DOWN;SOFT;7;(Host Check Timed Out)
[1342829651] HOST ALERT: wdcod001;DOWN;SOFT;7;(Host Check Timed Out)
[1342829651] HOST ALERT: wdcms002;DOWN;SOFT;7;(Host Check Timed Out)
[1342829651] HOST ALERT: wdcod002;DOWN;SOFT;7;(Host Check Timed Out)
[1342829681] HOST ALERT: wdcps002;DOWN;SOFT;4;(Host Check Timed Out)
[1342829701] HOST ALERT: wdcms001;DOWN;SOFT;8;(Host Check Timed Out)
[1342829701] HOST ALERT: wdcpf001;DOWN;SOFT;8;(Host Check Timed Out)
[1342829711] HOST ALERT: wdcpf002;DOWN;SOFT;8;(Host Check Timed Out)
[1342829721] HOST ALERT: wdchd028;DOWN;HARD;10;(Host Check Timed Out)
[1342829721] HOST ALERT: wdchd029;DOWN;HARD;10;(Host Check Timed Out)
[1342829721] HOST ALERT: wdchd030;DOWN;HARD;10;(Host Check Timed Out)
[1342829741] HOST ALERT: wdchd031;DOWN;HARD;10;(Host Check Timed Out)
[1342829741] HOST ALERT: wdcps001;DOWN;HARD;10;(Host Check Timed Out)
[1342829741] HOST ALERT: wdcms003;DOWN;SOFT;8;(Host Check Timed Out)
[1342829741] HOST ALERT: wdcod001;DOWN;SOFT;8;(Host Check Timed Out)
[1342829741] HOST ALERT: wdcms002;DOWN;SOFT;8;(Host Check Timed Out)
[1342829741] HOST ALERT: wdcod002;DOWN;SOFT;8;(Host Check Timed Out)
[1342829761] HOST ALERT: wdchd005;DOWN;SOFT;1;(Host Check Timed Out)
[1342829781] HOST ALERT: wdcps002;DOWN;SOFT;5;(Host Check Timed Out)
[1342829791] HOST ALERT: wdcms001;DOWN;SOFT;9;(Host Check Timed Out)
[1342829791] HOST ALERT: wdcpf001;DOWN;SOFT;9;(Host Check Timed Out)
[1342829801] HOST ALERT: wdcpf002;DOWN;SOFT;9;(Host Check Timed Out)
[1342829841] HOST ALERT: wdcms003;DOWN;SOFT;9;(Host Check Timed Out)
[1342829841] HOST ALERT: wdcod001;DOWN;SOFT;9;(Host Check Timed Out)
[1342829841] HOST ALERT: wdcms002;DOWN;SOFT;9;(Host Check Timed Out)
[1342829841] HOST ALERT: wdcod002;DOWN;SOFT;9;(Host Check Timed Out)
[1342829861] HOST ALERT: wdchd005;DOWN;SOFT;2;(Host Check Timed Out)
[1342829881] HOST ALERT: wdcps002;DOWN;SOFT;6;(Host Check Timed Out)
[1342829891] HOST ALERT: wdcms001;DOWN;HARD;10;(Host Check Timed Out)
[1342829891] HOST ALERT: wdcpf001;DOWN;HARD;10;(Host Check Timed Out)
[1342829891] HOST ALERT: wdcpf002;DOWN;HARD;10;(Host Check Timed Out)
[1342829931] HOST ALERT: wdcms003;DOWN;HARD;10;(Host Check Timed Out)
[1342829931] HOST ALERT: wdcod001;DOWN;HARD;10;(Host Check Timed Out)
[1342829931] HOST ALERT: wdcms002;DOWN;HARD;10;(Host Check Timed Out)
[1342829931] HOST ALERT: wdcod002;DOWN;HARD;10;(Host Check Timed Out)
[1342829951] HOST ALERT: wdchd005;DOWN;SOFT;3;(Host Check Timed Out)
[1342829981] HOST ALERT: wdcps002;DOWN;SOFT;7;(Host Check Timed Out)
[1342830051] HOST ALERT: wdchd005;DOWN;SOFT;4;(Host Check Timed Out)
[1342830071] HOST ALERT: wdcps002;DOWN;SOFT;8;(Host Check Timed Out)
[1342830151] HOST ALERT: wdchd005;DOWN;SOFT;5;(Host Check Timed Out)
[1342830171] HOST ALERT: wdcps002;DOWN;SOFT;9;(Host Check Timed Out)
[1342830251] HOST ALERT: wdchd005;DOWN;SOFT;6;(Host Check Timed Out)
[1342830261] HOST ALERT: wdcps002;DOWN;HARD;10;(Host Check Timed Out)
[1342830351] HOST ALERT: wdchd005;DOWN;SOFT;7;(Host Check Timed Out)
[1342830451] HOST ALERT: wdchd005;DOWN;SOFT;8;(Host Check Timed Out)
[1342830551] HOST ALERT: wdchd005;DOWN;SOFT;9;(Host Check Timed Out)
[1342830641] HOST ALERT: wdchd005;DOWN;HARD;10;(Host Check Timed Out)
[1342832521] SERVICE ALERT: wdcps001;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342832531] HOST ALERT: wdcps001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342832791] HOST ALERT: wdchd003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342836671] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342836701] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1342836731] SERVICE ALERT: wdcps001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1342836801] HOST ALERT: wdcps001;DOWN;SOFT;2;(Host Check Timed Out)
[1342836841] HOST ALERT: wdchd003;DOWN;SOFT;1;(Host Check Timed Out)
[1342836891] HOST ALERT: wdcps001;DOWN;SOFT;3;(Host Check Timed Out)
[1342836941] HOST ALERT: wdchd003;DOWN;SOFT;2;(Host Check Timed Out)
[1342836981] HOST ALERT: wdcps001;DOWN;SOFT;4;(Host Check Timed Out)
[1342837021] HOST ALERT: wdcod002;UP;HARD;1;PING OK - Packet loss = 79%, RTA = 6.56 ms
[1342837021] SERVICE ALERT: wdcod002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342837021] HOST ALERT: wdcod001;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 6.59 ms
[1342837021] SERVICE ALERT: wdcod001;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 28%, RTA = 6.64 ms
[1342837031] HOST ALERT: wdchd003;DOWN;SOFT;3;(Host Check Timed Out)
[1342837081] HOST ALERT: wdcps001;DOWN;SOFT;5;(Host Check Timed Out)
[1342837081] SERVICE ALERT: wdcod001;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342837131] HOST ALERT: wdchd003;DOWN;SOFT;4;(Host Check Timed Out)
[1342837181] HOST ALERT: wdcps001;DOWN;SOFT;6;(Host Check Timed Out)
[1342837221] HOST ALERT: wdchd003;DOWN;SOFT;5;(Host Check Timed Out)
[1342837281] HOST ALERT: wdcps001;DOWN;SOFT;7;(Host Check Timed Out)
[1342837321] HOST ALERT: wdchd003;DOWN;SOFT;6;(Host Check Timed Out)
[1342837381] HOST ALERT: wdcps001;DOWN;SOFT;8;(Host Check Timed Out)
[1342837381] SERVICE ALERT: wdcpf001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1342837381] SERVICE ALERT: wdcps001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1342837381] SERVICE ALERT: wdcps003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1342837381] SERVICE ALERT: wdcpf002;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 6.64 ms
[1342837391] HOST ALERT: wdcpf001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1342837391] HOST ALERT: wdcps001;UP;SOFT;9;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342837391] HOST ALERT: wdcps003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1342837391] HOST ALERT: wdcpf002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342837411] HOST ALERT: wdchd003;DOWN;SOFT;7;(Host Check Timed Out)
[1342837441] SERVICE ALERT: wdcms001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1342837441] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.52 ms
[1342837441] SERVICE ALERT: wdcms002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1342837441] SERVICE ALERT: wdcps002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342837441] SERVICE ALERT: wdcps004;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342837451] HOST ALERT: wdcms001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1342837451] HOST ALERT: wdcms003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342837451] HOST ALERT: wdcms002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342837451] HOST ALERT: wdcps002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1342837451] HOST ALERT: wdcps004;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342837501] HOST ALERT: wdchd003;DOWN;SOFT;8;(Host Check Timed Out)
[1342837601] HOST ALERT: wdchd003;DOWN;SOFT;9;(Host Check Timed Out)
[1342837701] HOST ALERT: wdchd003;DOWN;HARD;10;(Host Check Timed Out)
[1342838011] HOST ALERT: wdchd005;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1342838111] HOST ALERT: wdchd028;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342838111] HOST ALERT: wdchd029;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342838111] HOST ALERT: wdchd007;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342838111] HOST ALERT: wdchd008;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1342838111] HOST ALERT: wdchd009;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342838111] HOST ALERT: wdchd010;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.56 ms
[1342838111] HOST ALERT: wdchd011;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1342838121] HOST ALERT: wdchd031;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342838121] HOST ALERT: wdchd012;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342838181] HOST ALERT: wdchd013;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.70 ms
[1342838181] HOST ALERT: wdchd014;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1342838181] HOST ALERT: wdchd015;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.56 ms
[1342838181] HOST ALERT: wdchd016;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342838181] HOST ALERT: wdchd017;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1342838181] HOST ALERT: wdchd018;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342838181] HOST ALERT: wdchd019;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342838181] HOST ALERT: wdchd020;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.53 ms
[1342838181] HOST ALERT: wdchd021;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1342838181] HOST ALERT: wdchd022;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342838181] HOST ALERT: wdchd023;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1342838181] HOST ALERT: wdchd024;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1342838181] HOST ALERT: wdchd025;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1342838181] HOST ALERT: wdchd026;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342838201] HOST ALERT: wdchd027;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1342838201] HOST ALERT: wdchd001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342838201] HOST ALERT: wdchd002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.68 ms
[1342838201] HOST ALERT: wdchd004;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342838201] HOST ALERT: wdchd006;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1342839011] HOST ALERT: wdchd003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1342839121] HOST ALERT: wdchd030;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1342840261] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1342840321] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms