[1395985710] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 122.24 ms
[1395985770] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 82.83 ms
[1395985880] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 110.64 ms
[1395985940] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 104.92 ms
[1395985950] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 133.79 ms
[1395986000] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 112.27 ms
[1395986010] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 78.16 ms
[1395986060] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 64.85 ms
[1395986130] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 131.96 ms
[1395986190] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.13 ms
[1395986240] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 117.21 ms
[1395986300] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 100.33 ms
[1395986360] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1395996570] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 121.81 ms
[1395996630] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 92.16 ms
[1395996680] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 129.51 ms
[1395996690] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 113.86 ms
[1395996740] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 135.50 ms
[1395996750] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 78.32 ms
[1395996800] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 111.46 ms
[1395996810] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.96 ms
[1395996870] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 87.83 ms
[1395996920] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 94.57 ms
[1395996980] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 129.19 ms
[1395997040] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 89.25 ms
[1395997100] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 100.87 ms
[1395997160] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.95 ms
[1396007370] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 103.27 ms
[1396007430] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 126.66 ms
[1396007480] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 123.95 ms
[1396007490] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 116.17 ms
[1396007540] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 128.36 ms
[1396007600] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 98.75 ms
[1396007610] SERVICE ALERT: rt-nyccfcp-1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 79.22 ms
[1396007660] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 109.52 ms
[1396007670] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 104.37 ms
[1396007710] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1396007720] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 74.54 ms
[1396007730] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 135.80 ms
[1396007770] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1396007780] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 122.63 ms
[1396007790] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1396007840] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 78.42 ms
[1396008770] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1396008830] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1396009030] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1396009090] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1396010030] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;HARD;2;AME is idle:
[1396010290] SERVICE ALERT: nycnn-adbeame4;AME Job status;OK;HARD;2;AME is idle:
[1396010470] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;HARD;2;AME is idle:
[1396018230] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 114.01 ms
[1396018280] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 102.73 ms
[1396018290] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 135.76 ms
[1396018340] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 107.47 ms
[1396018350] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 79.67 ms
[1396018400] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 100.46 ms
[1396018410] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 154.13 ms
[1396018470] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 38.70 ms
[1396018530] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 140.00 ms
[1396018590] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.80 ms
[1396018640] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 61.60 ms
[1396018700] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 127.38 ms
[1396018760] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 92.19 ms
[1396024633] HOST DOWNTIME ALERT: nycnn-ipvstrm4;STARTED; Host has entered a period of scheduled downtime
[1396025210] SERVICE ALERT: nycnn-ipvstrm4;SVServer;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1396025270] SERVICE ALERT: nycnn-ipvstrm4;SVServer;CRITICAL;HARD;2;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1396025530] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1396025530] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1396025540] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1396025540] SERVICE ALERT: nycnn-ipvstrm4;MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1396025560] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1396025590] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1396025590] SERVICE ALERT: nycnn-ipvstrm4;SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1396025590] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1396025590] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396025600] SERVICE ALERT: nycnn-ipvstrm4;MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1396025600] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1396025660] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1396025710] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1396025720] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1396025870] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1396025870] SERVICE ALERT: nycnn-ipvstrm4;SVServer;CRITICAL;HARD;2;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1396025880] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1196.74GB (85.7%), Free=200.45GB (14.3%)
[1396025880] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1396025880] SERVICE ALERT: nycnn-ipvstrm4;MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1396026070] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1396026070] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1396026170] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1396026190] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026190] SERVICE ALERT: nycnn-ipvstrm4;SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1396026190] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;SVServer;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1396026190] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026200] SERVICE ALERT: nycnn-ipvstrm4;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026200] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026260] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING WARNING - Packet loss = 82%, RTA = 0.22 ms
[1396026260] HOST FLAPPING ALERT: nycnn-ipvstrm4;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1396026310] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1396026310] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;PING;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1396026470] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1396026470] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;FTP Service;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1396026470] SERVICE ALERT: nycnn-ipvstrm4;SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1396026480] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1196.74GB (85.7%), Free=200.45GB (14.3%)
[1396026480] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;D:\ Drive Space;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1396026480] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1396026480] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;SSH Service;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1396026480] SERVICE ALERT: nycnn-ipvstrm4;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1396026480] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;MMSService;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1396026670] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1396026710] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1396026730] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1396026790] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026790] SERVICE ALERT: nycnn-ipvstrm4;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026790] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026800] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026800] SERVICE ALERT: nycnn-ipvstrm4;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1396026810] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1396026880] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1396026910] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1396027070] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1396027070] SERVICE ALERT: nycnn-ipvstrm4;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1396027080] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1196.74GB (85.7%), Free=200.45GB (14.3%)
[1396027080] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1396027080] SERVICE ALERT: nycnn-ipvstrm4;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1396027930] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1396029980] HOST FLAPPING ALERT: nycnn-ipvstrm4;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1396032770] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;FTP Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1396032770] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;SVServer;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1396032780] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;D:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1396032780] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;SSH Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1396032780] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;MMSService;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1396035393] HOST DOWNTIME ALERT: nycnn-ipvstrm4;STOPPED; Host has exited from a period of scheduled downtime