[1401705856] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1401705916] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 9.84%
[1401705966] SERVICE ALERT: nycnn-ipvx18;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1401706026] SERVICE ALERT: nycnn-ipvx18;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 9.89%
[1401731436] HOST DOWNTIME ALERT: nycnn-ipvx35;STARTED; Host has entered a period of scheduled downtime
[1401731458] HOST DOWNTIME ALERT: nycnn-ipvx35e1;STARTED; Host has entered a period of scheduled downtime
[1401731956] SERVICE ALERT: nycnn-ipvx35;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1401731956] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1401731966] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1401731976] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1401731986] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1401731996] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1401731996] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1401732016] SERVICE ALERT: nycnn-ipvx35;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1401732016] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1401732026] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1401732036] SERVICE ALERT: nycnn-ipvx35;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1401732046] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1401732046] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1401732046] SERVICE ALERT: nycnn-ipvx35e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1401732046] SERVICE ALERT: nycnn-ipvx35e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1401732066] SERVICE ALERT: nycnn-ipvx35e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1401732076] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 11.85 ms
[1401732076] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.48 ms
[1401732096] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1401732096] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1401732106] SERVICE ALERT: nycnn-ipvx35e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1401732146] SERVICE ALERT: nycnn-ipvx35e1;Process: Siphon;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1401732296] SERVICE ALERT: nycnn-ipvx35;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1401732296] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1401732306] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1401732326] SERVICE ALERT: nycnn-ipvx35;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1401732336] SERVICE ALERT: nycnn-ipvx35e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1401732336] SERVICE ALERT: nycnn-ipvx35e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.822GB (41%) - Free: 1.168GB (59%)
[1401732356] SERVICE ALERT: nycnn-ipvx35e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1401734967] HOST DOWNTIME ALERT: nycnn-ipvx35;STOPPED; Host has exited from a period of scheduled downtime
[1401735048] HOST DOWNTIME ALERT: nycnn-ipvx35e1;STOPPED; Host has exited from a period of scheduled downtime
[1401739116] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job 10184715 remains in a queued state
[1401739176] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;SOFT;2;Job 10184715 progress is  0 :
[1401744336] SERVICE ALERT: nycnn-adbeame9;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports that the encode failed.
[1401744396] SERVICE ALERT: nycnn-adbeame9;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports that the encode failed.
[1401746196] SERVICE ALERT: nycnn-adbeame9;AME Job status;OK;HARD;2;AME is idle:
[1401757536] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;SOFT;1;Job 10185516 remains in a queued state
[1401757596] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;SOFT;2;AME is idle:
[1401759096] SERVICE ALERT: nycnn-adbeame9;AME Job status;CRITICAL;SOFT;1;Job 10185532 remains in a queued state
[1401759156] SERVICE ALERT: nycnn-adbeame9;AME Job status;OK;SOFT;2;AME is idle: