[1403094226] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403094266] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403094286] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403094356] HOST ALERT: omnmgsm-ny;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.28 ms
[1403094396] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1403102686] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.22 ms
[1403102716] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403102746] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403102756] HOST ALERT: omnmgsm-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403102846] HOST ALERT: omnmgsm-ny;UP;SOFT;3;PING WARNING - Packet loss = 81%, RTA = 0.41 ms
[1403102856] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1403122656] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;SOFT;1;Job 10226440 remains in a queued state
[1403122716] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;SOFT;2;Job 10226440 progress is  0 :
[1403124036] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1403124096] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1403124396] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;HARD;2;AME is idle:
[1403124946] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403124986] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403125006] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403125086] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403125116] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1403125126] HOST ALERT: omnmas1app2-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1403125906] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403125936] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403125966] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403125996] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403126096] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1403126176] HOST ALERT: omnmas1app1-ny;UP;SOFT;4;PING OK - Packet loss = 76%, RTA = 0.58 ms
[1403126196] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1403134226] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;SOFT;1;Job 10226568 remains in a queued state
[1403134276] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;SOFT;2;AME is idle:
[1403144276] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1403144286] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1403144336] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.51%
[1403144346] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.20%
[1403144456] SERVICE ALERT: nycnn-ipvsphn5;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1403144516] SERVICE ALERT: nycnn-ipvsphn5;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 3.69%