[1336452962] HOST ALERT: wdcmd001;DOWN;SOFT;1;(Host Check Timed Out)
[1336453032] HOST ALERT: wdcmd001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336454702] HOST ALERT: wdcps003;UP;HARD;1;PING OK - Packet loss = 28%, RTA = 6.66 ms
[1336454702] SERVICE ALERT: wdcps003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.91 ms
[1336454702] SERVICE ALERT: wdcps004;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1336454702] SERVICE ALERT: wdcps002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.75 ms
[1336454712] SERVICE ALERT: wdcps001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336454712] HOST ALERT: wdcps004;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1336454712] HOST ALERT: wdcps002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1336454722] HOST ALERT: wdcps001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.54 ms
[1336455192] SERVICE ALERT: wdcps004;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 6.62 ms
[1336455232] HOST ALERT: wdcps004;DOWN;SOFT;1;(Host Check Timed Out)
[1336455242] HOST ALERT: wdcps002;DOWN;SOFT;1;(Host Check Timed Out)
[1336455252] SERVICE ALERT: wdcps003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336455252] SERVICE ALERT: wdcps002;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1336455252] SERVICE ALERT: wdcps004;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336455252] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336455292] HOST ALERT: wdcps003;DOWN;SOFT;1;(Host Check Timed Out)
[1336455292] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1336455312] SERVICE ALERT: wdcps003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336455312] SERVICE ALERT: wdcps001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336455332] HOST ALERT: wdcps004;DOWN;SOFT;2;(Host Check Timed Out)
[1336455332] HOST ALERT: wdcps002;DOWN;SOFT;2;(Host Check Timed Out)
[1336455352] HOST ALERT: wdcps003;DOWN;SOFT;2;(Host Check Timed Out)
[1336455382] HOST ALERT: wdcps001;DOWN;SOFT;2;(Host Check Timed Out)
[1336455432] SERVICE ALERT: wdcod002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1336455432] SERVICE ALERT: wdcod001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.56 ms
[1336455432] HOST ALERT: wdcps004;DOWN;SOFT;3;(Host Check Timed Out)
[1336455432] HOST ALERT: wdcps002;DOWN;SOFT;3;(Host Check Timed Out)
[1336455432] HOST ALERT: wdcod001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1336455442] HOST ALERT: wdcod002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1336455452] HOST ALERT: wdcps003;DOWN;SOFT;3;(Host Check Timed Out)
[1336455482] HOST ALERT: wdcps001;DOWN;SOFT;3;(Host Check Timed Out)
[1336455532] HOST ALERT: wdcps004;DOWN;SOFT;4;(Host Check Timed Out)
[1336455532] HOST ALERT: wdcps002;DOWN;SOFT;4;(Host Check Timed Out)
[1336455552] HOST ALERT: wdcps003;DOWN;SOFT;4;(Host Check Timed Out)
[1336455552] SERVICE ALERT: wdcps001;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 6.77 ms
[1336455552] HOST ALERT: wdcps001;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336455602] HOST ALERT: wdcps004;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1336455602] HOST ALERT: wdcps002;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1336455602] SERVICE ALERT: wdcps003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.71 ms
[1336455602] SERVICE ALERT: wdcps002;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1336455602] SERVICE ALERT: wdcps004;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336455612] SERVICE ALERT: wdcps001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.73 ms
[1336455612] SERVICE FLAPPING ALERT: wdcps001;PING;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1336455612] HOST ALERT: wdcps003;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1336456442] HOST ALERT: wdcpf001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336456442] SERVICE ALERT: wdcms001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336456442] SERVICE ALERT: wdcms002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.69 ms
[1336456442] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1336456452] SERVICE ALERT: wdcpf002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336456452] SERVICE ALERT: wdcpf001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336456452] HOST ALERT: wdcms001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1336456452] HOST ALERT: wdcms002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1336456452] HOST ALERT: wdcms003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1336456462] HOST ALERT: wdcpf002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1336456752] SERVICE FLAPPING ALERT: wdcps001;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1336456922] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336456932] SERVICE ALERT: wdcms001;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1336456932] SERVICE ALERT: wdcms002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336456932] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336456962] HOST ALERT: wdcms002;DOWN;SOFT;1;(Host Check Timed Out)
[1336456962] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1336456992] SERVICE ALERT: wdcms002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336456992] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336457022] HOST ALERT: wdcms001;DOWN;SOFT;2;(Host Check Timed Out)
[1336457062] HOST ALERT: wdcms002;DOWN;SOFT;2;(Host Check Timed Out)
[1336457062] HOST ALERT: wdcms003;DOWN;SOFT;2;(Host Check Timed Out)
[1336457122] HOST ALERT: wdcms001;DOWN;SOFT;3;(Host Check Timed Out)
[1336457152] HOST ALERT: wdcms002;DOWN;SOFT;3;(Host Check Timed Out)
[1336457152] HOST ALERT: wdcms003;DOWN;SOFT;3;(Host Check Timed Out)
[1336457222] HOST ALERT: wdcms001;DOWN;SOFT;4;(Host Check Timed Out)
[1336457242] HOST ALERT: wdcms002;DOWN;SOFT;4;(Host Check Timed Out)
[1336457242] HOST ALERT: wdcms003;DOWN;SOFT;4;(Host Check Timed Out)
[1336457322] HOST ALERT: wdcms001;DOWN;SOFT;5;(Host Check Timed Out)
[1336457342] HOST ALERT: wdcms002;DOWN;SOFT;5;(Host Check Timed Out)
[1336457342] HOST ALERT: wdcms003;DOWN;SOFT;5;(Host Check Timed Out)
[1336457422] HOST ALERT: wdcms001;DOWN;SOFT;6;(Host Check Timed Out)
[1336457442] HOST ALERT: wdcms002;DOWN;SOFT;6;(Host Check Timed Out)
[1336457442] HOST ALERT: wdcms003;DOWN;SOFT;6;(Host Check Timed Out)
[1336457522] HOST ALERT: wdcms001;DOWN;SOFT;7;(Host Check Timed Out)
[1336457542] HOST ALERT: wdcms002;DOWN;SOFT;7;(Host Check Timed Out)
[1336457542] HOST ALERT: wdcms003;DOWN;SOFT;7;(Host Check Timed Out)
[1336457622] HOST ALERT: wdcms001;DOWN;SOFT;8;(Host Check Timed Out)
[1336457642] HOST ALERT: wdcms002;DOWN;SOFT;8;(Host Check Timed Out)
[1336457642] HOST ALERT: wdcms003;DOWN;SOFT;8;(Host Check Timed Out)
[1336457722] HOST ALERT: wdcms001;DOWN;SOFT;9;(Host Check Timed Out)
[1336457732] HOST ALERT: wdcms002;DOWN;SOFT;9;(Host Check Timed Out)
[1336457732] HOST ALERT: wdcms003;DOWN;SOFT;9;(Host Check Timed Out)
[1336457822] HOST ALERT: wdcms001;DOWN;HARD;10;(Host Check Timed Out)
[1336457822] HOST ALERT: wdcms002;DOWN;HARD;10;(Host Check Timed Out)
[1336457822] HOST ALERT: wdcms003;DOWN;HARD;10;(Host Check Timed Out)
[1336458542] SERVICE ALERT: wdcms001;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336458552] HOST ALERT: wdcms001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336459992] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336460022] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336460052] SERVICE ALERT: wdcms001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336460082] HOST ALERT: wdcms001;DOWN;SOFT;2;(Host Check Timed Out)
[1336460092] HOST ALERT: wdcmd002;DOWN;SOFT;1;(Host Check Timed Out)
[1336460162] HOST ALERT: wdcmd002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336460182] HOST ALERT: wdcms001;DOWN;SOFT;3;(Host Check Timed Out)
[1336460282] HOST ALERT: wdcms001;DOWN;SOFT;4;(Host Check Timed Out)
[1336460382] HOST ALERT: wdcms001;DOWN;SOFT;5;(Host Check Timed Out)
[1336460482] HOST ALERT: wdcms001;DOWN;SOFT;6;(Host Check Timed Out)
[1336460582] HOST ALERT: wdcms001;DOWN;SOFT;7;(Host Check Timed Out)
[1336460672] HOST ALERT: wdcms001;DOWN;SOFT;8;(Host Check Timed Out)
[1336460772] HOST ALERT: wdcms001;DOWN;SOFT;9;(Host Check Timed Out)
[1336460872] HOST ALERT: wdcms001;DOWN;HARD;10;(Host Check Timed Out)
[1336461272] HOST ALERT: wdchd018;UP;HARD;1;PING OK - Packet loss = 76%, RTA = 6.54 ms
[1336461292] HOST ALERT: wdchd019;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336461292] HOST ALERT: wdchd020;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1336461302] HOST ALERT: wdchd002;UP;HARD;1;PING OK - Packet loss = 66%, RTA = 7.82 ms
[1336461302] HOST ALERT: wdchd006;UP;HARD;1;PING OK - Packet loss = 70%, RTA = 6.61 ms
[1336461302] HOST ALERT: wdchd005;UP;HARD;1;PING OK - Packet loss = 72%, RTA = 6.60 ms
[1336461302] HOST ALERT: wdchd007;UP;HARD;1;PING OK - Packet loss = 73%, RTA = 7.12 ms
[1336461312] HOST ALERT: wdchd008;UP;HARD;1;PING OK - Packet loss = 76%, RTA = 6.87 ms
[1336461342] HOST ALERT: wdchd022;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.53 ms
[1336461342] HOST ALERT: wdchd021;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1336461342] HOST ALERT: wdchd009;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1336461362] HOST ALERT: wdchd023;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1336461362] HOST ALERT: wdchd013;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336461362] HOST ALERT: wdchd025;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336461382] HOST ALERT: wdchd024;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336461442] HOST ALERT: wdchd026;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1336461442] HOST ALERT: wdchd010;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.69 ms
[1336461442] HOST ALERT: wdchd015;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1336461442] HOST ALERT: wdchd028;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1336461472] HOST ALERT: wdchd011;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336461472] HOST ALERT: wdchd016;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.68 ms
[1336461472] HOST ALERT: wdchd003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.56 ms
[1336461472] HOST ALERT: wdchd014;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1336461472] HOST ALERT: wdchd017;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1336461472] HOST ALERT: wdchd012;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.54 ms
[1336461512] HOST ALERT: wdchd029;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336461552] HOST ALERT: wdchd004;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.69 ms
[1336461552] HOST ALERT: wdchd027;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1336461592] HOST ALERT: wdchd001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.54 ms
[1336462142] SERVICE ALERT: wdcpf03;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.36 ms
[1336462152] HOST ALERT: wdcpf03;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.35 ms
[1336462192] HOST ALERT: wdchd030;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.68 ms
[1336462192] HOST ALERT: wdchd031;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336463342] SERVICE ALERT: wdcms001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.52 ms
[1336463342] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1336463342] SERVICE ALERT: wdcms002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1336463352] HOST ALERT: wdcms001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336463352] HOST ALERT: wdcms003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336463352] HOST ALERT: wdcms002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336476492] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336476522] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336476542] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.67 ms
[1336476552] HOST ALERT: wdcms001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336483072] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336483092] HOST ALERT: wdcms001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1336483692] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336483692] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336483692] SERVICE ALERT: wdcms002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336483732] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336483732] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1336483732] HOST ALERT: wdcms002;DOWN;SOFT;1;(Host Check Timed Out)
[1336483752] SERVICE ALERT: wdcms001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336483752] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336483752] SERVICE ALERT: wdcms002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336483782] HOST ALERT: wdcms001;DOWN;SOFT;2;(Host Check Timed Out)
[1336483822] HOST ALERT: wdcms003;DOWN;SOFT;2;(Host Check Timed Out)
[1336483822] HOST ALERT: wdcms002;DOWN;SOFT;2;(Host Check Timed Out)
[1336483882] HOST ALERT: wdcms001;DOWN;SOFT;3;(Host Check Timed Out)
[1336483922] HOST ALERT: wdcms003;DOWN;SOFT;3;(Host Check Timed Out)
[1336483922] HOST ALERT: wdcms002;DOWN;SOFT;3;(Host Check Timed Out)
[1336483982] HOST ALERT: wdcms001;DOWN;SOFT;4;(Host Check Timed Out)
[1336484022] HOST ALERT: wdcms003;DOWN;SOFT;4;(Host Check Timed Out)
[1336484022] HOST ALERT: wdcms002;DOWN;SOFT;4;(Host Check Timed Out)
[1336484072] HOST ALERT: wdcms001;DOWN;SOFT;5;(Host Check Timed Out)
[1336484122] HOST ALERT: wdcms003;DOWN;SOFT;5;(Host Check Timed Out)
[1336484122] HOST ALERT: wdcms002;DOWN;SOFT;5;(Host Check Timed Out)
[1336484172] HOST ALERT: wdcms001;DOWN;SOFT;6;(Host Check Timed Out)
[1336484222] HOST ALERT: wdcms003;DOWN;SOFT;6;(Host Check Timed Out)
[1336484222] HOST ALERT: wdcms002;DOWN;SOFT;6;(Host Check Timed Out)
[1336484262] HOST ALERT: wdcms001;DOWN;SOFT;7;(Host Check Timed Out)
[1336484322] HOST ALERT: wdcms003;DOWN;SOFT;7;(Host Check Timed Out)
[1336484322] HOST ALERT: wdcms002;DOWN;SOFT;7;(Host Check Timed Out)
[1336484362] HOST ALERT: wdcms001;DOWN;SOFT;8;(Host Check Timed Out)
[1336484422] HOST ALERT: wdcms003;DOWN;SOFT;8;(Host Check Timed Out)
[1336484422] HOST ALERT: wdcms002;DOWN;SOFT;8;(Host Check Timed Out)
[1336484462] HOST ALERT: wdcms001;DOWN;SOFT;9;(Host Check Timed Out)
[1336484522] HOST ALERT: wdcms003;DOWN;SOFT;9;(Host Check Timed Out)
[1336484522] HOST ALERT: wdcms002;DOWN;SOFT;9;(Host Check Timed Out)
[1336484562] HOST ALERT: wdcms001;DOWN;HARD;10;(Host Check Timed Out)
[1336484622] HOST ALERT: wdcms003;DOWN;HARD;10;(Host Check Timed Out)
[1336484622] HOST ALERT: wdcms002;DOWN;HARD;10;(Host Check Timed Out)
[1336485252] HOST ALERT: wdchd013;DOWN;SOFT;1;(Host Check Timed Out)
[1336485322] HOST ALERT: wdchd013;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336485602] SERVICE ALERT: wdcms001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1336485612] HOST ALERT: wdcms001;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.68 ms
[1336485722] SERVICE ALERT: wdcms002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1336485722] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.53 ms
[1336485732] HOST ALERT: wdcms002;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.52 ms
[1336485732] HOST ALERT: wdcms003;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336486072] HOST ALERT: wdchd004;DOWN;SOFT;1;(Host Check Timed Out)
[1336486132] HOST ALERT: wdchd005;DOWN;SOFT;1;(Host Check Timed Out)
[1336486142] HOST ALERT: wdchd004;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.64 ms
[1336486222] HOST ALERT: wdchd005;DOWN;SOFT;2;(Host Check Timed Out)
[1336486322] HOST ALERT: wdchd005;DOWN;SOFT;3;(Host Check Timed Out)
[1336486332] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336486382] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1336486422] HOST ALERT: wdchd005;DOWN;SOFT;4;(Host Check Timed Out)
[1336486522] HOST ALERT: wdchd005;DOWN;SOFT;5;(Host Check Timed Out)
[1336486622] HOST ALERT: wdchd005;DOWN;SOFT;6;(Host Check Timed Out)
[1336486722] HOST ALERT: wdchd005;DOWN;SOFT;7;(Host Check Timed Out)
[1336486822] HOST ALERT: wdchd005;DOWN;SOFT;8;(Host Check Timed Out)
[1336486922] HOST ALERT: wdchd005;DOWN;SOFT;9;(Host Check Timed Out)
[1336487012] HOST ALERT: wdchd005;DOWN;HARD;10;(Host Check Timed Out)
[1336487112] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336487152] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1336487172] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336487222] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1336487222] HOST ALERT: omnmgsm-ny;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.28 ms
[1336487832] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336487862] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336487882] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1336487892] HOST ALERT: wdcms001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1336490712] HOST ALERT: wdchd005;UP;HARD;1;PING WARNING - Packet loss = 80%, RTA = 6.82 ms
[1336496412] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336496412] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336496432] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336496452] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1336496472] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336496472] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336496472] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336496502] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336496522] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1336496522] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;2;(Host Check Timed Out)
[1336496532] HOST ALERT: wdcms001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1336496532] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.28 ms
[1336496582] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336496582] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1336496592] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1336497562] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336497562] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336497602] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336497602] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1336497622] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336497622] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336497682] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.49 ms
[1336497682] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.46 ms
[1336497682] SERVICE ALERT: nycnn-ipvx8e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 28%, RTA = 0.29 ms
[1336497682] SERVICE ALERT: nycnn-ipvx8;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 28%, RTA = 0.27 ms
[1336497732] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336497732] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1336499962] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336499962] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500002] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1336500002] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336500022] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500022] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500022] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500022] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500052] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1336500052] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336500082] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.42 ms
[1336500082] SERVICE ALERT: nycnn-ipvx6;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.30 ms
[1336500082] SERVICE ALERT: nycnn-ipvx6e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 28%, RTA = 0.22 ms
[1336500082] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.29 ms
[1336500082] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500082] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500082] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500082] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500122] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1336500122] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336500122] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1336500122] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1336500132] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500132] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500132] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1336500132] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1336500132] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1336500132] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1336500142] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500142] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500172] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336500172] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1336500192] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500192] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500192] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500192] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1336500202] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1336500202] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1336500242] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1336500242] HOST ALERT: nycnn-ipvx10;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500242] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1336500242] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500312] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500312] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500342] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1336500342] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336500372] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500372] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500372] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500372] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1336500412] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336500412] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1336500412] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500412] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500422] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1336500422] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1336500432] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500432] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500472] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500472] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500482] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1336500482] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1336500502] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500502] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336500542] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1336500542] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336500562] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500562] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336500612] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1336500612] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1336500612] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1336500612] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1336501702] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336501702] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336501742] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1336501742] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1336501762] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336501762] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336501782] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;2;(Host Check Timed Out)
[1336501782] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;2;(Host Check Timed Out)
[1336501822] SERVICE ALERT: nycnn-ipvx4;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.31 ms
[1336501822] SERVICE ALERT: nycnn-ipvx4e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 37%, RTA = 0.49 ms
[1336501852] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1336501852] HOST ALERT: nycnn-ipvx4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1336501872] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1336501872] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1336504992] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 6.61 ms
[1336505032] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1336505052] SERVICE ALERT: wdcps001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 81%, RTA = 6.62 ms
[1336505102] HOST ALERT: wdcps001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 7.50 ms
[1336505112] SERVICE ALERT: wdcps001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.60 ms
[1336516872] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336516922] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1336518792] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1336518832] HOST ALERT: wdcms001;DOWN;SOFT;1;(Host Check Timed Out)
[1336518852] SERVICE ALERT: wdcms001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1336518892] HOST ALERT: wdcms001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.71 ms
[1336518902] SERVICE ALERT: wdcms001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.64 ms