[1408771428] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1408771478] SERVICE ALERT: nycnn-ipvx23;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1408771488] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.00%
[1408771538] SERVICE ALERT: nycnn-ipvx23;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1408786078] SERVICE ALERT: nycnn-ipvx2;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1408786138] SERVICE ALERT: nycnn-ipvx2;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1408786208] SERVICE ALERT: nycnn-ipvx2e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1408786268] SERVICE ALERT: nycnn-ipvx2e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1408790038] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1408790098] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1408790098] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1408790098] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridA/Temp/.gridamounted not found or not a regular file
[1408790158] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1408790158] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;2;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1408790158] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - File /Volumes/GridA/Temp/.gridamounted not found or not a regular file
[1408790168] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1408790178] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1408790188] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1408790198] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;1;(Host Check Timed Out)
[1408790208] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;1;(Host Check Timed Out)
[1408790218] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;3;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1408790218] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - File /Volumes/GridA/Temp/.gridamounted not found or not a regular file
[1408790228] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;1;(Host Check Timed Out)
[1408790228] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1408790238] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1408790248] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1408790298] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;2;(Host Check Timed Out)
[1408790308] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;2;(Host Check Timed Out)
[1408790318] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;2;(Host Check Timed Out)
[1408790347] HOST DOWNTIME ALERT: msny;STARTED; Host has entered a period of scheduled downtime
[1408790388] HOST ALERT: nycnn-ms2fsd1;UP;SOFT;3;PING WARNING - Packet loss = 81%, RTA = 0.64 ms
[1408790398] HOST ALERT: nycnn-ms2fsd2;UP;SOFT;3;PING OK - Packet loss = 61%, RTA = 0.54 ms
[1408790398] SERVICE ALERT: nycnn-ms2fsd1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1408790408] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;3;(Host Check Timed Out)
[1408790418] SERVICE ALERT: nycnn-ms2fsd2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1408790478] SERVICE ALERT: nycnn-ms2fsd;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.27 ms
[1408790478] HOST ALERT: nycnn-ms2fsd;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1408790518] SERVICE ALERT: msny;MS2 - Grid B Mount;UNKNOWN;HARD;3;Remote command execution failed: ssh: connect to host 10.164.219.196 port 22: Connection refused
[1408790518] SERVICE ALERT: msny;MS2 - Grid A Mount;UNKNOWN;HARD;3;Remote command execution failed: ssh: connect to host 10.164.219.196 port 22: Connection refused
[1408790518] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.219.196 port 22: Connection refused
[1408790588] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1408790588] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1408790588] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1408790588] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1408790588] SERVICE ALERT: msny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1408790628] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;1;(Host Check Timed Out)
[1408790628] HOST ALERT: msny;DOWN;SOFT;1;(Host Check Timed Out)
[1408790648] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1408790648] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1408790648] SERVICE ALERT: msny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1408790728] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;2;(Host Check Timed Out)
[1408790728] HOST ALERT: msny;DOWN;SOFT;2;(Host Check Timed Out)
[1408790758] SERVICE ALERT: nycnn-ms2app3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1408790768] HOST ALERT: nycnn-ms2app3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1408790798] HOST ALERT: msny;UP;SOFT;3;PING OK - Packet loss = 16%, RTA = 0.23 ms
[1408790818] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1408790818] SERVICE FLAPPING ALERT: msny;MS2 - Grid A Mount;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1408790818] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;3;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1408790818] SERVICE FLAPPING ALERT: msny;MS2 - Grid B Mount;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1408790818] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;3;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1408790828] SERVICE ALERT: msny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1408791058] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.035 second response time
[1408791778] SERVICE FLAPPING ALERT: msny;MS2 - Grid A Mount;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1408791778] SERVICE FLAPPING ALERT: msny;MS2 - Grid B Mount;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1408797517] HOST DOWNTIME ALERT: msny;STOPPED; Host has exited from a period of scheduled downtime
[1408802408] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1408802448] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1408802468] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1408802518] HOST ALERT: omnmgsm-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1408802528] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1408804148] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1408804178] HOST ALERT: omnspsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1408804208] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1408804278] HOST ALERT: omnspsm-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1408804318] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1408804328] HOST ALERT: omnspsm-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1408816728] SERVICE ALERT: nycnn-ipvx26;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1408816788] SERVICE ALERT: nycnn-ipvx26;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.03%
[1408816868] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1408816928] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1408819848] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1408819908] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.03%
[1408831518] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 104.44 ms
[1408831578] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 5.47 ms