[1395899310] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.52 ms
[1395899370] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 110.86 ms
[1395899430] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 68.69 ms
[1395899490] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 107.14 ms
[1395899550] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 94.76 ms
[1395910110] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 100.16 ms
[1395910170] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 108.58 ms
[1395910230] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 87.92 ms
[1395910350] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 107.31 ms
[1395910410] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 113.04 ms
[1395910470] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 102.17 ms
[1395910590] SERVICE ALERT: rt-nyccfcp-1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.77 ms
[1395920910] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 127.20 ms
[1395920970] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 116.66 ms
[1395921030] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 65.41 ms
[1395921330] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 122.59 ms
[1395921390] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1395921870] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;SOFT;1;Job  failed with error HASH(0x2014c250)
[1395921940] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;SOFT;2;Job 10021304 progress is  18.56 :
[1395927240] HOST DOWNTIME ALERT: nycnn-ipvstrm2;STARTED; Host has entered a period of scheduled downtime
[1395927310] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1395927340] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1395927350] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395927350] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395927370] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1395927380] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395927380] SERVICE ALERT: nycnn-ipvstrm2;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395927380] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395927440] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1395927510] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1395927550] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1395927630] SERVICE ALERT: nycnn-ipvstrm2;SVServer;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1395927630] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1395927660] SERVICE ALERT: nycnn-ipvstrm2;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1395927660] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1395927660] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1168.98GB (83.7%), Free=228.21GB (16.3%)
[1395927690] SERVICE ALERT: nycnn-ipvstrm2;SVServer;CRITICAL;HARD;2;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1395927730] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1395927770] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1395927790] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1395927850] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1395927920] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 44%, RTA = 0.24 ms
[1395927920] HOST FLAPPING ALERT: nycnn-ipvstrm2;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1395927940] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "iftpsvc.exe" running (0 excluded).
[1395927970] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1395927970] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;PING;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1395927990] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;SOFT;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1395928260] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1395928270] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1395928280] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395928280] SERVICE ALERT: nycnn-ipvstrm2;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395928280] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395928310] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395928310] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395928350] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1395928430] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 66%, RTA = 0.32 ms
[1395928450] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1395928560] SERVICE ALERT: nycnn-ipvstrm2;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1395928560] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;MMSService;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1395928560] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1395928560] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;SSH Service;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1395928560] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1168.98GB (83.7%), Free=228.21GB (16.3%)
[1395928560] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;D:\ Drive Space;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1395928590] SERVICE ALERT: nycnn-ipvstrm2;SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1395928590] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;SVServer;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1395928590] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1395928590] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;FTP Service;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1395928690] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1395928730] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1395928750] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1395928770] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1395928850] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 64%, RTA = 0.25 ms
[1395928870] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1395929170] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1395929180] SERVICE ALERT: nycnn-ipvstrm2;MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1395929180] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1395929180] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1395929190] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1395929210] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395929210] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395929230] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1395929240] SERVICE ALERT: nycnn-ipvstrm2;MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395929240] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395929240] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395929290] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1395929350] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1395929360] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1395929490] SERVICE ALERT: nycnn-ipvstrm2;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1395929490] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1395929520] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1395929520] SERVICE ALERT: nycnn-ipvstrm2;MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1395929520] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;2;OK - Proxy Total=1397.19GB, Used=1168.98GB (83.7%), Free=228.21GB (16.3%)
[1395930430] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1395930624] HOST DOWNTIME ALERT: nycnn-ipvstrm3;STARTED; Host has entered a period of scheduled downtime
[1395931650] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 113.75 ms
[1395931710] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 71.86 ms
[1395931770] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 122.87 ms
[1395931830] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 105.95 ms
[1395931890] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 76.30 ms
[1395932010] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 124.83 ms
[1395932070] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 166.26 ms
[1395932130] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 85.37 ms
[1395932140] HOST FLAPPING ALERT: nycnn-ipvstrm2;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1395932330] SERVICE ALERT: nycnn-ipvstrm3;SVServer;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1395932390] SERVICE ALERT: nycnn-ipvstrm3;SVServer;CRITICAL;HARD;2;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1395932590] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1395932630] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1395932650] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395932650] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1395932710] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1395932710] SERVICE ALERT: nycnn-ipvstrm3;SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395932770] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1395932770] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1395932930] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1395932990] SERVICE ALERT: nycnn-ipvstrm3;SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1395933670] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1395933680] SERVICE ALERT: nycnn-ipvstrm3;MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1395933680] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1395933680] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1395933710] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1395933730] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1395933740] SERVICE ALERT: nycnn-ipvstrm3;MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395933740] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395933740] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1395933800] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1395933850] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1395933870] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1395933870] HOST FLAPPING ALERT: nycnn-ipvstrm3;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1395933910] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1395934020] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1395934020] SERVICE ALERT: nycnn-ipvstrm3;MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1395934030] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;OK;HARD;2;OK - New Volume Total=1397.19GB, Used=1152.17GB (82.5%), Free=245.02GB (17.5%)
[1395934130] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1395934130] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;FTP Service;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1395934437] HOST DOWNTIME ALERT: nycnn-ipvstrm2;STOPPED; Host has exited from a period of scheduled downtime
[1395935190] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;SVServer;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1395935190] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;FTP Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1395935220] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;SSH Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1395935220] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;MMSService;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1395935220] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;D:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1395936660] HOST FLAPPING ALERT: nycnn-ipvstrm3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1395939830] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;FTP Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1395941383] HOST DOWNTIME ALERT: nycnn-ipvstrm3;STOPPED; Host has exited from a period of scheduled downtime
[1395942510] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 132.38 ms
[1395942570] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 78.66 ms
[1395942680] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 133.05 ms
[1395942740] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 101.07 ms
[1395942800] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 117.52 ms
[1395942810] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.93 ms
[1395942860] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 44.31 ms
[1395942870] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 159.02 ms
[1395942920] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 131.34 ms
[1395942930] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 128.40 ms
[1395942980] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 139.18 ms
[1395942990] SERVICE ALERT: rt-nyccfcp-1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.24 ms
[1395943040] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 91.75 ms
[1395943100] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 139.73 ms
[1395943160] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 124.48 ms
[1395943220] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.93 ms
[1395945710] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 10022453 remains in a queued state
[1395945770] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;AME is idle:
[1395951430] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job 10023251 remains in a queued state
[1395951490] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;SOFT;2;Job 10023251 progress is  0 :
[1395953430] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 125.48 ms
[1395953480] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 113.85 ms
[1395953490] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 90.43 ms
[1395953540] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 89.03 ms
[1395953550] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.30 ms
[1395953610] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 99.68 ms
[1395953670] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 109.98 ms
[1395953720] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.50 ms
[1395953730] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 90.87 ms
[1395953780] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 120.59 ms
[1395953840] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 110.14 ms
[1395953960] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 71.04 ms
[1395963940] SERVICE ALERT: nycnn-adbeame10;AME Job status;CRITICAL;SOFT;1;Job 10023738 remains in a queued state
[1395963990] SERVICE ALERT: nycnn-adbeame10;AME Job status;OK;SOFT;2;AME is idle:
[1395964110] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 119.20 ms
[1395964170] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 69.30 ms
[1395964290] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 123.27 ms
[1395964340] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 126.04 ms
[1395964350] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 75.62 ms
[1395964400] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 99.20 ms
[1395964410] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 115.82 ms
[1395964460] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 111.99 ms
[1395964470] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 110.58 ms
[1395964520] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 119.64 ms
[1395964530] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 62.03 ms
[1395964580] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 89.35 ms
[1395964640] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 101.27 ms
[1395964700] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 128.48 ms
[1395964760] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 120.65 ms
[1395964820] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.06 ms
[1395974910] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 111.89 ms
[1395974970] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 104.02 ms
[1395975030] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 80.79 ms
[1395975090] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 126.51 ms
[1395975140] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 115.25 ms
[1395975150] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 42.91 ms
[1395975200] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 105.95 ms
[1395975260] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 102.29 ms
[1395975270] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 113.71 ms
[1395975330] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 119.73 ms
[1395975390] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1395975440] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 99.71 ms
[1395975500] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 104.82 ms
[1395975560] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 65.94 ms