[1352791368] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 302 processes with STATE = RSZDT
[1352791428] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 152 processes with STATE = RSZDT
[1352793228] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 302 processes with STATE = RSZDT
[1352793288] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 152 processes with STATE = RSZDT
[1352797788] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352797848] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 147 processes with STATE = RSZDT
[1352799648] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1352799708] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 146 processes with STATE = RSZDT
[1352801508] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352801568] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1352803368] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1352803428] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 147 processes with STATE = RSZDT
[1352805228] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1352805288] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 144 processes with STATE = RSZDT
[1352816388] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352816448] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 146 processes with STATE = RSZDT
[1352818248] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1352818308] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1352818788] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352818788] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352818828] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1352818828] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352818848] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352818848] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352818898] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352818898] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1352818898] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1352818898] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352820108] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1352820168] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 144 processes with STATE = RSZDT
[1352820648] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352820648] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352820688] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1352820688] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352820708] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352820708] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352820768] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 1.68 ms
[1352820768] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.54 ms
[1352820768] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.34 ms
[1352820768] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.27 ms
[1352820828] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352820828] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352820858] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1352820868] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352820888] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352820888] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352820938] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.42 ms
[1352820938] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1352820938] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352820938] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352821968] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 348 processes with STATE = RSZDT
[1352821968] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352821968] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352821998] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352821998] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1352822028] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 156 processes with STATE = RSZDT
[1352822028] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352822028] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352822038] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;2;(Host Check Timed Out)
[1352822038] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;2;(Host Check Timed Out)
[1352822078] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1352822078] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1352822088] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1352822088] HOST ALERT: nycnn-ipvx33;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352822268] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352822268] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352822308] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1352822308] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352822328] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352822328] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352822378] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352822378] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352822378] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352822378] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352823828] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1352823888] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 147 processes with STATE = RSZDT
[1352824428] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352824428] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352824458] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1352824458] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352824488] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352824488] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352824528] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.46 ms
[1352824528] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.47 ms
[1352824538] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1352824538] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1352824788] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352824788] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352824828] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1352824828] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352824848] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352824848] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352824888] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;2;(Host Check Timed Out)
[1352824888] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;2;(Host Check Timed Out)
[1352824958] HOST ALERT: nycnn-ipvx34;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1352824958] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352824958] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1352824958] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1352825688] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1352825748] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 148 processes with STATE = RSZDT
[1352826048] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352826048] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352826078] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1352826078] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352826108] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352826108] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352826158] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.31 ms
[1352826158] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.45 ms
[1352826158] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1352826158] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352826828] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352826828] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352826858] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352826858] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1352826888] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352826888] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352826928] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1352826928] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1352826938] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1352826938] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1352827548] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 361 processes with STATE = RSZDT
[1352827548] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352827548] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352827568] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1352827568] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352827608] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 63 processes with STATE = RSZDT
[1352827608] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352827608] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352827648] HOST ALERT: nycnn-ipvx31;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.29 ms
[1352827648] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.34 ms
[1352827658] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1352827658] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352828748] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352828748] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352828788] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1352828788] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352828808] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352828808] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352828858] HOST ALERT: nycnn-ipvx31;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352828858] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352828858] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1352828858] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1352829408] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 348 processes with STATE = RSZDT
[1352829408] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352829408] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352829438] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1352829438] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352829468] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 150 processes with STATE = RSZDT
[1352829468] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352829468] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352829498] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.45 ms
[1352829498] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.50 ms
[1352829518] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1352829518] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1352830608] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352830608] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352830648] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1352830648] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352830668] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352830668] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352830718] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1352830718] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352830718] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1352830718] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352831268] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1352831328] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 144 processes with STATE = RSZDT
[1352833128] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352833188] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1352834988] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1352835048] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 148 processes with STATE = RSZDT
[1352836848] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 346 processes with STATE = RSZDT
[1352836908] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 144 processes with STATE = RSZDT
[1352838708] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352838768] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1352840568] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352840628] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 145 processes with STATE = RSZDT
[1352841228] SERVICE ALERT: omnmas1app-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352841288] SERVICE ALERT: omnmas1app-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1352841348] SERVICE ALERT: omnmas1app-ny;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1352841428] HOST ALERT: omnmas1app-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1352841528] HOST ALERT: omnmas1app-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1352841628] HOST ALERT: omnmas1app-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1352841728] HOST ALERT: omnmas1app-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1352841828] HOST ALERT: omnmas1app-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1352841928] HOST ALERT: omnmas1app-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1352842028] HOST ALERT: omnmas1app-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1352842118] HOST ALERT: omnmas1app-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1352842218] HOST ALERT: omnmas1app-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1352842318] HOST ALERT: omnmas1app-ny;DOWN;HARD;10;(Host Check Timed Out)
[1352842428] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 339 processes with STATE = RSZDT
[1352842488] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 146 processes with STATE = RSZDT
[1352844288] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1352844348] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 155 processes with STATE = RSZDT
[1352845178] SERVICE ALERT: omnmas1app-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1352845188] HOST ALERT: omnmas1app-ny;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1352845918] SERVICE ALERT: omnmdir01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352845918] SERVICE ALERT: omnmdir01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352845958] HOST ALERT: omnmdir01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1352845958] HOST ALERT: omnmdir01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1352845968] SERVICE ALERT: omnmdir01e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352845968] SERVICE ALERT: omnmdir01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352846028] HOST ALERT: omnmdir01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1352846028] SERVICE ALERT: omnmdir01e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1352846028] SERVICE ALERT: omnmdir01-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1352846028] HOST ALERT: omnmdir01e1-ny;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.19 ms
[1352846148] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1352846208] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1352848008] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352848068] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 146 processes with STATE = RSZDT
[1352849868] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352849928] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1352851728] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352851788] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1352853588] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352853648] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 146 processes with STATE = RSZDT
[1352855448] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 346 processes with STATE = RSZDT
[1352855508] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 141 processes with STATE = RSZDT
[1352857308] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352857368] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1352865768] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 279 processes with STATE = RSZDT
[1352865828] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 153 processes with STATE = RSZDT
[1352867628] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 278 processes with STATE = RSZDT
[1352867688] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 152 processes with STATE = RSZDT