[1340081314] SERVICE ALERT: wdcps001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 6.65 ms
[1340081354] HOST ALERT: wdcps001;DOWN;SOFT;1;(Host Check Timed Out)
[1340081374] SERVICE ALERT: wdcps001;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1340081424] HOST ALERT: wdcps001;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1340081424] SERVICE ALERT: wdcps001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1340106994] SERVICE ALERT: wdcps004;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340107034] HOST ALERT: wdcps004;DOWN;SOFT;1;(Host Check Timed Out)
[1340107054] SERVICE ALERT: wdcps004;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 6.64 ms
[1340107064] HOST ALERT: wdcps004;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1340107104] SERVICE ALERT: wdcps004;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.72 ms
[1340114644] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1340114704] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1340114754] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;SOFT;3;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1340115094] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1340115144] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;SOFT;2;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1340119044] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 236.40 ms
[1340119104] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 200.40 ms
[1340119164] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 187.25 ms
[1340120184] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 7.22 ms
[1340122284] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 233.48 ms
[1340122344] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 206.72 ms
[1340122404] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 206.56 ms
[1340123004] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1340123064] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 224.94 ms
[1340123124] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 191.18 ms
[1340123184] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 230.04 ms
[1340124634] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1340124684] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;SOFT;2;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1340126604] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 98.65 ms
[1340126664] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 214.93 ms
[1340126724] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 225.80 ms
[1340126784] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 186.90 ms
[1340127384] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 8.10 ms
[1340129554] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340129554] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1340129594] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1340129594] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1340129614] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340129614] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1340129664] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1340129664] HOST ALERT: nycnn-ipvx13;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1340129664] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1340129664] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1340140764] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 184.45 ms
[1340140824] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 7.52 ms
[1340140944] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 220.24 ms
[1340141004] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 210.22 ms
[1340141064] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 225.98 ms
[1340141484] SERVICE ALERT: nycnn-ipvstrm4;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 37%, RTA = 0.26 ms
[1340141544] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1340141664] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1340141724] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 178.17 ms
[1340141784] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 16%, RTA = 236.23 ms
[1340141844] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 230.33 ms
[1340143764] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1340143884] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 228.83 ms
[1340143944] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 223.93 ms
[1340144004] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 188.63 ms
[1340145084] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1340154445] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 233.90 ms
[1340154505] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1340157625] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 216.72 ms
[1340157685] SERVICE ALERT: nycnn-gridpushpull2;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 234.62 ms
[1340157745] SERVICE ALERT: nycnn-gridpushpull2;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms