[1403786166] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1403786226] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1403791219] SERVICE DOWNTIME ALERT: nycnn-adbeame6;PING;STARTED; Service has entered a period of scheduled downtime
[1403791219] SERVICE DOWNTIME ALERT: nycnn-adbeame6;AME Job status;STARTED; Service has entered a period of scheduled downtime
[1403795826] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;HARD;2;AME is idle:
[1403798410] SERVICE DOWNTIME ALERT: nycnn-adbeame6;PING;STOPPED; Service has exited from a period of scheduled downtime
[1403798410] SERVICE DOWNTIME ALERT: nycnn-adbeame6;AME Job status;STOPPED; Service has exited from a period of scheduled downtime
[1403808776] SERVICE ALERT: nycnn-adbeame4;AME Job status;CRITICAL;SOFT;1;Job 10246988 remains in a queued state
[1403808826] SERVICE ALERT: nycnn-adbeame4;AME Job status;OK;SOFT;2;AME is idle:
[1403814886] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 10247468 remains in a queued state
[1403814946] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;AME is idle:
[1403816226] SERVICE ALERT: nycnn-ipvapilic;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1403816286] SERVICE ALERT: nycnn-ipvapilic;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1403817406] SERVICE ALERT: nycnn-mamkey1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-mamdb2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-arvtrmops2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2db3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-mamkey2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2db1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-arvtrmops;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-mammq2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2app4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-mamfog;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1403817406] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1403817406] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1403817406] SERVICE ALERT: nymamkey;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1403817406] SERVICE ALERT: nymq;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-mammq1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2stills1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: rt-nyccfcp-2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ipvapilic;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1403817406] SERVICE ALERT: nycnn-ms2stills2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2db2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: msny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-mamdb1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: rt-nyccfcp-1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403817406] SERVICE ALERT: nycnn-ipvapilic;Process: sntlkeys;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1403817416] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817416] SERVICE ALERT: nycnn-ipvapilic;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1403817436] HOST ALERT: nycnn-mamkey1;DOWN;SOFT;1;(Host Check Timed Out)
[1403817436] HOST ALERT: nycnn-mamdb2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817436] HOST ALERT: nycnn-arvtrmops2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-mamkey2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-mammq2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-mamfog;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: msny;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nymamkey;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nymq;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-mammq1;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2stills1;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: rt-nyccfcp-2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2stills2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: rt-nyccfcp-1;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: rt-nycc3750tgr-ilo;DOWN;SOFT;1;(Host Check Timed Out)
[1403817446] HOST ALERT: rt-nycc3750tgr-ms2;DOWN;SOFT;1;(Host Check Timed Out)
[1403817456] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;CRITICAL;HARD;1;(Service Check Timed Out)
[1403817456] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;HARD;1;(Service Check Timed Out)
[1403817456] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;HARD;1;(Service Check Timed Out)
[1403817466] SERVICE ALERT: nycnn-mamkey1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-mamdb2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-arvtrmops2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2db3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-mamkey2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2db1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-arvtrmops;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-mammq2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2app4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-mamfog;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;2;CRITICAL - Socket timeout after 10 seconds
[1403817466] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1403817466] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1403817466] SERVICE ALERT: nymamkey;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1403817466] SERVICE ALERT: nymq;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-mammq1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2stills1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: rt-nyccfcp-2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ipvapilic;Process: sntlsrt;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1403817466] SERVICE ALERT: nycnn-ipvapilic;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1403817466] SERVICE ALERT: nycnn-ms2stills2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2db2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: msny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-mamdb1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: rt-nyccfcp-1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403817466] SERVICE ALERT: nycnn-ipvapilic;Process: sntlkeys;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1403817476] SERVICE ALERT: nycnn-ipvapilic;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1403817486] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;2;(Host Check Timed Out)
[1403817486] SERVICE ALERT: nycnn-arvtrmops2;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1403817496] SERVICE ALERT: nycnn-ipvapilic;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1403817506] HOST ALERT: rt-nycc3750tgr-ilo;DOWN;SOFT;2;(Host Check Timed Out)
[1403817506] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817526] HOST ALERT: nycnn-mamkey1;DOWN;SOFT;2;(Host Check Timed Out)
[1403817526] HOST ALERT: nycnn-mamdb2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817526] HOST ALERT: nycnn-arvtrmops2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-mamkey2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-mammq2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-mamfog;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: msny;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nymamkey;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nymq;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-mammq1;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2stills1;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: rt-nyccfcp-2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2stills2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: rt-nyccfcp-1;DOWN;SOFT;2;(Host Check Timed Out)
[1403817546] HOST ALERT: rt-nycc3750tgr-ms2;DOWN;SOFT;2;(Host Check Timed Out)
[1403817576] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;3;(Host Check Timed Out)
[1403817606] HOST ALERT: rt-nycc3750tgr-ilo;DOWN;SOFT;3;(Host Check Timed Out)
[1403817606] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817626] HOST ALERT: nycnn-mamkey1;DOWN;SOFT;3;(Host Check Timed Out)
[1403817626] HOST ALERT: nycnn-mamdb2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817626] HOST ALERT: nycnn-arvtrmops2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817636] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-mamkey2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-mammq2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-mamfog;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: msny;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nymamkey;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nymq;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-mammq1;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ms2stills1;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: rt-nyccfcp-2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ms2stills2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: rt-nyccfcp-1;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] HOST ALERT: rt-nycc3750tgr-ms2;DOWN;SOFT;3;(Host Check Timed Out)
[1403817646] SERVICE ALERT: nycnn-arvtrmops;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1403817676] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;4;(Host Check Timed Out)
[1403817686] SERVICE ALERT: nycnn-arvtrtape;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1403817686] SERVICE ALERT: nycnn-ipvapilic;DISK_C:;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1403817706] HOST ALERT: rt-nycc3750tgr-ilo;DOWN;SOFT;4;(Host Check Timed Out)
[1403817706] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817726] HOST ALERT: nycnn-mamkey1;DOWN;SOFT;4;(Host Check Timed Out)
[1403817726] HOST ALERT: nycnn-mamdb2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817726] HOST ALERT: nycnn-arvtrmops2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817736] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-mamkey2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-mammq2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-mamfog;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: msny;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nymamkey;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nymq;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-mammq1;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ms2stills1;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: rt-nyccfcp-2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ms2stills2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: rt-nyccfcp-1;DOWN;SOFT;4;(Host Check Timed Out)
[1403817746] HOST ALERT: rt-nycc3750tgr-ms2;DOWN;SOFT;4;(Host Check Timed Out)
[1403817776] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;5;(Host Check Timed Out)
[1403817796] HOST ALERT: rt-nycc3750tgr-ilo;DOWN;SOFT;5;(Host Check Timed Out)
[1403817796] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817826] HOST ALERT: nycnn-mamkey1;DOWN;SOFT;5;(Host Check Timed Out)
[1403817826] HOST ALERT: nycnn-mamdb2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817826] HOST ALERT: nycnn-arvtrmops2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817836] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;5;(Host Check Timed Out)
[1403817836] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;5;(Host Check Timed Out)
[1403817836] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-mamkey2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-mammq2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-mamfog;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: msny;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nymamkey;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nymq;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-mammq1;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-ms2stills1;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: rt-nyccfcp-2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-ms2stills2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: rt-nyccfcp-1;DOWN;SOFT;5;(Host Check Timed Out)
[1403817846] HOST ALERT: rt-nycc3750tgr-ms2;DOWN;SOFT;5;(Host Check Timed Out)
[1403817876] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;6;(Host Check Timed Out)
[1403817896] HOST ALERT: rt-nycc3750tgr-ilo;DOWN;SOFT;6;(Host Check Timed Out)
[1403817896] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817926] HOST ALERT: nycnn-mamkey1;DOWN;SOFT;6;(Host Check Timed Out)
[1403817926] HOST ALERT: nycnn-mamdb2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817926] HOST ALERT: nycnn-arvtrmops2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817936] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;6;(Host Check Timed Out)
[1403817936] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;6;(Host Check Timed Out)
[1403817936] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;6;(Host Check Timed Out)
[1403817936] HOST ALERT: nycnn-mamkey2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817936] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;6;(Host Check Timed Out)
[1403817936] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;6;(Host Check Timed Out)
[1403817936] HOST ALERT: nycnn-mammq2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-mamfog;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: msny;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1403817946] HOST ALERT: nymamkey;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1403817946] HOST ALERT: nymq;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-mammq1;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-ms2stills1;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: rt-nyccfcp-2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-ms2stills2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-ms2db2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: rt-nyccfcp-1;DOWN;SOFT;6;(Host Check Timed Out)
[1403817946] HOST ALERT: rt-nycc3750tgr-ms2;DOWN;SOFT;6;(Host Check Timed Out)
[1403817966] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;OK;HARD;1;DISK OK - free space: /oracle 779480 MB (79% inode=99%):
[1403817966] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;OK;HARD;1;DISK OK - free space: /oracle 859365 MB (87% inode=99%):
[1403817966] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;OK;HARD;1;DISK OK - free space: /oracle 897032 MB (91% inode=99%):
[1403817966] HOST ALERT: nycnn-ms2fsd1;UP;SOFT;7;PING WARNING - Packet loss = 80%, RTA = 0.28 ms
[1403817966] HOST ALERT: nycnn-ms2fsd2;UP;SOFT;7;PING OK - Packet loss = 16%, RTA = 0.13 ms
[1403817966] HOST ALERT: rt-nycc3750tgr-ilo;UP;SOFT;7;PING OK - Packet loss = 16%, RTA = 1.86 ms
[1403817976] SERVICE ALERT: nycnn-ipvapilic;DISK_C:;OK;HARD;1;OK - C: Total=136.69GB, Used=8.74GB (6.4%), Free=127.94GB (93.6%)
[1403817976] HOST ALERT: nycnn-ms2db2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1403817976] HOST ALERT: nycnn-ms2db3;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1403817976] HOST ALERT: nycnn-ms2db1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1403817986] HOST ALERT: nycnn-ipvapilic;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1403817996] SERVICE ALERT: msny;HTTP;OK;HARD;2;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.002 second response time
[1403817996] HOST ALERT: nycnn-mamkey1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1403817996] HOST ALERT: nycnn-mamdb2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1403817996] SERVICE ALERT: nycnn-mamkey1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1403817996] SERVICE ALERT: nycnn-mamdb2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1403817996] SERVICE ALERT: nycnn-ms2fsd1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1403817996] SERVICE ALERT: nycnn-ms2db3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1403817996] SERVICE ALERT: nycnn-mamkey2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1403817996] SERVICE ALERT: nycnn-ms2db1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1403817996] SERVICE ALERT: nycnn-mammq2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1403817996] SERVICE ALERT: nycnn-ms2app3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1403817996] SERVICE ALERT: nycnn-ms2app4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1403817996] SERVICE ALERT: nycnn-mamfog;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1403817996] SERVICE ALERT: nycnn-ipvapilic;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1403817996] SERVICE ALERT: nymamkey;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1403817996] SERVICE ALERT: nymq;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1403817996] SERVICE ALERT: nycnn-mammq1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1403817996] SERVICE ALERT: nycnn-ms2stills1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1403817996] SERVICE ALERT: rt-nyccfcp-2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1403818006] SERVICE ALERT: nycnn-ms2stills2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1403818006] SERVICE ALERT: nycnn-ms2db2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1403818006] SERVICE ALERT: nycnn-ms2fsd2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1403818006] SERVICE ALERT: msny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1403818006] SERVICE ALERT: nycnn-ms2fsd;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1403818006] SERVICE ALERT: nycnn-mamdb1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1403818006] SERVICE ALERT: rt-nyccfcp-1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.94 ms
[1403818006] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 1.15 ms
[1403818006] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1403818006] HOST ALERT: msny;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1403818006] HOST ALERT: nycnn-mamkey2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1403818006] HOST ALERT: nycnn-mammq2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1403818006] HOST ALERT: nycnn-ms2app3;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1403818006] HOST ALERT: nycnn-ms2app4;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1403818006] HOST ALERT: nycnn-mamfog;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1403818006] HOST ALERT: nymamkey;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1403818006] HOST ALERT: nymq;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1403818006] HOST ALERT: nycnn-mammq1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1403818006] HOST ALERT: nycnn-ms2stills1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1403818006] HOST ALERT: rt-nyccfcp-2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.94 ms
[1403818016] HOST ALERT: nycnn-ms2stills2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1403818016] HOST ALERT: nycnn-ms2fsd;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1403818016] HOST ALERT: nycnn-mamdb1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1403818016] HOST ALERT: rt-nyccfcp-1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1403818016] HOST ALERT: rt-nycc3750tgr-ms2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1403818026] HOST ALERT: nycnn-arvtrmops2;DOWN;SOFT;7;(Host Check Timed Out)
[1403818036] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;7;(Host Check Timed Out)
[1403818036] HOST ALERT: nycnn-arvtrmops;DOWN;SOFT;7;(Host Check Timed Out)
[1403818056] SERVICE ALERT: nycnn-ipvapilic;Process: sntlsrt;OK;HARD;1;OK - Found 1 Instance(s) of "sntlsrtsrvr.exe" running (0 excluded).  (List is on next line)
[1403818056] SERVICE ALERT: nycnn-ipvapilic;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1403818056] SERVICE ALERT: nycnn-ipvapilic;Process: sntlkeys;OK;HARD;2;OK - Found 1 Instance(s) of "sntlkeyssrvr.exe" running (0 excluded).  (List is on next line)
[1403818056] SERVICE ALERT: nycnn-arvtrmops2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1403818056] SERVICE ALERT: nycnn-arvtrmops;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1403818066] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1403818066] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1403818066] HOST ALERT: nycnn-arvtrmops2;UP;SOFT;8;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1403818066] HOST ALERT: nycnn-arvtrmops;UP;SOFT;8;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1403818066] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1403818066] SERVICE ALERT: nycnn-ipvapilic;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 3.248GB - Used: 0.530GB (16%) - Free: 2.718GB (84%)
[1403818076] SERVICE ALERT: nycnn-arvtrmops2;SSH;OK;HARD;1;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1403818086] SERVICE ALERT: nycnn-ipvapilic;CPU Load;OK;HARD;1;OK (Sample Period 900 sec) - Average CPU Utilisation 0.06%
[1403818106] HOST ALERT: nycnn-arvtrtape;UP;SOFT;8;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1403818116] SERVICE ALERT: nycnn-arvtrtape;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1403818232] SERVICE ALERT: nycnn-arvtrmops;SSH;OK;HARD;1;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1403818272] SERVICE ALERT: nycnn-arvtrtape;SSH;OK;HARD;1;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1403820412] SERVICE ALERT: rt-nycchdagg-a;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 159.24 ms
[1403820472] SERVICE ALERT: rt-nycchdagg-a;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 45.76 ms
[1403821922] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403821962] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403821982] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403821982] SERVICE ALERT: omnmas1app-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403821982] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403822022] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403822022] HOST ALERT: omnmas1app-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403822042] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403822042] SERVICE ALERT: omnmas1app-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403822052] HOST ALERT: omnmas1app1-ny;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.23 ms
[1403822062] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403822092] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1403822122] HOST ALERT: omnmas1app-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403822142] HOST ALERT: omnmas1app2-ny;UP;SOFT;3;PING OK - Packet loss = 66%, RTA = 0.23 ms
[1403822152] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1403822192] HOST ALERT: omnmas1app-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1403822212] SERVICE ALERT: omnmas1app-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1403822562] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1403822622] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.04%
[1403824202] SERVICE ALERT: omnclda02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403824202] SERVICE ALERT: omnclda02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403824232] HOST ALERT: omnclda02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403824242] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403824262] SERVICE ALERT: omnclda02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403824262] SERVICE ALERT: omnclda02e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403824332] HOST ALERT: omnclda02-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403824342] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403824432] HOST ALERT: omnclda02-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1403824442] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1403824532] HOST ALERT: omnclda02-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1403824542] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1403824632] HOST ALERT: omnclda02-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1403824642] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1403824732] HOST ALERT: omnclda02-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1403824742] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1403824802] SERVICE ALERT: omnclda03e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403824802] SERVICE ALERT: omnclda03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403824832] HOST ALERT: omnclda02-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1403824842] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403824842] HOST ALERT: omnclda03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403824842] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1403824862] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403824862] SERVICE ALERT: omnclda03e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403824862] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1403824862] SERVICE ALERT: omnclda03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403824902] HOST ALERT: omnclda01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403824902] HOST ALERT: omnclda01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1403824922] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403824922] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1403824922] HOST ALERT: omnclda02-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1403824942] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403824942] HOST ALERT: omnclda03-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1403824942] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1403824972] HOST ALERT: omnclda01-ny;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.51 ms
[1403824972] SERVICE ALERT: omnclda01-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1403824972] HOST ALERT: omnclda01e1-ny;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.22 ms
[1403824972] SERVICE ALERT: omnclda01e1-ny;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 44%, RTA = 0.23 ms
[1403825012] HOST ALERT: omnclda02-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1403825032] SERVICE ALERT: omnclda01e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1403825042] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1403825042] HOST ALERT: omnclda03-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1403825042] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1403825112] HOST ALERT: omnclda02-ny;DOWN;HARD;10;(Host Check Timed Out)
[1403825140] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1403825140] HOST ALERT: omnclda03-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1403825140] HOST ALERT: omnclda02e1-ny;DOWN;HARD;10;(Host Check Timed Out)
[1403825238] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1403825238] HOST ALERT: omnclda03-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1403825382] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1403825382] HOST ALERT: omnclda03-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1403825482] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1403825482] HOST ALERT: omnclda03-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1403825572] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1403825572] HOST ALERT: omnclda03-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1403825672] HOST ALERT: omnclda03e1-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1403825672] HOST ALERT: omnclda03-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1403825772] HOST ALERT: omnclda03e1-ny;DOWN;HARD;10;(Host Check Timed Out)
[1403825792] HOST ALERT: omnclda03-ny;DOWN;HARD;10;(Host Check Timed Out)
[1403826180] SERVICE ALERT: omnclda03-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1403826180] SERVICE ALERT: omnclda03e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1403826190] HOST ALERT: omnclda03-ny;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1403826208] HOST ALERT: omnclda03e1-ny;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1403830328] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1403830328] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;3;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1403830328] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;3;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1403830748] SERVICE ALERT: omnclda02e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1403830748] SERVICE ALERT: omnclda02-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1403830758] HOST ALERT: omnclda02e1-ny;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1403830758] HOST ALERT: omnclda02-ny;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1403836388] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1403836438] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1403836498] SERVICE ALERT: msny;HTTP;OK;SOFT;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.009 second response time
[1403837278] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 321 processes with STATE = RSZDT
[1403837338] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 62 processes with STATE = RSZDT
[1403839138] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 326 processes with STATE = RSZDT
[1403839198] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 61 processes with STATE = RSZDT
[1403840998] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 321 processes with STATE = RSZDT
[1403841058] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 57 processes with STATE = RSZDT