[1353052608] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1353052668] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353054468] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1353054528] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 136 processes with STATE = RSZDT
[1353056328] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1353056388] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353058188] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1353058248] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1353060048] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 349 processes with STATE = RSZDT
[1353060108] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353061908] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1353061968] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353063768] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1353063828] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353065628] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 333 processes with STATE = RSZDT
[1353065688] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 141 processes with STATE = RSZDT
[1353074328] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353074328] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353074358] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1353074358] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353074388] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353074388] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353074428] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353074428] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353074438] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353074438] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353074568] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353074568] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353074608] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1353074608] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353074628] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353074628] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353074678] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1353074678] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1353074678] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353074678] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353074988] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353074988] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353075028] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1353075028] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353075048] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353075048] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353075078] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;2;(Host Check Timed Out)
[1353075078] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353075098] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1353075098] SERVICE FLAPPING ALERT: nycnn-ipvx7;PING;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1353075098] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.84 ms
[1353075098] SERVICE FLAPPING ALERT: nycnn-ipvx7e1;PING;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1353075108] HOST ALERT: nycnn-ipvx7;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353075108] HOST FLAPPING ALERT: nycnn-ipvx7;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1353075108] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353075108] HOST FLAPPING ALERT: nycnn-ipvx7e1;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1353075408] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353075408] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353075448] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1353075448] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353075468] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353075468] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353075518] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.36 ms
[1353075518] HOST FLAPPING ALERT: nycnn-ipvx8;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1353075518] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.47 ms
[1353075518] HOST FLAPPING ALERT: nycnn-ipvx8e1;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1353075518] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1353075518] SERVICE FLAPPING ALERT: nycnn-ipvx8;PING;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1353075518] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1353075518] SERVICE FLAPPING ALERT: nycnn-ipvx8e1;PING;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1353076178] HOST FLAPPING ALERT: nycnn-ipvx7;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353076178] HOST FLAPPING ALERT: nycnn-ipvx7e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353076188] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353076188] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353076218] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353076218] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1353076238] SERVICE FLAPPING ALERT: nycnn-ipvx7;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353076238] SERVICE FLAPPING ALERT: nycnn-ipvx7e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353076248] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353076248] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353076278] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.54 ms
[1353076278] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.48 ms
[1353076298] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353076298] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1353076658] SERVICE FLAPPING ALERT: nycnn-ipvx8;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353076658] HOST FLAPPING ALERT: nycnn-ipvx8;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353076658] SERVICE FLAPPING ALERT: nycnn-ipvx8e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353076658] HOST FLAPPING ALERT: nycnn-ipvx8e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353076788] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1353076848] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353077028] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077028] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077058] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353077058] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1353077088] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077088] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077128] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1353077128] HOST FLAPPING ALERT: nycnn-ipvx6e1;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1353077128] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1353077128] HOST FLAPPING ALERT: nycnn-ipvx6;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1353077138] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353077138] SERVICE FLAPPING ALERT: nycnn-ipvx6e1;PING;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1353077138] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353077138] SERVICE FLAPPING ALERT: nycnn-ipvx6;PING;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1353077288] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077328] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1353077348] SERVICE ALERT: wdcms002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077348] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077398] HOST ALERT: wdcms003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1353077398] SERVICE ALERT: wdcms002;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1353077398] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.66 ms
[1353077868] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077868] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077908] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1353077908] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353077928] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077928] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077928] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077928] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353077928] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077928] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077958] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353077958] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1353077958] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353077958] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1353077988] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077988] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077988] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077988] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353077988] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.30 ms
[1353077988] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.28 ms
[1353078028] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353078028] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.07 ms
[1353078028] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.51 ms
[1353078028] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.47 ms
[1353078038] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353078038] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1353078038] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353078038] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353078038] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353078038] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353078278] SERVICE FLAPPING ALERT: nycnn-ipvx6e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353078278] HOST FLAPPING ALERT: nycnn-ipvx6e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353078278] SERVICE FLAPPING ALERT: nycnn-ipvx6;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353078278] HOST FLAPPING ALERT: nycnn-ipvx6;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353078588] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353078588] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353078618] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1353078618] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353078648] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 369 processes with STATE = RSZDT
[1353078648] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353078648] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353078708] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 156 processes with STATE = RSZDT
[1353078708] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.52 ms
[1353078708] HOST FLAPPING ALERT: nycnn-ipvx5;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1353078708] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.52 ms
[1353078708] HOST FLAPPING ALERT: nycnn-ipvx5e1;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1353078758] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353078758] SERVICE FLAPPING ALERT: nycnn-ipvx5;PING;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1353078758] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353078758] SERVICE FLAPPING ALERT: nycnn-ipvx5e1;PING;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1353078828] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353078828] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353078858] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1353078858] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353078888] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353078888] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353078938] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.49 ms
[1353078938] HOST FLAPPING ALERT: nycnn-ipvx4;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1353078938] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.52 ms
[1353078938] HOST FLAPPING ALERT: nycnn-ipvx4e1;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1353078938] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353078938] SERVICE FLAPPING ALERT: nycnn-ipvx4;PING;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1353078938] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353078938] SERVICE FLAPPING ALERT: nycnn-ipvx4e1;PING;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1353079838] SERVICE FLAPPING ALERT: nycnn-ipvx5;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353079838] HOST FLAPPING ALERT: nycnn-ipvx5;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353079838] SERVICE FLAPPING ALERT: nycnn-ipvx5e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353079838] HOST FLAPPING ALERT: nycnn-ipvx5e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353080078] SERVICE FLAPPING ALERT: nycnn-ipvx4;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353080078] HOST FLAPPING ALERT: nycnn-ipvx4;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353080078] SERVICE FLAPPING ALERT: nycnn-ipvx4e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353080078] HOST FLAPPING ALERT: nycnn-ipvx4e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353080508] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 354 processes with STATE = RSZDT
[1353080568] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 150 processes with STATE = RSZDT
[1353080568] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353080568] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353080608] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353080608] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1353080628] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353080628] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353080678] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1353080678] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1353080678] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353080678] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1353082368] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1353082428] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 144 processes with STATE = RSZDT
[1353083048] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353083098] SERVICE ALERT: wdcms003;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.59 ms
[1353084128] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 6.62 ms
[1353084158] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1353084188] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353084218] HOST ALERT: wdcms003;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.55 ms
[1353084228] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1353084238] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1353084288] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 55 processes with STATE = RSZDT
[1353084428] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353084458] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1353084488] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353084548] HOST ALERT: wdcms003;DOWN;SOFT;2;(Host Check Timed Out)
[1353084598] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1353084598] SERVICE FLAPPING ALERT: wdcms003;PING;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1353084608] HOST ALERT: wdcms003;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1353084608] HOST FLAPPING ALERT: wdcms003;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1353085618] HOST FLAPPING ALERT: wdcms003;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1353085678] SERVICE FLAPPING ALERT: wdcms003;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353086088] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1353086148] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 147 processes with STATE = RSZDT
[1353087948] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 348 processes with STATE = RSZDT
[1353088008] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 145 processes with STATE = RSZDT
[1353089808] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 349 processes with STATE = RSZDT
[1353089868] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353091668] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1353091728] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 56 processes with STATE = RSZDT
[1353091908] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.16 ms
[1353091908] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353091908] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353091908] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353091938] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1353091938] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353091938] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1353091938] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353091968] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353091968] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353091968] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353091968] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353092008] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.51 ms
[1353092008] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.47 ms
[1353092018] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353092018] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1353092038] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;2;(Host Check Timed Out)
[1353092038] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353092078] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353092078] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353092088] HOST ALERT: nycnn-ipvx10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353092088] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1353092628] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353092628] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353092648] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353092648] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1353092688] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353092688] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353092728] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.47 ms
[1353092728] HOST FLAPPING ALERT: nycnn-ipvx9e1;STARTED; Host appears to have started flapping (21.6% change > 20.0% threshold)
[1353092728] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.49 ms
[1353092728] HOST FLAPPING ALERT: nycnn-ipvx9;STARTED; Host appears to have started flapping (21.6% change > 20.0% threshold)
[1353092738] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1353092738] SERVICE FLAPPING ALERT: nycnn-ipvx9e1;PING;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1353092738] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1353092738] SERVICE FLAPPING ALERT: nycnn-ipvx9;PING;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1353092808] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353092808] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353092838] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1353092838] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353092868] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353092868] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353092908] HOST ALERT: nycnn-ipvx10;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353092908] HOST FLAPPING ALERT: nycnn-ipvx10;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1353092908] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1353092908] HOST FLAPPING ALERT: nycnn-ipvx10e1;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1353092918] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353092918] SERVICE FLAPPING ALERT: nycnn-ipvx10;PING;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1353092918] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353092918] SERVICE FLAPPING ALERT: nycnn-ipvx10e1;PING;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1353093528] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 333 processes with STATE = RSZDT
[1353093588] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 146 processes with STATE = RSZDT
[1353093828] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353093828] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353093828] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353093828] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353093858] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1353093858] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1353093858] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353093858] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353093878] SERVICE FLAPPING ALERT: nycnn-ipvx9e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353093878] HOST FLAPPING ALERT: nycnn-ipvx9e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353093878] SERVICE FLAPPING ALERT: nycnn-ipvx9;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353093878] HOST FLAPPING ALERT: nycnn-ipvx9;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353093888] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353093888] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353093888] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353093888] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353093928] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353093928] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353093938] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353093938] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1353093938] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.55 ms
[1353093938] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.50 ms
[1353093948] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.33 ms
[1353093948] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.31 ms
[1353094058] SERVICE FLAPPING ALERT: nycnn-ipvx10;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353094058] HOST FLAPPING ALERT: nycnn-ipvx10;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353094058] SERVICE FLAPPING ALERT: nycnn-ipvx10e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353094058] HOST FLAPPING ALERT: nycnn-ipvx10e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353094548] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353094548] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353094578] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1353094578] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353094608] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353094608] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353094648] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.50 ms
[1353094648] HOST FLAPPING ALERT: nycnn-ipvx17;STARTED; Host appears to have started flapping (21.6% change > 20.0% threshold)
[1353094648] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.51 ms
[1353094648] HOST FLAPPING ALERT: nycnn-ipvx17e1;STARTED; Host appears to have started flapping (21.6% change > 20.0% threshold)
[1353094658] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1353094658] SERVICE FLAPPING ALERT: nycnn-ipvx17;PING;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1353094658] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353094658] SERVICE FLAPPING ALERT: nycnn-ipvx17e1;PING;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1353094728] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353094728] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353094758] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1353094758] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353094788] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353094788] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353094838] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.47 ms
[1353094838] HOST FLAPPING ALERT: nycnn-ipvx16;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1353094838] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 1.09 ms
[1353094838] HOST FLAPPING ALERT: nycnn-ipvx16e1;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1353094838] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1353094838] SERVICE FLAPPING ALERT: nycnn-ipvx16;PING;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1353094838] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1353094838] SERVICE FLAPPING ALERT: nycnn-ipvx16e1;PING;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1353095388] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353095448] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 61 processes with STATE = RSZDT
[1353095798] SERVICE FLAPPING ALERT: nycnn-ipvx17;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353095798] HOST FLAPPING ALERT: nycnn-ipvx17;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353095798] SERVICE FLAPPING ALERT: nycnn-ipvx17e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353095798] HOST FLAPPING ALERT: nycnn-ipvx17e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353095978] SERVICE FLAPPING ALERT: nycnn-ipvx16;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353095978] HOST FLAPPING ALERT: nycnn-ipvx16;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353095978] SERVICE FLAPPING ALERT: nycnn-ipvx16e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353095978] HOST FLAPPING ALERT: nycnn-ipvx16e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353097248] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 358 processes with STATE = RSZDT
[1353097308] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 147 processes with STATE = RSZDT
[1353097788] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353097788] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353097828] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1353097828] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353097848] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353097848] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353097898] HOST ALERT: nycnn-ipvx31;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353097898] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1353097898] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353097898] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353097978] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353098018] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;1;(Host Check Timed Out)
[1353098038] SERVICE ALERT: nycnn-adbeame3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353098118] HOST ALERT: nycnn-adbeame3;DOWN;SOFT;2;(Host Check Timed Out)
[1353098158] SERVICE ALERT: nycnn-adbeame3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1353098168] HOST ALERT: nycnn-adbeame3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1353099108] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353099168] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353100968] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353101028] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353102828] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353102888] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353104688] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 349 processes with STATE = RSZDT
[1353104748] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 60 processes with STATE = RSZDT
[1353106548] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 354 processes with STATE = RSZDT
[1353106608] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1353108408] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 354 processes with STATE = RSZDT
[1353108468] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 55 processes with STATE = RSZDT
[1353110268] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 349 processes with STATE = RSZDT
[1353110328] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1353112128] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 349 processes with STATE = RSZDT
[1353112188] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 55 processes with STATE = RSZDT
[1353113988] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 349 processes with STATE = RSZDT
[1353114048] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 150 processes with STATE = RSZDT
[1353115848] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 354 processes with STATE = RSZDT
[1353115908] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1353117708] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 353 processes with STATE = RSZDT
[1353117768] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 147 processes with STATE = RSZDT
[1353118488] SERVICE ALERT: wdcps002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353118528] HOST ALERT: wdcps002;DOWN;SOFT;1;(Host Check Timed Out)
[1353118538] SERVICE ALERT: wdcps002;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.63 ms
[1353118548] HOST ALERT: wdcps002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.56 ms
[1353119568] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353119628] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353121428] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353121488] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 59 processes with STATE = RSZDT
[1353123288] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353123348] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 148 processes with STATE = RSZDT
[1353125148] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 355 processes with STATE = RSZDT
[1353125208] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353127008] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 351 processes with STATE = RSZDT
[1353127068] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 144 processes with STATE = RSZDT