[1415601966] SERVICE ALERT: nycnn-ipvx25;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1415602026] SERVICE ALERT: nycnn-ipvx25;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.07%
[1415602176] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1415602236] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.04%
[1415603246] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;WARNING;SOFT;1;WARNING: at least one queue with depth 59, all MS2_EVENT queues in RUNNING state
[1415603306] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;SOFT;2;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1415632616] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1415632656] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1415632676] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1415632756] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1415632856] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1415632906] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1415632916] HOST ALERT: omnmas1app2-ny;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1415633926] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1415633966] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1415633986] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1415634066] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1415634166] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1415634226] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1415634236] HOST ALERT: omnmas1app1-ny;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1415651894] SERVICE ALERT: spider-tmp;/nymoneygfx Free Space;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1415652014] SERVICE ALERT: spider-tmp;/nymoneygfx Free Space;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1415652134] SERVICE ALERT: spider-tmp;/nymoneygfx Free Space;CRITICAL;HARD;3;(Return code of 127 is out of bounds - plugin may be missing)
[1415652962] SERVICE ALERT: spider-tmp;/nymoneygfx Free Space;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1415653082] SERVICE ALERT: spider-tmp;/nymoneygfx Free Space;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1415653202] SERVICE ALERT: spider-tmp;/nymoneygfx Free Space;CRITICAL;HARD;3;(Return code of 127 is out of bounds - plugin may be missing)