[1406438731] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1406438791] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1406438941] SERVICE ALERT: nycnn-ipvx18;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1406439001] SERVICE ALERT: nycnn-ipvx18;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.04%
[1406459761] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1406459821] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1406459881] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1406459951] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1406459951] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1406459961] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1406459961] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1406459981] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;1;(Host Check Timed Out)
[1406459991] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;1;(Host Check Timed Out)
[1406459991] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;1;(Host Check Timed Out)
[1406460011] SERVICE ALERT: nycnn-ms2fsd1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1406460011] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1406460021] SERVICE ALERT: nycnn-ms2fsd;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1406460021] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1406460071] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1406460081] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;2;(Host Check Timed Out)
[1406460081] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;2;(Host Check Timed Out)
[1406460081] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;2;(Host Check Timed Out)
[1406460171] HOST ALERT: nycnn-ms2fsd2;UP;SOFT;3;PING OK - Packet loss = 78%, RTA = 0.83 ms
[1406460171] HOST ALERT: nycnn-ms2fsd1;DOWN;SOFT;3;(Host Check Timed Out)
[1406460171] HOST ALERT: nycnn-ms2fsd;DOWN;SOFT;3;(Host Check Timed Out)
[1406460201] SERVICE ALERT: nycnn-ms2fsd2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1406460241] HOST ALERT: nycnn-ms2fsd1;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1406460241] HOST ALERT: nycnn-ms2fsd;UP;SOFT;4;PING OK - Packet loss = 28%, RTA = 0.23 ms
[1406460241] SERVICE ALERT: nycnn-ms2fsd1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1406460251] SERVICE ALERT: nycnn-ms2fsd;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1406460301] SERVICE ALERT: msny;MS2 - Grid A Mount;UNKNOWN;HARD;3;Remote command execution failed: Connection to 10.164.219.196 closed by remote host.
[1406460301] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;UNKNOWN;SOFT;1;Remote command execution failed: Connection to 10.164.219.196 closed by remote host.
[1406460371] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1406460371] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1406460371] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1406460371] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;SOFT;2;CRITICAL - Plugin timed out after 10 seconds
[1406460371] SERVICE ALERT: msny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1406460381] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;1;(Host Check Timed Out)
[1406460401] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;1;(Host Check Timed Out)
[1406460401] HOST ALERT: msny;DOWN;SOFT;1;(Host Check Timed Out)
[1406460431] SERVICE ALERT: nycnn-ms2app3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1406460431] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1406460431] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1406460431] SERVICE ALERT: msny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1406460481] HOST ALERT: nycnn-ms2app4;DOWN;SOFT;2;(Host Check Timed Out)
[1406460501] HOST ALERT: nycnn-ms2app3;DOWN;SOFT;2;(Host Check Timed Out)
[1406460501] HOST ALERT: msny;DOWN;SOFT;2;(Host Check Timed Out)
[1406460541] SERVICE ALERT: nycnn-ms2app3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1406460551] HOST ALERT: nycnn-ms2app3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1406460551] HOST ALERT: nycnn-ms2app4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1406460591] HOST ALERT: msny;DOWN;SOFT;3;(Host Check Timed Out)
[1406460661] HOST ALERT: msny;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1406460661] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;2;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1406460661] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;HARD;3;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state
[1406460661] SERVICE ALERT: msny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1406460841] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;HARD;3;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1406460841] SERVICE FLAPPING ALERT: msny;MS2 - Grid A Mount;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1406461021] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.082 second response time
[1406461561] SERVICE FLAPPING ALERT: msny;MS2 - Grid A Mount;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1406465051] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1406465071] SERVICE ALERT: nycnn-ipvx35;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1406465111] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.02%
[1406465131] SERVICE ALERT: nycnn-ipvx35;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.02%
[1406513391] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1406513451] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.03%