[1402054456] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports that the encode failed.
[1402054516] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports that the encode failed.
[1402057816] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;HARD;2;AME is idle:
[1402059456] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports that the encode failed.
[1402059516] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports that the encode failed.
[1402060116] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;HARD;2;AME is idle:
[1402061506] SERVICE ALERT: nycnn-adbeame8;AME Job status;CRITICAL;SOFT;1;Job 10195678 remains in a queued state
[1402061566] SERVICE ALERT: nycnn-adbeame8;AME Job status;OK;SOFT;2;Job 10195678 progress is  32.39 :
[1402069816] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402069876] SERVICE ALERT: nycnn-ipvx18e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 9.56%
[1402069926] SERVICE ALERT: nycnn-ipvx18;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402069986] SERVICE ALERT: nycnn-ipvx18;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 9.54%
[1402096226] SERVICE ALERT: nycnn-ipvx35;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402096286] SERVICE ALERT: nycnn-ipvx35;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1402096336] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1402096396] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.04%