[1383828267] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job 09681223 remains in a queued state
[1383828327] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;SOFT;2;AME is idle:
[1383835590] SERVICE ALERT: nycnn-ipvapilic;SVServer;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1383835650] SERVICE ALERT: nycnn-ipvapilic;SVServer;CRITICAL;HARD;2;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "SVServer.exe" running (0 excluded).
[1383841246] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;CRITICAL;SOFT;1;Connection refused
[1383841246] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;CRITICAL;SOFT;1;Connection refused
[1383841366] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;CRITICAL;HARD;2;Connection refused
[1383841366] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;CRITICAL;HARD;2;Connection refused
[1383841406] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;CRITICAL;SOFT;1;Connection refused
[1383841416] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;CRITICAL;SOFT;1;Connection refused
[1383842271] SERVICE ALERT: nycnn-ipvapilic;MMSService;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "MMSService.exe" running (0 excluded).
[1383842271] SERVICE ALERT: nycnn-ipvapilic;SSH Service;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "sshserver.exe" running (0 excluded).
[1383842914] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383842914] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383842934] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;SOFT;1;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383842974] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383842974] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383842994] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;SOFT;2;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383843034] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;HARD;3;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383843034] SERVICE ALERT: nycnn-ms2db2;Disk Space on /oracle;UNKNOWN;HARD;3;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383843044] SERVICE ALERT: nycnn-ipvx1;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.212.50 port 22: Connection refused
[1383843044] SERVICE ALERT: nycnn-ipvx10e1;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.212.179 port 22: Connection refused
[1383843044] SERVICE ALERT: nycnn-ipvx11e1;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.212.180 port 22: Connection refused
[1383843044] SERVICE ALERT: nycnn-ipvx12e1;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.212.181 port 22: Connection refused
[1383843044] SERVICE ALERT: nycnn-ipvx13e1;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.212.182 port 22: Connection refused
[1383843044] SERVICE ALERT: nycnn-ipvx15;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.212.64 port 22: Connection refused
[1383843044] SERVICE ALERT: nycnn-ipvx16;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.212.65 port 22: Connection refused
[1383843062] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;HARD;3;Remote command execution failed: /usr/local/nagios/libexec/check_disk: option requires an argument -- p
[1383843349] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;SOFT;1;Remote command execution failed: ssh: connect to host 10.164.213.70 port 22: Connection refused
[1383843409] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;HARD;2;Remote command execution failed: ssh: connect to host 10.164.213.70 port 22: Connection refused
[1383844069] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383844109] HOST ALERT: omnspsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1383844129] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383844199] HOST ALERT: omnspsm-ny;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.41 ms
[1383844239] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1383846192] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383846212] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1383846232] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1383846252] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383846272] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1383846332] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1383846427] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1383846447] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1383846497] HOST ALERT: omnmas1app2-ny;UP;SOFT;3;PING OK - Packet loss = 16%, RTA = 0.74 ms
[1383846547] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1383846567] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;OK;HARD;2;OK - C: Total=68.33GB, Used=27.30GB (40.0%), Free=41.02GB (60.0%)
[1383847395] SERVICE ALERT: nycnn-adbeame9;AME Job status;CRITICAL;SOFT;1;Job 09681967 remains in a queued state
[1383847405] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1383847435] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1383847445] SERVICE ALERT: nycnn-adbeame9;AME Job status;OK;SOFT;2;AME is idle:
[1383847465] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1383847475] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;WARNING;SOFT;1;WARNING - [Triggered by _Used%>80] - New Volume Total=1397.19GB, Used=1135.05GB (81.2%), Free=262.14GB (18.8%)
[1383847475] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;WARNING;SOFT;1;WARNING - [Triggered by _Used%>80] - New Volume Total=1397.19GB, Used=1127.16GB (80.7%), Free=270.03GB (19.3%)
[1383847525] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1383847535] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;WARNING;HARD;2;WARNING - [Triggered by _Used%>80] - New Volume Total=1397.19GB, Used=1127.18GB (80.7%), Free=270.01GB (19.3%)
[1383847535] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;WARNING;HARD;2;WARNING - [Triggered by _Used%>80] - New Volume Total=1397.19GB, Used=1135.07GB (81.2%), Free=262.12GB (18.8%)
[1383847625] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1383847645] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;WARNING;SOFT;1;WARNING - [Triggered by _Used%>80] - Proxy Total=1397.19GB, Used=1132.36GB (81.0%), Free=264.83GB (19.0%)
[1383847645] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;WARNING;SOFT;1;WARNING - [Triggered by _Used%>80] - Proxy Total=1397.19GB, Used=1134.98GB (81.2%), Free=262.21GB (18.8%)
[1383847715] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1383847814] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1383847824] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1383847914] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1383847944] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1383847954] HOST ALERT: omnmas1app1-ny;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1383848654] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job 09681981 remains in a queued state
[1383848714] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;SOFT;2;AME is idle:
[1383851301] SERVICE ALERT: nycnn-adbeame1;AME Job status;CRITICAL;SOFT;1;Job 09682031 remains in a queued state
[1383851361] SERVICE ALERT: nycnn-adbeame1;AME Job status;OK;SOFT;2;Job 09682031 progress is  25.15 :
[1383867629] SERVICE ALERT: omnmdk07h0-ny;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 44%, RTA = 0.50 ms
[1383867689] SERVICE ALERT: omnmdk07h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1383870609] SERVICE ALERT: nycnn-adbeame5;AME Job status;CRITICAL;SOFT;1;Job 09684103 remains in a queued state
[1383870659] SERVICE ALERT: nycnn-adbeame5;AME Job status;OK;SOFT;2;AME is idle:
[1383883439] SERVICE ALERT: nycnn-ipvstrm1;MMSService;CRITICAL;SOFT;1;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "MMSService.exe" running (0 excluded).
[1383883499] SERVICE ALERT: nycnn-ipvstrm1;MMSService;CRITICAL;HARD;2;CRITICAL - [Triggered by _ItemCount<1] - Found 0 Instance(s) of "MMSService.exe" running (0 excluded).