[1352955888] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1352955948] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 141 processes with STATE = RSZDT
[1352957268] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1352957298] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1352957328] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352957358] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1352957458] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1352957528] HOST ALERT: omnmas1app1-ny;UP;SOFT;4;PING OK - Packet loss = 50%, RTA = 0.34 ms
[1352957558] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1352957748] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1352957808] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352959608] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352959668] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352961468] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352961528] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352963328] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352963388] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 136 processes with STATE = RSZDT
[1352965188] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352965248] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 141 processes with STATE = RSZDT
[1352967048] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1352967108] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352968908] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352968968] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352977368] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1352977428] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 152 processes with STATE = RSZDT
[1352979228] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1352979288] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 152 processes with STATE = RSZDT
[1352981088] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1352981148] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 61 processes with STATE = RSZDT
[1352985648] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1352985708] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 136 processes with STATE = RSZDT
[1352987508] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1352987568] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 55 processes with STATE = RSZDT
[1352989368] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1352989428] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352991048] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1352991048] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352991048] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1352991048] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1352991148] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;2;(Host Check Timed Out)
[1352991148] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;2;(Host Check Timed Out)
[1352991158] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1352991158] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1352991168] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1352991168] HOST ALERT: nycnn-ipvx14;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1352991228] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352991288] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352991948] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 0.23 ms
[1352991948] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 80%, RTA = 0.20 ms
[1352991988] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1352991988] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1352992008] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352992008] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1352992078] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.50 ms
[1352992078] HOST FLAPPING ALERT: nycnn-ipvx14e1;STARTED; Host appears to have started flapping (20.1% change > 20.0% threshold)
[1352992078] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.54 ms
[1352992078] HOST FLAPPING ALERT: nycnn-ipvx14;STARTED; Host appears to have started flapping (20.1% change > 20.0% threshold)
[1352992118] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1352992118] SERVICE FLAPPING ALERT: nycnn-ipvx14e1;PING;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1352992118] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1352992118] SERVICE FLAPPING ALERT: nycnn-ipvx14;PING;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1352993088] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352993148] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1352993198] SERVICE FLAPPING ALERT: nycnn-ipvx14e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1352993198] HOST FLAPPING ALERT: nycnn-ipvx14e1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1352993198] SERVICE FLAPPING ALERT: nycnn-ipvx14;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1352993198] HOST FLAPPING ALERT: nycnn-ipvx14;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1352994948] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 346 processes with STATE = RSZDT
[1352995008] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352996808] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1352996868] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1352998668] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1352998728] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353000528] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1353000588] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353002388] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1353002448] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1353003648] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.26 ms
[1353003648] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.28 ms
[1353003678] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353003678] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1353003708] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353003708] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353003768] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353003768] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;2;(Host Check Timed Out)
[1353003818] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353003818] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353003828] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353003828] HOST ALERT: nycnn-ipvx19;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353003948] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353003948] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353003988] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353003988] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1353004008] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353004008] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353004048] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353004048] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;2;(Host Check Timed Out)
[1353004058] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1353004058] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353004068] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1353004068] HOST ALERT: nycnn-ipvx20;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1353004248] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 347 processes with STATE = RSZDT
[1353004308] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 138 processes with STATE = RSZDT
[1353004728] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353004728] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353004768] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353004768] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1353004788] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353004788] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353004838] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353004838] HOST ALERT: nycnn-ipvx19;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353004838] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353004838] SERVICE FLAPPING ALERT: nycnn-ipvx19e1;PING;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1353004838] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353004838] SERVICE FLAPPING ALERT: nycnn-ipvx19;PING;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1353005028] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353005028] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353005058] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353005058] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1353005088] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353005088] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353005128] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353005128] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353005138] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353005138] SERVICE FLAPPING ALERT: nycnn-ipvx20e1;PING;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1353005138] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353005138] SERVICE FLAPPING ALERT: nycnn-ipvx20;PING;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1353005388] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353005388] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353005418] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353005418] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1353005448] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353005448] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353005518] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353005518] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;2;(Host Check Timed Out)
[1353005558] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353005558] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353005568] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1353005568] HOST ALERT: nycnn-ipvx21;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1353005978] SERVICE FLAPPING ALERT: nycnn-ipvx19e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353005978] SERVICE FLAPPING ALERT: nycnn-ipvx19;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353005988] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353005988] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353006028] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1353006028] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353006048] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353006048] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353006108] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 352 processes with STATE = RSZDT
[1353006118] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.54 ms
[1353006118] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.46 ms
[1353006158] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353006158] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353006168] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 136 processes with STATE = RSZDT
[1353006278] SERVICE FLAPPING ALERT: nycnn-ipvx20e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353006278] SERVICE FLAPPING ALERT: nycnn-ipvx20;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353006408] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353006408] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353006438] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353006438] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1353006468] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353006468] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353006518] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.49 ms
[1353006518] HOST FLAPPING ALERT: nycnn-ipvx21e1;STARTED; Host appears to have started flapping (20.3% change > 20.0% threshold)
[1353006518] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.46 ms
[1353006518] HOST FLAPPING ALERT: nycnn-ipvx21;STARTED; Host appears to have started flapping (20.3% change > 20.0% threshold)
[1353006518] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1353006518] SERVICE FLAPPING ALERT: nycnn-ipvx21e1;PING;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1353006518] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353006518] SERVICE FLAPPING ALERT: nycnn-ipvx21;PING;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1353007188] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353007188] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353007228] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1353007228] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353007248] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353007248] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353007298] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1353007298] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1353007298] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353007298] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353007368] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353007368] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353007398] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353007398] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1353007428] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353007428] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353007468] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;2;(Host Check Timed Out)
[1353007468] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353007488] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.84 ms
[1353007488] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.46 ms
[1353007498] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1353007498] HOST ALERT: nycnn-ipvx23;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353007658] SERVICE FLAPPING ALERT: nycnn-ipvx21e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353007658] HOST FLAPPING ALERT: nycnn-ipvx21e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353007658] SERVICE FLAPPING ALERT: nycnn-ipvx21;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353007658] HOST FLAPPING ALERT: nycnn-ipvx21;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353007668] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353007668] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353007708] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353007708] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1353007728] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353007728] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353007788] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;2;(Host Check Timed Out)
[1353007788] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353007838] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353007838] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353007848] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353007848] HOST ALERT: nycnn-ipvx24;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353007968] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1353008028] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353008328] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353008328] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353008358] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1353008358] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353008388] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353008388] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353008438] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.34 ms
[1353008438] HOST FLAPPING ALERT: nycnn-ipvx23;STARTED; Host appears to have started flapping (20.4% change > 20.0% threshold)
[1353008438] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.29 ms
[1353008438] HOST FLAPPING ALERT: nycnn-ipvx23e1;STARTED; Host appears to have started flapping (20.4% change > 20.0% threshold)
[1353008438] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1353008438] SERVICE FLAPPING ALERT: nycnn-ipvx23;PING;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1353008438] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353008438] SERVICE FLAPPING ALERT: nycnn-ipvx23e1;PING;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1353008988] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.17 ms
[1353008988] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.30 ms
[1353009018] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353009018] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1353009048] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353009048] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353009118] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353009118] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;2;(Host Check Timed Out)
[1353009158] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353009158] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353009168] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1353009168] HOST ALERT: nycnn-ipvx24;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353009578] SERVICE FLAPPING ALERT: nycnn-ipvx23;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353009578] HOST FLAPPING ALERT: nycnn-ipvx23;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353009578] SERVICE FLAPPING ALERT: nycnn-ipvx23e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353009578] HOST FLAPPING ALERT: nycnn-ipvx23e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353009828] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1353009888] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 139 processes with STATE = RSZDT
[1353014148] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353014148] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353014178] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1353014178] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353014208] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353014208] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353014238] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1353014248] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353014258] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353014258] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353014448] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353014448] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353014478] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1353014478] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353014508] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353014508] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353014548] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1353014548] HOST FLAPPING ALERT: nycnn-ipvx1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1353014548] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353014548] HOST FLAPPING ALERT: nycnn-ipvx1e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1353014558] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353014558] SERVICE FLAPPING ALERT: nycnn-ipvx1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1353014558] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353014558] SERVICE FLAPPING ALERT: nycnn-ipvx1e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1353015168] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353015168] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353015198] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353015198] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1353015228] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353015228] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353015268] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1353015268] HOST ALERT: nycnn-ipvx2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1353015278] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353015278] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353015468] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353015468] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353015508] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353015508] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1353015528] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353015528] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353015578] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1353015578] HOST FLAPPING ALERT: nycnn-ipvx2e1;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1353015578] HOST ALERT: nycnn-ipvx2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1353015578] HOST FLAPPING ALERT: nycnn-ipvx2;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1353015578] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353015578] SERVICE FLAPPING ALERT: nycnn-ipvx2e1;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1353015578] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1353015578] SERVICE FLAPPING ALERT: nycnn-ipvx2;PING;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1353015698] SERVICE FLAPPING ALERT: nycnn-ipvx1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353015698] HOST FLAPPING ALERT: nycnn-ipvx1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353015698] SERVICE FLAPPING ALERT: nycnn-ipvx1e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353015698] HOST FLAPPING ALERT: nycnn-ipvx1e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353016718] SERVICE FLAPPING ALERT: nycnn-ipvx2e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353016718] HOST FLAPPING ALERT: nycnn-ipvx2e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353016718] SERVICE FLAPPING ALERT: nycnn-ipvx2;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353016718] HOST FLAPPING ALERT: nycnn-ipvx2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353017268] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353017268] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353017308] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1353017308] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353017328] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353017328] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353017388] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.50 ms
[1353017388] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.50 ms
[1353017438] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1353017438] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353018048] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353018048] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353018078] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353018078] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1353018108] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353018108] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353018148] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.47 ms
[1353018148] HOST FLAPPING ALERT: nycnn-ipvx12;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1353018148] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.50 ms
[1353018148] HOST FLAPPING ALERT: nycnn-ipvx12e1;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1353018158] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353018158] SERVICE FLAPPING ALERT: nycnn-ipvx12e1;PING;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1353018158] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1353018158] SERVICE FLAPPING ALERT: nycnn-ipvx12;PING;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1353018348] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353018348] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353018388] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353018388] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1353018408] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353018408] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353018458] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1353018458] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1353018458] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353018458] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353018828] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353018828] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353018868] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353018868] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1353018888] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353018888] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353018948] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.37 ms
[1353018948] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.29 ms
[1353018998] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1353018998] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353019598] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1353019598] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1353019598] SERVICE FLAPPING ALERT: nycnn-ipvx12e1;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353019598] HOST FLAPPING ALERT: nycnn-ipvx12e1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353019598] SERVICE FLAPPING ALERT: nycnn-ipvx12;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353019598] HOST FLAPPING ALERT: nycnn-ipvx12;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1353019608] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1353019608] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1353019698] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;2;(Host Check Timed Out)
[1353019698] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;2;(Host Check Timed Out)
[1353019718] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353019718] SERVICE FLAPPING ALERT: nycnn-ipvx11e1;PING;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1353019718] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1353019718] SERVICE FLAPPING ALERT: nycnn-ipvx11;PING;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1353019728] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1353019728] HOST FLAPPING ALERT: nycnn-ipvx11e1;STARTED; Host appears to have started flapping (20.9% change > 20.0% threshold)
[1353019728] HOST ALERT: nycnn-ipvx11;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1353019728] HOST FLAPPING ALERT: nycnn-ipvx11;STARTED; Host appears to have started flapping (20.9% change > 20.0% threshold)
[1353020698] HOST FLAPPING ALERT: nycnn-ipvx11;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1353020698] HOST FLAPPING ALERT: nycnn-ipvx11e1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1353020798] SERVICE FLAPPING ALERT: nycnn-ipvx11e1;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353020798] SERVICE FLAPPING ALERT: nycnn-ipvx11;PING;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1353020988] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 342 processes with STATE = RSZDT
[1353021048] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 143 processes with STATE = RSZDT
[1353022848] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 345 processes with STATE = RSZDT
[1353022908] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 139 processes with STATE = RSZDT
[1353024168] SERVICE ALERT: wdcpf001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353024188] SERVICE ALERT: wdcms003;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353024198] HOST ALERT: wdcpf001;DOWN;SOFT;1;(Host Check Timed Out)
[1353024218] HOST ALERT: wdcms003;DOWN;SOFT;1;(Host Check Timed Out)
[1353024228] SERVICE ALERT: wdcpf001;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353024248] SERVICE ALERT: wdcms003;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1353024298] HOST ALERT: wdcpf001;DOWN;SOFT;2;(Host Check Timed Out)
[1353024318] HOST ALERT: wdcms003;DOWN;SOFT;2;(Host Check Timed Out)
[1353024368] HOST ALERT: wdcpf001;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 6.72 ms
[1353024388] HOST ALERT: wdcms003;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 6.53 ms
[1353024398] SERVICE ALERT: wdcpf001;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.51 ms
[1353024418] SERVICE ALERT: wdcms003;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.61 ms
[1353024708] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 340 processes with STATE = RSZDT
[1353024768] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 136 processes with STATE = RSZDT
[1353024828] SERVICE ALERT: wdcps002;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353024858] HOST ALERT: wdcps002;DOWN;SOFT;1;(Host Check Timed Out)
[1353024888] SERVICE ALERT: wdcps002;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1353024958] HOST ALERT: wdcps002;DOWN;SOFT;2;(Host Check Timed Out)
[1353025058] HOST ALERT: wdcps002;DOWN;SOFT;3;(Host Check Timed Out)
[1353025148] HOST ALERT: wdcps002;DOWN;SOFT;4;(Host Check Timed Out)
[1353025238] HOST ALERT: wdcps002;DOWN;SOFT;5;(Host Check Timed Out)
[1353025338] HOST ALERT: wdcps002;DOWN;SOFT;6;(Host Check Timed Out)
[1353025438] HOST ALERT: wdcps002;DOWN;SOFT;7;(Host Check Timed Out)
[1353025478] SERVICE ALERT: wdcps002;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 6.56 ms
[1353025488] HOST ALERT: wdcps002;UP;SOFT;8;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1353026568] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 341 processes with STATE = RSZDT
[1353026628] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 137 processes with STATE = RSZDT
[1353028428] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 323 processes with STATE = RSZDT
[1353028488] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 140 processes with STATE = RSZDT
[1353030288] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 344 processes with STATE = RSZDT
[1353030348] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 145 processes with STATE = RSZDT
[1353032148] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 349 processes with STATE = RSZDT
[1353032208] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 141 processes with STATE = RSZDT
[1353034008] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1353034068] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1353035868] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1353035928] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 139 processes with STATE = RSZDT
[1353037728] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 333 processes with STATE = RSZDT
[1353037788] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 142 processes with STATE = RSZDT
[1353039588] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 343 processes with STATE = RSZDT
[1353039648] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 63 processes with STATE = RSZDT
[1353040208] SERVICE ALERT: wdcms001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353040248] HOST ALERT: wdcms002;DOWN;SOFT;1;(Host Check Timed Out)
[1353040258] SERVICE ALERT: wdcms001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.57 ms
[1353040268] SERVICE ALERT: wdcms002;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1353040318] HOST ALERT: wdcms002;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.62 ms
[1353040318] SERVICE ALERT: wdcms002;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 6.65 ms
[1353040328] SERVICE ALERT: wdcod001;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1353040378] SERVICE ALERT: wdcod001;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.58 ms
[1353041448] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 350 processes with STATE = RSZDT
[1353041508] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 141 processes with STATE = RSZDT