[1409841778] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;SOFT;1;Job 10419512 remains in a queued state
[1409841838] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;SOFT;2;Job 10419512 progress is  65.79 :
[1409845658] SERVICE ALERT: nycnn-ardb1;get_cnnid_count;CRITICAL;SOFT;1;(Return code of 127 is out of bounds - plugin may be missing)
[1409845718] SERVICE ALERT: nycnn-ardb1;get_cnnid_count;CRITICAL;SOFT;2;(Return code of 127 is out of bounds - plugin may be missing)
[1409845778] SERVICE ALERT: nycnn-ardb1;get_cnnid_count;CRITICAL;HARD;3;(Return code of 127 is out of bounds - plugin may be missing)
[1409849506] SERVICE ALERT: nycnn-ipvstrm3;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409849566] SERVICE ALERT: nycnn-ipvstrm4;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409849566] SERVICE ALERT: nycnn-ipvstrm3;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.37%
[1409849626] SERVICE ALERT: nycnn-ipvstrm4;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.24%
[1409849756] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409849806] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409849816] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1409849866] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.34%
[1409863616] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409863656] SERVICE ALERT: nycnn-ipvx23;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409863676] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1409863716] SERVICE ALERT: nycnn-ipvx23;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1409863976] SERVICE ALERT: nycnn-ipvapilic;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409864036] SERVICE ALERT: nycnn-ipvapilic;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.01%
[1409864046] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;SOFT;1;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1409864106] SERVICE ALERT: nycnn-adbeame6;AME Job status;CRITICAL;HARD;2;Job  failed with error AME reports the source as being invalid.  Please check that the source media/project is valid and/or network accessible.
[1409864706] SERVICE ALERT: nycnn-adbeame6;AME Job status;OK;HARD;2;AME is idle:
[1409865056] SERVICE ALERT: nycnn-adbeame7;AME Job status;CRITICAL;SOFT;1;Job 10421396 remains in a queued state
[1409865116] SERVICE ALERT: nycnn-adbeame7;AME Job status;OK;SOFT;2;Job 10421396 progress is  24 :
[1409869498] HOST DOWNTIME ALERT: nycnn-ipvx30;STARTED; Host has entered a period of scheduled downtime
[1409869536] SERVICE ALERT: nycnn-ipvx30;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1409869536] SERVICE ALERT: nycnn-ipvx30e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1409869536] SERVICE ALERT: nycnn-ipvx30;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1409869566] SERVICE ALERT: nycnn-ipvx30;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1409869566] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1409869566] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1409869596] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1409869596] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1409869606] SERVICE ALERT: nycnn-ipvx30;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1409869616] SERVICE ALERT: nycnn-ipvx30e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1409869616] SERVICE ALERT: nycnn-ipvx30;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1409869626] SERVICE ALERT: nycnn-ipvx30;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1409869626] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1409869626] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1409869666] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1409869686] SERVICE ALERT: nycnn-ipvx30e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1409869686] SERVICE ALERT: nycnn-ipvx30;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1409869686] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.45 ms
[1409869686] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.47 ms
[1409869686] SERVICE ALERT: nycnn-ipvx30e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 54%, RTA = 0.32 ms
[1409869686] SERVICE ALERT: nycnn-ipvx30;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 54%, RTA = 0.30 ms
[1409869746] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1409869746] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1409869896] SERVICE ALERT: nycnn-ipvx30;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.800GB (40%) - Free: 1.190GB (60%)
[1409869896] SERVICE ALERT: nycnn-ipvx30e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1409869896] SERVICE ALERT: nycnn-ipvx30;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1409869906] SERVICE ALERT: nycnn-ipvx30;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1409869956] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1409869986] SERVICE ALERT: nycnn-ipvx30e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1409869986] SERVICE ALERT: nycnn-ipvx30;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1409869986] SERVICE ALERT: nycnn-ipvx12e1;Memory Usage;WARNING;SOFT;1;WARNING - [Triggered by _MemUsed%>90] - Physical Memory: Total: 1.990GB - Used: 1.862GB (94%) - Free: 0.128GB (6%)
[1409870046] SERVICE ALERT: nycnn-ipvx12e1;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 1.990GB - Used: 0.878GB (44%) - Free: 1.112GB (56%)
[1409870667] HOST DOWNTIME ALERT: nycnn-ipvx30;STOPPED; Host has exited from a period of scheduled downtime
[1409878256] SERVICE ALERT: nycnn-ipvx2;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409878316] SERVICE ALERT: nycnn-ipvx2;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 9.84%
[1409878386] SERVICE ALERT: nycnn-ipvx2e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1409878446] SERVICE ALERT: nycnn-ipvx2e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 9.24%