[1413544366] SERVICE ALERT: nycnn-nag1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413544406] HOST ALERT: nycnn-nag1;DOWN;SOFT;1;(Host Check Timed Out)
[1413544426] SERVICE ALERT: nycnn-nag1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413544506] HOST ALERT: nycnn-nag1;DOWN;SOFT;2;(Host Check Timed Out)
[1413544606] HOST ALERT: nycnn-nag1;DOWN;SOFT;3;(Host Check Timed Out)
[1413544696] HOST ALERT: nycnn-nag1;DOWN;SOFT;4;(Host Check Timed Out)
[1413544796] HOST ALERT: nycnn-nag1;DOWN;SOFT;5;(Host Check Timed Out)
[1413544896] HOST ALERT: nycnn-nag1;DOWN;SOFT;6;(Host Check Timed Out)
[1413544996] HOST ALERT: nycnn-nag1;DOWN;SOFT;7;(Host Check Timed Out)
[1413545096] HOST ALERT: nycnn-nag1;DOWN;SOFT;8;(Host Check Timed Out)
[1413545196] HOST ALERT: nycnn-nag1;DOWN;SOFT;9;(Host Check Timed Out)
[1413545286] HOST ALERT: nycnn-nag1;DOWN;HARD;10;(Host Check Timed Out)
[1413556846] SERVICE ALERT: nycnn-nag1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1413556856] HOST ALERT: nycnn-nag1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1413569407] HOST DOWNTIME ALERT: nycnn-ipvsphn1;STARTED; Host has entered a period of scheduled downtime
[1413569424] HOST DOWNTIME ALERT: nycnn-ipvsphn2;STARTED; Host has entered a period of scheduled downtime
[1413569437] HOST DOWNTIME ALERT: nycnn-ipvsphn3;STARTED; Host has entered a period of scheduled downtime
[1413569451] HOST DOWNTIME ALERT: nycnn-ipvsphn4;STARTED; Host has entered a period of scheduled downtime
[1413569464] HOST DOWNTIME ALERT: nycnn-ipvsphn5;STARTED; Host has entered a period of scheduled downtime
[1413569475] HOST DOWNTIME ALERT: nycnn-ipvstrm1;STARTED; Host has entered a period of scheduled downtime
[1413569488] HOST DOWNTIME ALERT: nycnn-ipvstrm2;STARTED; Host has entered a period of scheduled downtime
[1413569505] HOST DOWNTIME ALERT: nycnn-ipvstrm3;STARTED; Host has entered a period of scheduled downtime
[1413569518] HOST DOWNTIME ALERT: nycnn-ipvstrm4;STARTED; Host has entered a period of scheduled downtime
[1413569986] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413569986] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413569986] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413569996] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570006] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570016] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1413570046] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570046] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570046] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413570046] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413570056] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570056] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570066] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1413570066] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570066] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570066] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570076] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570086] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1413570106] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413570116] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570126] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570126] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570136] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570136] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570136] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570146] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570146] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570146] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING OK - Packet loss = 70%, RTA = 0.50 ms
[1413570146] SERVICE ALERT: nycnn-ipvstrm1;Process: MMSService;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413570146] SERVICE ALERT: nycnn-ipvstrm1;Process: IPV_FPT;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1413570166] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1413570166] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570186] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1413570206] SERVICE ALERT: nycnn-ipvstrm1;Process: MMSService;OK;SOFT;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1413570206] SERVICE ALERT: nycnn-ipvstrm1;Process: IPV_FPT;OK;SOFT;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1413570256] SERVICE ALERT: nycnn-ipvstrm4;Process: MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570266] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING OK - Packet loss = 75%, RTA = 0.26 ms
[1413570286] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413570326] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413570336] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;OK;HARD;2;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1413570346] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;OK;HARD;2;OK - New Volume Total=1397.19GB, Used=1133.14GB (81.1%), Free=264.05GB (18.9%)
[1413570346] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1413570396] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1413570406] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;2;OK - Proxy Total=1397.19GB, Used=1133.38GB (81.1%), Free=263.81GB (18.9%)
[1413570416] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;OK;HARD;2;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1413570426] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1413570426] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1413570426] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1413570426] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1413570426] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1133.38GB (81.1%), Free=263.81GB (18.9%)
[1413570446] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.798GB (20%) - Free: 3.202GB (80%)
[1413570536] SERVICE ALERT: nycnn-ipvstrm4;Process: MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1413570736] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570766] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413570766] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570776] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1413570796] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570826] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413570826] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413570836] SERVICE ALERT: nycnn-ipvsphn1;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413570876] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;2;(Host Check Timed Out)
[1413570886] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1413570896] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1413571076] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1413571116] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413571116] SERVICE ALERT: nycnn-ipvsphn1;C:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=1.62GB (1.2%), Free=135.07GB (98.8%)
[1413571186] SERVICE ALERT: nycnn-ipvsphn5;C:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413571206] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413571216] HOST ALERT: nycnn-ipvsphn5;DOWN;SOFT;1;(Host Check Timed Out)
[1413571236] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1413571246] SERVICE ALERT: nycnn-ipvsphn3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413571256] SERVICE ALERT: nycnn-ipvsphn3;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571256] SERVICE ALERT: nycnn-ipvsphn5;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571256] SERVICE ALERT: nycnn-ipvsphn5;C:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413571266] SERVICE ALERT: nycnn-ipvsphn3;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571266] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413571286] HOST ALERT: nycnn-ipvsphn3;DOWN;SOFT;1;(Host Check Timed Out)
[1413571296] SERVICE ALERT: nycnn-ipvsphn4;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571306] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413571306] SERVICE ALERT: nycnn-ipvsphn3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413571316] HOST ALERT: nycnn-ipvsphn5;DOWN;SOFT;2;(Host Check Timed Out)
[1413571316] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1413571316] SERVICE ALERT: nycnn-ipvsphn3;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413571326] SERVICE ALERT: nycnn-ipvsphn3;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413571326] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;1;(Host Check Timed Out)
[1413571336] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571336] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571346] SERVICE ALERT: nycnn-ipvsphn5;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571356] SERVICE ALERT: nycnn-ipvsphn4;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1413571366] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413571376] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1413571386] HOST ALERT: nycnn-ipvsphn3;DOWN;SOFT;2;(Host Check Timed Out)
[1413571386] SERVICE ALERT: nycnn-ipvsphn5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1413571386] HOST ALERT: nycnn-ipvsphn5;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413571396] SERVICE ALERT: nycnn-ipvsphn4;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571416] SERVICE ALERT: nycnn-ipvsphn3;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571416] SERVICE ALERT: nycnn-ipvsphn4;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1413571426] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;2;(Host Check Timed Out)
[1413571426] SERVICE ALERT: nycnn-ipvsphn4;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.27 ms
[1413571436] HOST ALERT: nycnn-ipvsphn4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413571486] HOST ALERT: nycnn-ipvsphn3;DOWN;SOFT;3;(Host Check Timed Out)
[1413571486] SERVICE ALERT: nycnn-ipvsphn3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1413571496] HOST ALERT: nycnn-ipvsphn3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1413571536] SERVICE ALERT: nycnn-ipvsphn5;D:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.63GB (100.0%)
[1413571546] SERVICE ALERT: nycnn-ipvsphn5;C:\ Drive Space;OK;HARD;2;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.63GB (100.0%)
[1413571596] SERVICE ALERT: nycnn-ipvsphn3;D:\ Drive Space;OK;HARD;2;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.62GB (100.0%)
[1413571606] SERVICE ALERT: nycnn-ipvsphn3;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413571616] SERVICE ALERT: nycnn-ipvsphn3;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413571626] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1413571626] SERVICE ALERT: nycnn-ipvsphn5;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.601GB (15%) - Free: 3.398GB (85%)
[1413571636] SERVICE ALERT: nycnn-ipvsphn4;D:\ Drive Space;OK;HARD;2;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.62GB (100.0%)
[1413571656] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;OK;HARD;3;OK - Physical Memory: Total: 4.000GB - Used: 0.588GB (15%) - Free: 3.412GB (85%)
[1413571676] SERVICE ALERT: nycnn-ipvsphn4;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413571706] SERVICE ALERT: nycnn-ipvsphn3;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.629GB (16%) - Free: 3.370GB (84%)
[1413571706] SERVICE ALERT: nycnn-ipvsphn4;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1413573836] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1413573876] HOST ALERT: omnspsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1413573896] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1413573966] HOST ALERT: omnspsm-ny;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.46 ms
[1413574006] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1413576589] HOST DOWNTIME ALERT: nycnn-ipvsphn1;STOPPED; Host has exited from a period of scheduled downtime
[1413576619] HOST DOWNTIME ALERT: nycnn-ipvsphn2;STOPPED; Host has exited from a period of scheduled downtime
[1413576633] HOST DOWNTIME ALERT: nycnn-ipvsphn3;STOPPED; Host has exited from a period of scheduled downtime
[1413576647] HOST DOWNTIME ALERT: nycnn-ipvsphn4;STOPPED; Host has exited from a period of scheduled downtime
[1413576659] HOST DOWNTIME ALERT: nycnn-ipvsphn5;STOPPED; Host has exited from a period of scheduled downtime
[1413576671] HOST DOWNTIME ALERT: nycnn-ipvstrm1;STOPPED; Host has exited from a period of scheduled downtime
[1413576684] HOST DOWNTIME ALERT: nycnn-ipvstrm2;STOPPED; Host has exited from a period of scheduled downtime
[1413576698] HOST DOWNTIME ALERT: nycnn-ipvstrm3;STOPPED; Host has exited from a period of scheduled downtime
[1413576714] HOST DOWNTIME ALERT: nycnn-ipvstrm4;STOPPED; Host has exited from a period of scheduled downtime
[1413588816] SERVICE ALERT: nycnn-adbeame3;AME Job status;CRITICAL;SOFT;1;Job 10529071 remains in a queued state
[1413588876] SERVICE ALERT: nycnn-adbeame3;AME Job status;OK;SOFT;2;AME is idle: