[1385272758] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 132.70 ms
[1385272768] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.82 ms
[1385272818] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 112.05 ms
[1385272828] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 10.84 ms
[1385272878] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 17.74 ms
[1385276708] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276718] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276738] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.21 ms
[1385276738] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.24 ms
[1385276738] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276738] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276738] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276738] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276748] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276748] SERVICE ALERT: nycnn-ipvx24;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276758] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276758] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276758] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276758] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276758] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276758] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276768] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276768] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276768] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276768] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276768] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276768] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276768] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276778] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1385276778] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276778] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1385276778] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx15;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] SERVICE ALERT: nycnn-ipvx33;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] SERVICE ALERT: nycnn-ipvx15e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276778] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276778] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1385276778] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276778] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx4;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276788] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276788] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276788] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx5;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276788] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.22 ms
[1385276788] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276788] SERVICE ALERT: nycnn-ipvx2;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276788] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276788] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1385276798] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276798] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276798] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276798] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276798] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1385276798] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.20 ms
[1385276808] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276808] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1385276808] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276808] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276808] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276808] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276808] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276808] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276808] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276808] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276808] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1385276808] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1385276818] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276818] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276818] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276818] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276818] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276818] SERVICE ALERT: nycnn-ipvx27;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276828] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1385276828] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276828] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276828] SERVICE ALERT: nycnn-ipvx24;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276828] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276828] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276828] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276828] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276828] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276828] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276828] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276828] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276828] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] SERVICE ALERT: nycnn-ipvx35;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1385276838] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx9;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1385276838] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.38 ms
[1385276838] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvsphn1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 28%, RTA = 0.26 ms
[1385276838] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx15;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx33;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx15e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276838] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276838] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276848] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276848] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;2;(Host Check Timed Out)
[1385276848] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276848] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276848] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;2;(Host Check Timed Out)
[1385276848] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;2;(Host Check Timed Out)
[1385276848] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx4;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276848] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276848] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx5;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276848] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx2;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276848] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276848] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276848] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1385276848] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276858] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276858] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276858] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1385276868] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276868] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276868] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276868] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276868] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.27 ms
[1385276868] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.51 ms
[1385276868] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276878] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276878] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1385276878] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276878] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;2;(Host Check Timed Out)
[1385276878] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.50 ms
[1385276878] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276878] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276878] SERVICE ALERT: nycnn-ipvx27;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385276878] SERVICE ALERT: nycnn-ipvx14;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.47 ms
[1385276878] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.32 ms
[1385276878] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.44 ms
[1385276888] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1385276888] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.56 ms
[1385276888] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.46 ms
[1385276888] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.44 ms
[1385276888] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.44 ms
[1385276888] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.46 ms
[1385276888] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.45 ms
[1385276888] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.28 ms
[1385276888] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.46 ms
[1385276888] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276888] HOST ALERT: nycnn-ipvx2;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.50 ms
[1385276888] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.47 ms
[1385276888] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.59 ms
[1385276888] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276888] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276898] SERVICE ALERT: nycnn-ipvx29e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 0.31 ms
[1385276898] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276898] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;2;(Host Check Timed Out)
[1385276898] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276898] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;2;(Host Check Timed Out)
[1385276898] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276898] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276898] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276898] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.23 ms
[1385276898] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1385276898] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276898] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.46 ms
[1385276898] SERVICE ALERT: nycnn-ipvx12;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276898] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276898] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276898] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276898] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1385276898] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276898] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1385276898] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276898] HOST ALERT: nycnn-ipvx19;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385276898] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1385276898] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276898] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276898] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1385276898] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276898] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276898] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1385276898] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276898] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276898] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276898] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276898] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276898] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276898] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276908] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276908] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385276908] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385276908] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1385276908] HOST ALERT: nycnn-ipvx32;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385276908] HOST ALERT: nycnn-ipvx30;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276908] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1385276908] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276908] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276908] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385276908] HOST ALERT: nycnn-ipvx28;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276908] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276908] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276908] HOST ALERT: nycnn-ipvx25;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.26 ms
[1385276908] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1385276908] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276908] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385276908] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276908] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276908] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385276908] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276908] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276908] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1385276908] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 1.05 ms
[1385276908] SERVICE ALERT: nycnn-ipvx1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276908] SERVICE ALERT: nycnn-ipvx10;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276908] SERVICE ALERT: nycnn-ipvx1e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276908] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276908] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1385276908] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276908] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1385276908] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385276908] SERVICE ALERT: nycnn-ipvx1e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.48 ms
[1385276908] SERVICE ALERT: nycnn-ipvx11;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276908] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.28 ms
[1385276908] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.28 ms
[1385276918] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.25 ms
[1385276918] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.41 ms
[1385276918] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.45 ms
[1385276918] SERVICE ALERT: nycnn-ipvx6e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 50%, RTA = 0.29 ms
[1385276918] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1385276918] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1385276918] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276918] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276918] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276918] HOST ALERT: nycnn-ipvx14;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276918] HOST ALERT: nycnn-ipvx15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276918] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276918] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1385276918] HOST ALERT: nycnn-ipvx29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276918] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1385276918] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276918] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276918] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276918] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276918] SERVICE ALERT: nycnn-ipvx20;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276918] SERVICE ALERT: nycnn-ipvx24;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.30 ms
[1385276918] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.30 ms
[1385276918] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276918] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276928] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276928] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276928] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;2;(Host Check Timed Out)
[1385276928] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;2;(Host Check Timed Out)
[1385276928] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;2;(Host Check Timed Out)
[1385276928] SERVICE ALERT: nycnn-ipvx22e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276928] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276928] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276928] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1385276928] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276928] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276928] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276928] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276938] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.44 ms
[1385276938] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276938] SERVICE ALERT: nycnn-ipvx24e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385276948] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;2;(Host Check Timed Out)
[1385276948] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;2;(Host Check Timed Out)
[1385276948] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276948] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276948] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276948] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276948] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276948] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276948] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276948] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276948] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276948] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1385276958] HOST ALERT: nycnn-ipvx27;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276958] HOST ALERT: nycnn-ipvx12;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276958] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385276958] HOST ALERT: nycnn-ipvx22;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276958] HOST ALERT: nycnn-ipvx10;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276958] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276958] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276958] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1385276958] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1385276958] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276958] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276958] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1385276958] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276968] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385276968] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276968] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1385276968] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276968] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276978] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385276978] HOST ALERT: nycnn-ipvx20;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385276978] HOST ALERT: nycnn-ipvx31;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276978] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1385276978] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1385276978] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1385276988] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385276988] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385277068] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277078] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277088] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277098] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277098] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277098] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277098] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277098] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277108] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277108] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277108] SERVICE ALERT: nycnn-ipvx24;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx35;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx9;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277118] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx15;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx15e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx33;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx4;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277128] SERVICE ALERT: nycnn-ipvx5;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277138] SERVICE ALERT: nycnn-ipvx2;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277138] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277148] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277168] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277168] SERVICE ALERT: nycnn-ipvx27;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277188] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277188] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277198] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277198] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277198] SERVICE ALERT: nycnn-ipvx1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277198] SERVICE ALERT: nycnn-ipvx10;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277198] SERVICE ALERT: nycnn-ipvx11;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277198] SERVICE ALERT: nycnn-ipvx12;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277198] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277208] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277208] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277208] SERVICE ALERT: nycnn-ipvx22e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277208] SERVICE ALERT: nycnn-ipvx20;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277208] SERVICE ALERT: nycnn-ipvx1e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385277218] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277218] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277218] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385277228] SERVICE ALERT: nycnn-ipvx24e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1385279898] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 135.20 ms
[1385279958] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 52.32 ms
[1385280018] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 102.14 ms
[1385280078] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 87.74 ms
[1385280288] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280308] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385280308] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385280308] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385280308] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385280318] SERVICE ALERT: nycnn-ipvstrm1;MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280318] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280318] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280328] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1385280348] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1385280348] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1385280348] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1385280348] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385280368] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385280368] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385280368] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385280368] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1385280378] SERVICE ALERT: nycnn-ipvstrm1;MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385280378] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385280378] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1385280428] SERVICE ALERT: nycnn-ipvstrm1;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm3;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm1;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm4;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm2;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvstrm3;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280428] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1385280428] SERVICE ALERT: nycnn-ipvstrm4;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280438] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1385280438] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1385280438] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1385280468] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;1;(Host Check Timed Out)
[1385280488] SERVICE ALERT: nycnn-ipvapilic;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1385280488] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385280488] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385280498] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1385280498] HOST ALERT: nycnn-ipvapilic;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385280498] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1385280498] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1385280508] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1385280508] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1385280548] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385280548] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1385280638] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1385280658] SERVICE ALERT: nycnn-ipvstrm1;MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1385280668] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;2;OK - Proxy Total=1397.19GB, Used=1134.79GB (81.2%), Free=262.40GB (18.8%)
[1385280668] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;OK;HARD;2;OK - New Volume Total=1397.19GB, Used=1134.73GB (81.2%), Free=262.46GB (18.8%)
[1385280708] SERVICE ALERT: nycnn-ipvstrm1;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1385280708] SERVICE ALERT: nycnn-ipvstrm2;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1385280708] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1385280708] SERVICE ALERT: nycnn-ipvstrm3;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1385280708] SERVICE ALERT: nycnn-ipvstrm1;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385280718] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1385280718] SERVICE ALERT: nycnn-ipvstrm4;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1385280718] SERVICE ALERT: nycnn-ipvstrm2;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385280718] SERVICE ALERT: nycnn-ipvstrm3;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385280718] SERVICE ALERT: nycnn-ipvstrm4;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1385280778] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1134.73GB (81.2%), Free=262.46GB (18.8%)
[1385287168] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 118.36 ms
[1385287228] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 18.10 ms
[1385287338] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 107.99 ms
[1385287398] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 67.67 ms
[1385301438] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 103.32 ms
[1385301448] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 113.60 ms
[1385301498] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 98.71 ms
[1385301508] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 96.67 ms
[1385301558] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 123.13 ms
[1385301618] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 114.79 ms
[1385301678] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 91.69 ms
[1385301738] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 119.57 ms
[1385301798] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 116.11 ms
[1385301858] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.12 ms
[1385308638] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 112.36 ms
[1385308698] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 98.79 ms
[1385308758] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 119.79 ms
[1385308818] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 58.66 ms
[1385308878] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 122.65 ms
[1385308938] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 78.43 ms
[1385308998] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 122.99 ms
[1385309058] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1385315838] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 108.61 ms
[1385315898] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 62.15 ms
[1385323098] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 116.19 ms
[1385323108] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 114.93 ms
[1385323158] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 4.13 ms
[1385323168] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 71.17 ms
[1385323278] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 111.33 ms
[1385323338] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 96.28 ms
[1385330308] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 105.14 ms
[1385330368] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 75.54 ms
[1385330418] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 111.36 ms
[1385330478] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 47.21 ms
[1385334748] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 106.34 ms
[1385334808] SERVICE ALERT: rt-nycc3750tgr-ms2;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.52 ms
[1385337568] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 101.21 ms
[1385337618] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 104.84 ms
[1385337628] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.92 ms
[1385337678] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 102.50 ms
[1385337738] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 119.25 ms
[1385337858] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1385344638] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 113.52 ms
[1385344698] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 74.29 ms
[1385344758] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 108.69 ms
[1385344768] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 120.65 ms
[1385344818] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 74.25 ms
[1385344828] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 57.82 ms
[1385344938] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 137.40 ms
[1385344998] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 119.99 ms
[1385345058] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1385351838] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 130.17 ms
[1385351848] SERVICE ALERT: rt-nyccfcp-1;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 120.66 ms
[1385351898] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 128.58 ms
[1385351908] SERVICE ALERT: rt-nyccfcp-1;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 48.81 ms
[1385351958] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 68.25 ms
[1385352078] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 113.61 ms
[1385352138] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 118.10 ms
[1385352198] SERVICE ALERT: rt-nycc3750tgr-ilo;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 68.34 ms