[1376139917] SERVICE DOWNTIME ALERT: omnmdir01-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] SERVICE DOWNTIME ALERT: omnmdir01e1-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] SERVICE DOWNTIME ALERT: omnmdir02-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] SERVICE DOWNTIME ALERT: omnmdir02e1-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] SERVICE DOWNTIME ALERT: omnmdir03-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] SERVICE DOWNTIME ALERT: omnmdir03e1-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] SERVICE DOWNTIME ALERT: omnmdir04-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] SERVICE DOWNTIME ALERT: omnmdir04e1-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir01-ny;STARTED; Host has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir01e1-ny;STARTED; Host has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir02-ny;STARTED; Host has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir02e1-ny;STARTED; Host has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir03-ny;STARTED; Host has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir03e1-ny;STARTED; Host has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir04-ny;STARTED; Host has entered a period of scheduled downtime
[1376139917] HOST DOWNTIME ALERT: omnmdir04e1-ny;STARTED; Host has entered a period of scheduled downtime
[1376139973] SERVICE DOWNTIME ALERT: omnspsm-ny;PING;STARTED; Service has entered a period of scheduled downtime
[1376140031] HOST DOWNTIME ALERT: omnspsm-ny;STARTED; Host has entered a period of scheduled downtime
[1376140658] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376140688] HOST ALERT: omnspsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376140718] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376140788] HOST ALERT: omnspsm-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1376140828] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1376140838] HOST ALERT: omnspsm-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376142398] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1376142458] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;2;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1376142518] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;3;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1376147918] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376147948] HOST ALERT: omnspsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376147978] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376148018] HOST ALERT: omnspsm-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1376148028] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376148038] HOST ALERT: omnspsm-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1376149298] SERVICE ALERT: omnmdir01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376149298] SERVICE ALERT: omnmdir01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376149328] HOST ALERT: omnmdir01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376149328] HOST ALERT: omnmdir01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376149358] SERVICE ALERT: omnmdir01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376149358] SERVICE ALERT: omnmdir01e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376149398] HOST ALERT: omnmdir01-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376149398] HOST ALERT: omnmdir01e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376149418] SERVICE ALERT: omnmdir01-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1376149418] SERVICE ALERT: omnmdir01e1-ny;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1376149538] SERVICE ALERT: omnmdir01-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376149538] SERVICE ALERT: omnmdir01e1-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1376149778] SERVICE ALERT: omnmdir02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376149778] SERVICE ALERT: omnmdir02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376149818] HOST ALERT: omnmdir02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376149818] HOST ALERT: omnmdir02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376149838] SERVICE ALERT: omnmdir02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376149838] SERVICE ALERT: omnmdir02e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376149868] HOST ALERT: omnmdir02-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1376149888] HOST ALERT: omnmdir02e1-ny;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.26 ms
[1376149898] SERVICE ALERT: omnmdir02-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376149898] SERVICE ALERT: omnmdir02e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376149908] HOST ALERT: omnmdir02-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1376149958] SERVICE ALERT: omnmdir02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376149958] SERVICE ALERT: omnmdir02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376149958] HOST ALERT: omnmdir02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376149998] HOST ALERT: omnmdir02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150018] SERVICE ALERT: omnmdir02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150018] SERVICE ALERT: omnmdir02e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150028] HOST ALERT: omnmdir02-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1376150028] HOST FLAPPING ALERT: omnmdir02-ny;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1376150068] HOST ALERT: omnmdir02e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1376150068] HOST FLAPPING ALERT: omnmdir02e1-ny;STARTED; Host appears to have started flapping (23.9% change > 20.0% threshold)
[1376150078] SERVICE ALERT: omnmdir02-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376150078] SERVICE FLAPPING ALERT: omnmdir02-ny;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376150078] SERVICE ALERT: omnmdir02e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376150078] SERVICE FLAPPING ALERT: omnmdir02e1-ny;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376150258] SERVICE ALERT: omnmdir03e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150258] SERVICE ALERT: omnmdir03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150288] HOST ALERT: omnmdir03e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150288] HOST ALERT: omnmdir03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150318] SERVICE ALERT: omnmdir03e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150318] SERVICE ALERT: omnmdir03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150378] HOST ALERT: omnmdir03e1-ny;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.38 ms
[1376150378] HOST ALERT: omnmdir03-ny;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.20 ms
[1376150378] SERVICE ALERT: omnmdir03e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376150378] SERVICE ALERT: omnmdir03-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1376150438] SERVICE ALERT: omnmdir03e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150438] SERVICE ALERT: omnmdir03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150478] HOST ALERT: omnmdir03e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150478] HOST ALERT: omnmdir03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150498] SERVICE ALERT: omnmdir03e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150498] SERVICE ALERT: omnmdir03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150548] HOST ALERT: omnmdir03e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1376150548] HOST FLAPPING ALERT: omnmdir03e1-ny;STARTED; Host appears to have started flapping (23.9% change > 20.0% threshold)
[1376150548] HOST ALERT: omnmdir03-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376150548] HOST FLAPPING ALERT: omnmdir03-ny;STARTED; Host appears to have started flapping (23.9% change > 20.0% threshold)
[1376150558] SERVICE ALERT: omnmdir03e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1376150558] SERVICE FLAPPING ALERT: omnmdir03e1-ny;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376150558] SERVICE ALERT: omnmdir03-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376150558] SERVICE FLAPPING ALERT: omnmdir03-ny;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376150618] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;3;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1376150678] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1376150688] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150688] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150718] HOST ALERT: omnmdir04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150718] HOST ALERT: omnmdir04e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150738] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;2;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1376150748] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150748] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150788] HOST ALERT: omnmdir04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1376150788] HOST ALERT: omnmdir04e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376150798] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;3;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1376150798] SERVICE ALERT: omnmdir04-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376150798] SERVICE ALERT: omnmdir04e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376150868] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150868] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1376150898] HOST ALERT: omnmdir04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150898] HOST ALERT: omnmdir04e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1376150928] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150928] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1376150968] HOST ALERT: omnmdir04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1376150968] HOST FLAPPING ALERT: omnmdir04-ny;STARTED; Host appears to have started flapping (23.9% change > 20.0% threshold)
[1376150968] HOST ALERT: omnmdir04e1-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376150968] HOST FLAPPING ALERT: omnmdir04e1-ny;STARTED; Host appears to have started flapping (23.9% change > 20.0% threshold)
[1376150978] SERVICE ALERT: omnmdir04-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1376150978] SERVICE FLAPPING ALERT: omnmdir04-ny;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376150978] SERVICE ALERT: omnmdir04e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1376150978] SERVICE FLAPPING ALERT: omnmdir04e1-ny;PING;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1376151158] HOST FLAPPING ALERT: omnmdir02-ny;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1376151158] HOST FLAPPING ALERT: omnmdir02e1-ny;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376151218] SERVICE FLAPPING ALERT: omnmdir02-ny;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376151218] SERVICE FLAPPING ALERT: omnmdir02e1-ny;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376151638] HOST FLAPPING ALERT: omnmdir03e1-ny;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376151638] HOST FLAPPING ALERT: omnmdir03-ny;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376151698] SERVICE FLAPPING ALERT: omnmdir03e1-ny;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376151698] SERVICE FLAPPING ALERT: omnmdir03-ny;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376152118] SERVICE FLAPPING ALERT: omnmdir04e1-ny;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376152118] HOST FLAPPING ALERT: omnmdir04e1-ny;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376152118] SERVICE FLAPPING ALERT: omnmdir04-ny;PING;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376152118] HOST FLAPPING ALERT: omnmdir04-ny;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1376153918] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;3;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1376155631] HOST DOWNTIME ALERT: omnspsm-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155636] SERVICE DOWNTIME ALERT: omnspsm-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155642] HOST DOWNTIME ALERT: omnmdir04e1-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155650] HOST DOWNTIME ALERT: omnmdir04-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155655] HOST DOWNTIME ALERT: omnmdir01-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155658] SERVICE DOWNTIME ALERT: omnmdir01-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155662] SERVICE DOWNTIME ALERT: omnmdir04e1-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155666] HOST DOWNTIME ALERT: omnmdir03e1-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155669] HOST DOWNTIME ALERT: omnmdir01e1-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155673] HOST DOWNTIME ALERT: omnmdir03-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155680] HOST DOWNTIME ALERT: omnmdir02-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155701] HOST DOWNTIME ALERT: omnmdir02e1-ny;CANCELLED; Scheduled downtime for host has been cancelled.
[1376155738] SERVICE DOWNTIME ALERT: omnmdir04-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155741] SERVICE DOWNTIME ALERT: omnmdir01e1-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155745] SERVICE DOWNTIME ALERT: omnmdir03e1-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155748] SERVICE DOWNTIME ALERT: omnmdir02-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155753] SERVICE DOWNTIME ALERT: omnmdir03-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.
[1376155758] SERVICE DOWNTIME ALERT: omnmdir02e1-ny;PING;CANCELLED; Scheduled downtime for service has been cancelled.