[1393139312] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139312] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139372] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139412] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139482] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139492] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139522] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139522] SERVICE ALERT: nycnn-ipvx9;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139532] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139542] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139542] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139542] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139552] SERVICE ALERT: nycnn-ipvx5;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139552] SERVICE ALERT: nycnn-ipvx2;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139552] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 0.23 ms
[1393139552] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 0.26 ms
[1393139552] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 80%, RTA = 0.66 ms
[1393139552] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.21 ms
[1393139552] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139552] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139552] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1393139562] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139562] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139562] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139562] SERVICE ALERT: nycnn-ipvx27;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139562] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1393139572] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139572] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139572] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139572] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139572] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139572] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139572] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139572] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139572] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139582] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139582] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139582] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139582] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139582] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139582] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139582] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139582] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139582] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1393139582] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1393139592] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1393139592] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139592] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139592] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139592] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1393139592] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139592] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139592] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139592] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139592] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx35;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139592] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139592] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139592] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx10e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139592] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139592] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx9;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139592] SERVICE ALERT: nycnn-ipvx10;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139602] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139602] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139602] SERVICE ALERT: nycnn-ipvx11;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139602] SERVICE ALERT: nycnn-ipvx13;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139602] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139602] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139602] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx14e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx5;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx2;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx22e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139612] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx12;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx14;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139612] SERVICE ALERT: nycnn-ipvx14e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx16;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.22 ms
[1393139612] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx19;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139612] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.26 ms
[1393139612] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139612] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139612] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139612] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139612] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139622] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139622] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139622] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139622] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139622] SERVICE ALERT: nycnn-ipvx21;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139622] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 2.73 ms
[1393139622] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139622] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139622] SERVICE ALERT: nycnn-ipvx20;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139622] SERVICE ALERT: nycnn-ipvx1e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1393139632] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139632] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1393139632] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1393139632] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139632] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;2;(Host Check Timed Out)
[1393139632] SERVICE ALERT: nycnn-ipvx21;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx21e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx23;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139632] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139632] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139632] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139632] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139632] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139632] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139632] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139642] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139642] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1393139642] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139642] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139642] SERVICE ALERT: nycnn-ipvx27;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139642] SERVICE ALERT: nycnn-ipvx24e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139642] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139652] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139652] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139652] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139652] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139652] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139652] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139652] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139652] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139652] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1393139652] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139652] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139652] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139652] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139652] SERVICE ALERT: nycnn-ipvx3;PING;WARNING;HARD;2;PING WARNING - Packet loss = 54%, RTA = 0.60 ms
[1393139652] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.45 ms
[1393139652] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.34 ms
[1393139662] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1393139662] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1393139662] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;2;(Host Check Timed Out)
[1393139662] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139662] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139662] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.32 ms
[1393139662] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.46 ms
[1393139662] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139662] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 1.64 ms
[1393139662] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.30 ms
[1393139672] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1393139672] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1393139672] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139672] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1393139672] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1393139672] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1393139672] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1393139672] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139672] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1393139672] SERVICE ALERT: nycnn-ipvx10e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.98 ms
[1393139672] HOST ALERT: nycnn-ipvx7;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.30 ms
[1393139672] SERVICE ALERT: nycnn-ipvx10;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1393139672] SERVICE ALERT: nycnn-ipvx11;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx13;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx16;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139672] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139672] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1393139672] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1393139672] HOST ALERT: nycnn-ipvx28;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.33 ms
[1393139672] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.33 ms
[1393139682] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139682] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139682] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.46 ms
[1393139682] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;2;(Host Check Timed Out)
[1393139682] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139682] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139682] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139682] HOST ALERT: nycnn-ipvx33;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1393139682] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139682] HOST ALERT: nycnn-ipvx30;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139682] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.27 ms
[1393139682] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139682] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139682] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139682] SERVICE ALERT: nycnn-ipvx22e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139682] SERVICE ALERT: nycnn-ipvx20;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139682] SERVICE ALERT: nycnn-ipvx1e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139682] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1393139682] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1393139682] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1393139692] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;2;(Host Check Timed Out)
[1393139692] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139692] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139692] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;2;(Host Check Timed Out)
[1393139692] SERVICE ALERT: nycnn-ipvx25e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.47 ms
[1393139692] SERVICE ALERT: nycnn-ipvx25;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1393139692] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1393139692] SERVICE ALERT: nycnn-ipvx25e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1393139692] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139692] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139692] SERVICE ALERT: nycnn-ipvx35e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1393139692] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139702] SERVICE ALERT: nycnn-ipvx14;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.47 ms
[1393139702] HOST ALERT: nycnn-ipvx32;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139702] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1393139702] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139702] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139702] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1393139712] SERVICE ALERT: nycnn-ipvx29;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.50 ms
[1393139712] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139712] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139712] HOST ALERT: nycnn-ipvx29;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.27 ms
[1393139712] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.28 ms
[1393139712] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1393139712] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1393139712] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139712] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139712] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139712] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139712] HOST ALERT: nycnn-ipvx12;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139712] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139712] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139712] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139712] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139712] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139712] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.31 ms
[1393139712] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139722] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393139722] SERVICE ALERT: nycnn-ipvx8e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139722] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139722] HOST ALERT: nycnn-ipvx4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139722] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139722] HOST ALERT: nycnn-ipvx2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139722] HOST ALERT: nycnn-ipvx9;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139722] HOST ALERT: nycnn-ipvx15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1393139722] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139722] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1393139722] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139722] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1393139722] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.44 ms
[1393139722] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139722] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139722] SERVICE ALERT: nycnn-ipvx6;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139732] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;2;(Host Check Timed Out)
[1393139732] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139732] SERVICE ALERT: nycnn-ipvx8;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139732] SERVICE ALERT: nycnn-ipvx6e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139732] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139732] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1393139732] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139732] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139732] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139732] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139732] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139732] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1393139732] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139732] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139732] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139732] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139732] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139732] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139732] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139732] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139732] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139732] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1393139732] SERVICE ALERT: nycnn-ipvx8;PING;WARNING;HARD;2;PING WARNING - Packet loss = 37%, RTA = 0.46 ms
[1393139732] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.63 ms
[1393139732] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139732] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139732] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1393139732] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139732] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1393139732] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.31 ms
[1393139732] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.74 ms
[1393139742] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139742] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393139742] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139742] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139742] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.25 ms
[1393139742] SERVICE ALERT: nycnn-ipvx11e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139742] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139742] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139742] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139742] HOST ALERT: nycnn-ipvx1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139742] HOST ALERT: nycnn-ipvx25;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139742] HOST ALERT: nycnn-ipvx35;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139742] HOST ALERT: nycnn-ipvx19;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139742] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139742] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139742] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139742] HOST ALERT: nycnn-ipvx21;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139742] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139742] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139742] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139742] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139752] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139752] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1393139752] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;2;(Host Check Timed Out)
[1393139752] SERVICE ALERT: nycnn-ipvx31;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139752] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139752] SERVICE ALERT: nycnn-ipvx24;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139752] HOST ALERT: nycnn-ipvx23;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139752] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139752] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139752] HOST ALERT: nycnn-ipvx14;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139752] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139752] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139752] SERVICE ALERT: nycnn-ipvx31;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139762] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139762] SERVICE ALERT: nycnn-ipvx22;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139762] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 40.92 ms
[1393139762] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139762] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139762] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139772] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139772] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1393139772] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.34 ms
[1393139772] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.46 ms
[1393139782] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;2;(Host Check Timed Out)
[1393139782] SERVICE ALERT: nycnn-ipvx13e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393139792] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139792] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139792] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1393139792] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393139792] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1393139792] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139792] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139792] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139792] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139792] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139792] SERVICE ALERT: nycnn-ipvx24;XCode Service;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139802] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393139802] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139802] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139802] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139802] HOST ALERT: nycnn-ipvx10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139802] HOST ALERT: nycnn-ipvx31;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1393139812] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393139812] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139812] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393139822] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1393139832] HOST ALERT: nycnn-ipvx13;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393139852] SERVICE ALERT: nycnn-ipvx3;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139862] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139872] SERVICE ALERT: nycnn-ipvx23e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139872] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139872] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139882] SERVICE ALERT: nycnn-ipvx35;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139882] SERVICE ALERT: nycnn-ipvx2e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139882] SERVICE ALERT: nycnn-ipvx30e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139882] SERVICE ALERT: nycnn-ipvx9;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139882] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139882] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139882] SERVICE ALERT: nycnn-ipvx33e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139892] SERVICE ALERT: nycnn-ipvx30e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139892] SERVICE ALERT: nycnn-ipvx14e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139892] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139892] SERVICE ALERT: nycnn-ipvx5;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139892] SERVICE ALERT: nycnn-ipvx2;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139892] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139902] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139902] SERVICE ALERT: nycnn-ipvx12;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139902] SERVICE ALERT: nycnn-ipvx14;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139902] SERVICE ALERT: nycnn-ipvx14e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139902] SERVICE ALERT: nycnn-ipvx19;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139902] SERVICE ALERT: nycnn-ipvx17;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139902] SERVICE ALERT: nycnn-ipvx17e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393139902] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx21;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx21;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx21e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx25;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx23;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139912] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139922] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139922] SERVICE ALERT: nycnn-ipvx2;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139922] SERVICE ALERT: nycnn-ipvx12e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139922] SERVICE ALERT: nycnn-ipvx20e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139922] SERVICE ALERT: nycnn-ipvx27;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139932] SERVICE ALERT: nycnn-ipvx24e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139932] SERVICE ALERT: nycnn-ipvx27e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139942] SERVICE ALERT: nycnn-ipvx29e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139942] SERVICE ALERT: nycnn-ipvx10e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139942] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139942] SERVICE ALERT: nycnn-ipvx7e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139942] SERVICE ALERT: nycnn-ipvx31e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139942] SERVICE ALERT: nycnn-ipvx10;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx10e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx10;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139952] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139962] SERVICE ALERT: nycnn-ipvx11;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139962] SERVICE ALERT: nycnn-ipvx13;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139962] SERVICE ALERT: nycnn-ipvx16;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139962] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139962] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139962] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139972] SERVICE ALERT: nycnn-ipvx22e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139972] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139972] SERVICE ALERT: nycnn-ipvx20;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139972] SERVICE ALERT: nycnn-ipvx1e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139972] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139972] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139972] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139972] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393139972] SERVICE ALERT: nycnn-ipvx17e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139972] SERVICE ALERT: nycnn-ipvx17;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393139982] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139992] SERVICE ALERT: nycnn-ipvx25;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139992] SERVICE ALERT: nycnn-ipvx25e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393139992] SERVICE ALERT: nycnn-ipvx35e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393139992] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140002] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140002] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140002] SERVICE ALERT: nycnn-ipvx8e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140012] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1393140012] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1393140012] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140012] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140012] SERVICE ALERT: nycnn-ipvx6;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140012] SERVICE ALERT: nycnn-ipvx6e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140012] SERVICE ALERT: nycnn-ipvx8;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140022] SERVICE ALERT: nycnn-ipvx11e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140032] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140032] SERVICE ALERT: nycnn-ipvx31;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140032] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140032] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393140032] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393140032] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140032] SERVICE ALERT: nycnn-ipvx31;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140042] SERVICE ALERT: nycnn-ipvx22;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140072] SERVICE ALERT: nycnn-ipvx13e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140082] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1393140082] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393140092] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393140092] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393140252] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140262] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393140262] SERVICE ALERT: nycnn-ipvx17;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393140262] SERVICE ALERT: nycnn-ipvx17e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1393142942] SERVICE ALERT: nycnn-ipvstrm2;MMSService;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393142942] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1393142952] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.22 ms
[1393142972] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1393143012] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393143012] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393143012] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393143012] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393143012] SERVICE ALERT: nycnn-ipvstrm4;MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143012] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143012] SERVICE ALERT: nycnn-ipvstrm1;FTP Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143012] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143012] SERVICE ALERT: nycnn-ipvstrm2;MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393143012] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393143042] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1393143042] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1393143042] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1393143052] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143052] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143072] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1393143072] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393143072] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393143072] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393143072] SERVICE ALERT: nycnn-ipvstrm1;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143072] SERVICE ALERT: nycnn-ipvstrm4;MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393143072] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393143072] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393143072] SERVICE ALERT: nycnn-ipvstrm1;FTP Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1393143082] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143122] SERVICE ALERT: nycnn-ipvstrm1;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143122] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm3;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm1;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm2;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm4;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm3;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm4;Sentinel Protection Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143132] SERVICE ALERT: nycnn-ipvstrm3;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143142] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1393143142] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1393143142] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1393143142] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1393143152] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 66%, RTA = 0.25 ms
[1393143182] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [wmi/wmic.c:196:main()] ERROR: Login to remote object.
[1393143182] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1393143182] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393143182] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1393143182] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393143192] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1393143192] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1393143192] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1393143192] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1393143222] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;1;(Host Check Timed Out)
[1393143252] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1393143292] HOST ALERT: nycnn-ipvapilic;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1393143302] SERVICE ALERT: nycnn-ipvstrm2;MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1393143302] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1393143302] SERVICE ALERT: nycnn-ipvapilic;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1393143332] SERVICE ALERT: nycnn-ipvstrm2;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1393143332] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1393143362] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1393143362] SERVICE ALERT: nycnn-ipvstrm4;MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1393143362] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1393143362] SERVICE ALERT: nycnn-ipvstrm1;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1393143362] SERVICE ALERT: nycnn-ipvstrm1;FTP Service;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1393143362] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1133.37GB (81.1%), Free=263.82GB (18.9%)
[1393143412] SERVICE ALERT: nycnn-ipvstrm1;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1393143412] SERVICE ALERT: nycnn-ipvstrm3;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1393143412] SERVICE ALERT: nycnn-ipvstrm3;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1393143412] SERVICE ALERT: nycnn-ipvstrm1;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393143412] SERVICE ALERT: nycnn-ipvstrm4;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1393143412] SERVICE ALERT: nycnn-ipvstrm2;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393143412] SERVICE ALERT: nycnn-ipvstrm4;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1393143412] SERVICE ALERT: nycnn-ipvstrm3;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393143422] SERVICE ALERT: nycnn-ipvstrm4;Sentinel Protection Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1393143422] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1393143422] SERVICE ALERT: nycnn-ipvstrm3;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1393143422] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1134.16GB (81.2%), Free=263.02GB (18.8%)
[1393143482] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1134.51GB (81.2%), Free=262.68GB (18.8%)
[1393188162] SERVICE ALERT: rt-nycchdagg-a;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 127.24 ms
[1393188222] SERVICE ALERT: rt-nycchdagg-a;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.05 ms