[1397973851] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973851] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973861] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973861] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973861] SERVICE ALERT: nycnn-ipvx16;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973861] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973861] SERVICE ALERT: nycnn-ipvx17;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973861] SERVICE ALERT: nycnn-ipvx17e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx25;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx25e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx34e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx28;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx35e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973891] SERVICE ALERT: nycnn-ipvx3;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973901] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973901] SERVICE ALERT: nycnn-ipvx7;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973901] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973901] SERVICE ALERT: nycnn-ipvx8e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973911] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973911] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973911] SERVICE ALERT: nycnn-ipvx6;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973911] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.18 ms
[1397973911] SERVICE ALERT: nycnn-ipvx8;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973911] SERVICE ALERT: nycnn-ipvx6e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973911] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973911] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973921] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973921] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973921] SERVICE ALERT: nycnn-ipvx11e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973931] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973931] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973931] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973931] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973931] SERVICE ALERT: nycnn-ipvx31;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973931] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973931] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973931] SERVICE ALERT: nycnn-ipvx16;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973931] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.20 ms
[1397973931] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973931] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 72%, RTA = 0.22 ms
[1397973931] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 70%, RTA = 0.21 ms
[1397973931] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.22 ms
[1397973931] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.21 ms
[1397973931] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973931] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;1;(Host Check Timed Out)
[1397973941] SERVICE ALERT: nycnn-ipvx17e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973941] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx17;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973941] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.27 ms
[1397973941] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx7;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.25 ms
[1397973941] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973941] SERVICE ALERT: nycnn-ipvx22;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397973951] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1397973951] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1397973951] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973951] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973951] HOST ALERT: nycnn-ipvx2e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973951] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973951] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973951] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973951] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973951] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973961] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1397973961] HOST ALERT: nycnn-ipvx2;DOWN;SOFT;1;(Host Check Timed Out)
[1397973961] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973961] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx32e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx25;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx25e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973961] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] SERVICE ALERT: nycnn-ipvx34e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx28;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] SERVICE ALERT: nycnn-ipvx35e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx3;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx32;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx7;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx32e1;XCode Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx2e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973971] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973971] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx7e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1397973971] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973981] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973981] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1397973981] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973981] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1397973981] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973981] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973981] SERVICE ALERT: nycnn-ipvx8e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973981] SERVICE ALERT: nycnn-ipvx4;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx9e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx33;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx34;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx34e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx35e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx3e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx7e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx4e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973981] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] SERVICE ALERT: nycnn-ipvx5e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx6;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx7;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx8;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx6e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx9;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx9e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397973991] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397973991] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1397973991] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1397974001] SERVICE ALERT: nycnn-ipvx30;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974001] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx7e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx11e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974001] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974001] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974001] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974001] SERVICE ALERT: nycnn-ipvx19;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974001] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1397974001] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1397974001] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1397974001] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1397974001] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397974001] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1397974011] SERVICE ALERT: nycnn-ipvx30;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx23e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx31;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx29e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx29;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx31;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx35;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974011] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974011] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974011] SERVICE ALERT: nycnn-ipvx3e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974011] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974011] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974021] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974021] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974021] SERVICE ALERT: nycnn-ipvx22;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974021] SERVICE ALERT: nycnn-ipvx32e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974021] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974021] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397974021] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397974021] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397974021] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397974021] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1397974021] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1397974021] HOST ALERT: nycnn-ipvx7;DOWN;SOFT;2;(Host Check Timed Out)
[1397974031] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974031] HOST ALERT: nycnn-ipvx2e1;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.45 ms
[1397974031] SERVICE ALERT: nycnn-ipvx2e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 28%, RTA = 0.28 ms
[1397974031] HOST ALERT: nycnn-ipvx2;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 1.94 ms
[1397974031] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974031] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974031] SERVICE ALERT: nycnn-ipvx32;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974031] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;CRITICAL - Plugin timed out after 10 seconds
[1397974031] SERVICE ALERT: nycnn-ipvx2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974031] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974031] SERVICE ALERT: nycnn-ipvx32e1;XCode Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397974031] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974031] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974031] SERVICE ALERT: nycnn-ipvx15e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974041] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397974041] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1397974041] SERVICE ALERT: nycnn-ipvx15;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974041] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974041] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974041] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.12 ms
[1397974041] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974041] HOST ALERT: nycnn-ipvx16;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974041] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974041] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974041] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1397974041] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.29 ms
[1397974041] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 5.09 ms
[1397974041] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.28 ms
[1397974041] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1397974041] HOST ALERT: nycnn-ipvx3;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.26 ms
[1397974041] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.42 ms
[1397974051] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;2;(Host Check Timed Out)
[1397974051] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;2;(Host Check Timed Out)
[1397974051] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;2;(Host Check Timed Out)
[1397974051] SERVICE ALERT: nycnn-ipvx13e1;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974051] HOST ALERT: nycnn-ipvx7e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.50 ms
[1397974051] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.34 ms
[1397974051] HOST ALERT: nycnn-ipvx19;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.34 ms
[1397974051] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974051] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974051] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974051] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1397974051] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1397974051] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974051] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974051] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974051] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974051] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974051] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974051] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1397974051] SERVICE ALERT: nycnn-ipvx7e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1397974051] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974051] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1397974051] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.44 ms
[1397974051] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.50 ms
[1397974051] SERVICE ALERT: nycnn-ipvx34;PING;WARNING;HARD;2;PING WARNING - Packet loss = 37%, RTA = 0.49 ms
[1397974051] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.46 ms
[1397974051] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.30 ms
[1397974051] SERVICE ALERT: nycnn-ipvx34e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 44%, RTA = 0.33 ms
[1397974051] HOST ALERT: nycnn-ipvx28;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 1.49 ms
[1397974051] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.64 ms
[1397974051] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.48 ms
[1397974051] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974051] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974051] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974051] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974051] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974051] SERVICE ALERT: nycnn-ipvx30;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 54%, RTA = 0.35 ms
[1397974051] SERVICE ALERT: nycnn-ipvx30e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 54%, RTA = 0.29 ms
[1397974051] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974061] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.59 ms
[1397974061] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.23 ms
[1397974061] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1397974061] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974061] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1397974061] SERVICE ALERT: nycnn-ipvx7;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974061] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974061] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974061] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974061] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974061] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.50 ms
[1397974061] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.45 ms
[1397974061] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 1.07 ms
[1397974061] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.49 ms
[1397974061] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974061] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974061] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 48.09 ms
[1397974061] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974061] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.63 ms
[1397974061] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974061] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.70 ms
[1397974061] HOST ALERT: nycnn-ipvx15;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.75 ms
[1397974061] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.48 ms
[1397974061] SERVICE ALERT: nycnn-ipvx35e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 28%, RTA = 0.46 ms
[1397974061] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.56 ms
[1397974061] SERVICE ALERT: nycnn-ipvx27;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974061] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397974061] HOST ALERT: nycnn-ipvx25;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 1.45 ms
[1397974061] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.44 ms
[1397974061] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974061] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974061] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974061] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974071] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;2;(Host Check Timed Out)
[1397974071] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;2;(Host Check Timed Out)
[1397974071] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;2;(Host Check Timed Out)
[1397974071] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1397974071] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;2;(Host Check Timed Out)
[1397974071] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;2;(Host Check Timed Out)
[1397974071] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974071] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974071] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974071] SERVICE ALERT: nycnn-ipvx24;XCode Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397974071] HOST ALERT: nycnn-ipvx7;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974071] HOST ALERT: nycnn-ipvx29;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1397974071] HOST ALERT: nycnn-ipvx6;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974071] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974081] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 2.02 ms
[1397974081] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974081] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.29 ms
[1397974081] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1397974081] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974081] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974081] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974081] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974081] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974091] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974091] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1397974091] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974091] SERVICE ALERT: nycnn-ipvx2e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974091] SERVICE ALERT: nycnn-ipvx13;PING;WARNING;HARD;2;PING WARNING - Packet loss = 37%, RTA = 0.50 ms
[1397974091] HOST ALERT: nycnn-ipvx13;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.33 ms
[1397974091] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974091] HOST ALERT: nycnn-ipvx12;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974091] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1397974091] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1397974091] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974091] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1397974091] HOST ALERT: nycnn-ipvx20;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974091] HOST ALERT: nycnn-ipvx11;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1397974091] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974091] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1397974101] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;2;(Host Check Timed Out)
[1397974101] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974111] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974111] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974111] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974111] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974111] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974111] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1397974111] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974111] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974111] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974111] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1397974111] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974111] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974111] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1397974111] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974111] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397974111] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974111] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974111] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397974111] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974111] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1397974111] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974111] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1397974111] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974121] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1397974121] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974121] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1397974121] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974121] HOST ALERT: nycnn-ipvx10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974121] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974121] HOST ALERT: nycnn-ipvx8;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1397974131] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974131] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397974141] HOST ALERT: nycnn-ipvx31;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1397974151] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1397974151] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1397974161] HOST ALERT: nycnn-ipvx22;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397974211] SERVICE ALERT: nycnn-ipvx16;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974211] SERVICE ALERT: nycnn-ipvx16e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974221] SERVICE ALERT: nycnn-ipvx17;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974221] SERVICE ALERT: nycnn-ipvx17e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974221] SERVICE ALERT: nycnn-ipvx16;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974221] SERVICE ALERT: nycnn-ipvx16e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974221] SERVICE ALERT: nycnn-ipvx17;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974221] SERVICE ALERT: nycnn-ipvx17e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx32;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx25;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx25e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx33e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx34e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx28;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx35e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx3;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974251] SERVICE ALERT: nycnn-ipvx5;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974261] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974261] SERVICE ALERT: nycnn-ipvx7;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974261] SERVICE ALERT: nycnn-ipvx8;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974261] SERVICE ALERT: nycnn-ipvx8e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974261] SERVICE ALERT: nycnn-ipvx33;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974261] SERVICE ALERT: nycnn-ipvx4;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974261] SERVICE ALERT: nycnn-ipvx9e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx34;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx34e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx35e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx4e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx7e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx3e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx6;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx5e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx8e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx8;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx7;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx6e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974271] SERVICE ALERT: nycnn-ipvx9;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974281] SERVICE ALERT: nycnn-ipvx9e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974281] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974281] SERVICE ALERT: nycnn-ipvx30;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974281] SERVICE ALERT: nycnn-ipvx11e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974281] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx22;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx19;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx23e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx31;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx30;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974291] SERVICE ALERT: nycnn-ipvx29e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx31;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx29;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx35;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx3e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx28e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974301] SERVICE ALERT: nycnn-ipvx22;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974311] SERVICE ALERT: nycnn-ipvx32e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974321] SERVICE ALERT: nycnn-ipvx32;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974321] SERVICE ALERT: nycnn-ipvx32e1;XCode Service;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974321] SERVICE ALERT: nycnn-ipvx15e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974321] SERVICE ALERT: nycnn-ipvx15;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974331] SERVICE ALERT: nycnn-ipvx13e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974331] SERVICE ALERT: nycnn-ipvx19e1;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974351] SERVICE ALERT: nycnn-ipvx27;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974351] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974351] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974351] SERVICE ALERT: nycnn-ipvx24;XCode Service;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1397974381] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397974391] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1397977451] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.27 ms
[1397977451] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397977461] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1397977481] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1397977481] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1397977501] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1397977511] SERVICE ALERT: nycnn-ipvstrm3;SVServer;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977511] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1397977511] SERVICE ALERT: nycnn-ipvapilic;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397977511] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1397977511] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397977511] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397977531] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397977541] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1397977541] HOST ALERT: nycnn-ipvapilic;DOWN;SOFT;1;(Host Check Timed Out)
[1397977551] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1397977561] SERVICE ALERT: nycnn-ipvapilic;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397977571] SERVICE ALERT: nycnn-ipvstrm3;SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1397977571] HOST ALERT: nycnn-ipvapilic;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397977571] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1397977571] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1397977581] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977581] SERVICE ALERT: nycnn-ipvstrm1;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977581] SERVICE ALERT: nycnn-ipvstrm1;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977601] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1397977611] SERVICE ALERT: nycnn-ipvstrm2;SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977611] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977641] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1397977641] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977641] SERVICE ALERT: nycnn-ipvstrm2;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977641] SERVICE ALERT: nycnn-ipvstrm3;MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977641] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977641] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977641] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1397977651] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;3;(Host Check Timed Out)
[1397977661] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;3;PING WARNING - Packet loss = 80%, RTA = 0.27 ms
[1397977671] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397977681] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1397977681] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1397977681] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1397977691] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1397977691] SERVICE ALERT: nycnn-ipvstrm3;PING;WARNING;HARD;2;PING WARNING - Packet loss = 54%, RTA = 0.40 ms
[1397977711] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1397977741] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1397977821] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;OK;HARD;2;OK - New Volume Total=1397.19GB, Used=1119.03GB (80.1%), Free=278.15GB (19.9%)
[1397977851] SERVICE ALERT: nycnn-ipvstrm3;SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1397977861] SERVICE ALERT: nycnn-ipvstrm1;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1397977861] SERVICE ALERT: nycnn-ipvstrm1;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1397977861] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1397977891] SERVICE ALERT: nycnn-ipvstrm2;SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1397977891] SERVICE ALERT: nycnn-ipvstrm2;FTP Service;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1397977921] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1397977921] SERVICE ALERT: nycnn-ipvstrm2;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1397977921] SERVICE ALERT: nycnn-ipvstrm3;MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1397977921] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1397977921] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1121.45GB (80.3%), Free=275.74GB (19.7%)
[1397977931] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1125.14GB (80.5%), Free=272.05GB (19.5%)