[1407384598] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 311 processes with STATE = RSZDT
[1407384658] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 88 processes with STATE = RSZDT
[1407386458] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 291 processes with STATE = RSZDT
[1407386518] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 76 processes with STATE = RSZDT
[1407388318] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 288 processes with STATE = RSZDT
[1407388378] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 78 processes with STATE = RSZDT
[1407390178] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 294 processes with STATE = RSZDT
[1407390238] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 84 processes with STATE = RSZDT
[1407392038] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 303 processes with STATE = RSZDT
[1407392098] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 96 processes with STATE = RSZDT
[1407393898] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 301 processes with STATE = RSZDT
[1407393958] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 82 processes with STATE = RSZDT
[1407395758] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 292 processes with STATE = RSZDT
[1407395818] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 80 processes with STATE = RSZDT
[1407397618] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 294 processes with STATE = RSZDT
[1407397678] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 78 processes with STATE = RSZDT
[1407399478] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 294 processes with STATE = RSZDT
[1407399538] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 84 processes with STATE = RSZDT
[1407399958] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;SOFT;1;Job 10351434 remains in a queued state
[1407400008] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;SOFT;2;AME is idle:
[1407401338] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 303 processes with STATE = RSZDT
[1407401398] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 96 processes with STATE = RSZDT
[1407403198] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 300 processes with STATE = RSZDT
[1407403258] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 78 processes with STATE = RSZDT
[1407405058] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 312 processes with STATE = RSZDT
[1407405118] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 80 processes with STATE = RSZDT
[1407406918] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 288 processes with STATE = RSZDT
[1407406978] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 73 processes with STATE = RSZDT
[1407408778] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 288 processes with STATE = RSZDT
[1407408838] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 90 processes with STATE = RSZDT
[1407410638] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 302 processes with STATE = RSZDT
[1407410698] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 102 processes with STATE = RSZDT
[1407412498] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 300 processes with STATE = RSZDT
[1407412558] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 78 processes with STATE = RSZDT
[1407414358] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 297 processes with STATE = RSZDT
[1407414418] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 82 processes with STATE = RSZDT
[1407416218] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 288 processes with STATE = RSZDT
[1407416278] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 72 processes with STATE = RSZDT
[1407418078] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 287 processes with STATE = RSZDT
[1407418138] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 89 processes with STATE = RSZDT
[1407419938] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 292 processes with STATE = RSZDT
[1407419998] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 91 processes with STATE = RSZDT
[1407421378] SERVICE ALERT: nycnn-ipvapilic;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407421438] SERVICE ALERT: nycnn-ipvapilic;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.08%
[1407421798] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 303 processes with STATE = RSZDT
[1407421858] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 73 processes with STATE = RSZDT
[1407423658] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 293 processes with STATE = RSZDT
[1407423718] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 76 processes with STATE = RSZDT
[1407425518] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 295 processes with STATE = RSZDT
[1407425578] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 67 processes with STATE = RSZDT
[1407427378] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 283 processes with STATE = RSZDT
[1407427388] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407427428] HOST ALERT: omnmgsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1407427438] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 79 processes with STATE = RSZDT
[1407427448] SERVICE ALERT: omnmgsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407427498] HOST ALERT: omnmgsm-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1407427508] SERVICE ALERT: omnmgsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407429238] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 300 processes with STATE = RSZDT
[1407429298] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 91 processes with STATE = RSZDT
[1407431098] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 295 processes with STATE = RSZDT
[1407431158] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 77 processes with STATE = RSZDT
[1407432958] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 291 processes with STATE = RSZDT
[1407433018] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 69 processes with STATE = RSZDT
[1407433681] HOST DOWNTIME ALERT: nycnn-ipvx1;STARTED; Host has entered a period of scheduled downtime
[1407433706] HOST DOWNTIME ALERT: nycnn-ipvx10;STARTED; Host has entered a period of scheduled downtime
[1407433727] HOST DOWNTIME ALERT: nycnn-ipvx11;STARTED; Host has entered a period of scheduled downtime
[1407433753] HOST DOWNTIME ALERT: nycnn-ipvx12;STARTED; Host has entered a period of scheduled downtime
[1407433804] HOST DOWNTIME ALERT: nycnn-ipvx13;STARTED; Host has entered a period of scheduled downtime
[1407433832] HOST DOWNTIME ALERT: nycnn-ipvx14;STARTED; Host has entered a period of scheduled downtime
[1407433852] HOST DOWNTIME ALERT: nycnn-ipvx15;STARTED; Host has entered a period of scheduled downtime
[1407433876] HOST DOWNTIME ALERT: nycnn-ipvx16;STARTED; Host has entered a period of scheduled downtime
[1407433907] HOST DOWNTIME ALERT: nycnn-ipvx17;STARTED; Host has entered a period of scheduled downtime
[1407433927] HOST DOWNTIME ALERT: nycnn-ipvx18;STARTED; Host has entered a period of scheduled downtime
[1407433943] HOST DOWNTIME ALERT: nycnn-ipvx19;STARTED; Host has entered a period of scheduled downtime
[1407433968] HOST DOWNTIME ALERT: nycnn-ipvx2;STARTED; Host has entered a period of scheduled downtime
[1407433988] HOST DOWNTIME ALERT: nycnn-ipvx20;STARTED; Host has entered a period of scheduled downtime
[1407434007] HOST DOWNTIME ALERT: nycnn-ipvx21;STARTED; Host has entered a period of scheduled downtime
[1407434025] HOST DOWNTIME ALERT: nycnn-ipvx22;STARTED; Host has entered a period of scheduled downtime
[1407434043] HOST DOWNTIME ALERT: nycnn-ipvx23;STARTED; Host has entered a period of scheduled downtime
[1407434060] HOST DOWNTIME ALERT: nycnn-ipvx24;STARTED; Host has entered a period of scheduled downtime
[1407434078] HOST DOWNTIME ALERT: nycnn-ipvx25;STARTED; Host has entered a period of scheduled downtime
[1407434095] HOST DOWNTIME ALERT: nycnn-ipvx26;STARTED; Host has entered a period of scheduled downtime
[1407434162] HOST DOWNTIME ALERT: nycnn-ipvx27;STARTED; Host has entered a period of scheduled downtime
[1407434185] HOST DOWNTIME ALERT: nycnn-ipvx28;STARTED; Host has entered a period of scheduled downtime
[1407434200] HOST DOWNTIME ALERT: nycnn-ipvx29;STARTED; Host has entered a period of scheduled downtime
[1407434218] HOST DOWNTIME ALERT: nycnn-ipvx3;STARTED; Host has entered a period of scheduled downtime
[1407434236] HOST DOWNTIME ALERT: nycnn-ipvx30;STARTED; Host has entered a period of scheduled downtime
[1407434253] HOST DOWNTIME ALERT: nycnn-ipvx31;STARTED; Host has entered a period of scheduled downtime
[1407434268] HOST DOWNTIME ALERT: nycnn-ipvx32;STARTED; Host has entered a period of scheduled downtime
[1407434285] HOST DOWNTIME ALERT: nycnn-ipvx33;STARTED; Host has entered a period of scheduled downtime
[1407434302] HOST DOWNTIME ALERT: nycnn-ipvx34;STARTED; Host has entered a period of scheduled downtime
[1407434318] HOST DOWNTIME ALERT: nycnn-ipvx35;STARTED; Host has entered a period of scheduled downtime
[1407434340] HOST DOWNTIME ALERT: nycnn-ipvx4;STARTED; Host has entered a period of scheduled downtime
[1407434359] HOST DOWNTIME ALERT: nycnn-ipvx5;STARTED; Host has entered a period of scheduled downtime
[1407434377] HOST DOWNTIME ALERT: nycnn-ipvx6;STARTED; Host has entered a period of scheduled downtime
[1407434393] HOST DOWNTIME ALERT: nycnn-ipvx8;STARTED; Host has entered a period of scheduled downtime
[1407434415] HOST DOWNTIME ALERT: nycnn-ipvx9;STARTED; Host has entered a period of scheduled downtime
[1407434818] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 286 processes with STATE = RSZDT
[1407434878] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 77 processes with STATE = RSZDT
[1407434988] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407434988] SERVICE ALERT: nycnn-ipvx1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407434988] SERVICE ALERT: nycnn-ipvx1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407434988] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407434988] SERVICE ALERT: nycnn-ipvx1e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407434998] SERVICE ALERT: nycnn-ipvx1e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407435008] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 80%, RTA = 0.23 ms
[1407435008] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 80%, RTA = 0.20 ms
[1407435018] SERVICE ALERT: nycnn-ipvx1e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435048] HOST ALERT: nycnn-ipvx1;DOWN;SOFT;1;(Host Check Timed Out)
[1407435048] HOST ALERT: nycnn-ipvx1e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407435068] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435068] SERVICE ALERT: nycnn-ipvx1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435068] SERVICE ALERT: nycnn-ipvx1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435068] SERVICE ALERT: nycnn-ipvx1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435068] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435068] SERVICE ALERT: nycnn-ipvx1e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435068] SERVICE ALERT: nycnn-ipvx1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407435068] SERVICE ALERT: nycnn-ipvx1e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407435068] SERVICE ALERT: nycnn-ipvx1e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435078] SERVICE ALERT: nycnn-ipvx1e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435138] HOST ALERT: nycnn-ipvx1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 18.98 ms
[1407435138] HOST ALERT: nycnn-ipvx1e1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.46 ms
[1407435188] SERVICE ALERT: nycnn-ipvx1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407435188] SERVICE ALERT: nycnn-ipvx1e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1407435348] SERVICE ALERT: nycnn-ipvx1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407435348] SERVICE ALERT: nycnn-ipvx1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407435348] SERVICE ALERT: nycnn-ipvx1e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407435348] SERVICE ALERT: nycnn-ipvx1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407435348] SERVICE ALERT: nycnn-ipvx1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407435348] SERVICE ALERT: nycnn-ipvx1e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407435358] SERVICE ALERT: nycnn-ipvx1e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407435358] SERVICE ALERT: nycnn-ipvx1e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407435408] SERVICE ALERT: nycnn-ipvx1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.71%
[1407435418] SERVICE ALERT: nycnn-ipvx1e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.47%
[1407435428] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435438] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.113)
[1407435448] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435458] HOST ALERT: nycnn-editxfer;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.120)
[1407435478] HOST ALERT: nycnn-editxfer;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407435478] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407435498] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407435498] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407435548] HOST ALERT: nycnn-editxfer;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.120)
[1407435568] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.113)
[1407435598] SERVICE ALERT: nycnn-editxfer3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1407435608] HOST ALERT: nycnn-editxfer3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1407435618] HOST ALERT: nycnn-editxfer;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1407435618] SERVICE ALERT: nycnn-editxfer;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1407435688] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435698] HOST ALERT: nycnn-editxfer;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.120)
[1407435728] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435728] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407435738] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.113)
[1407435738] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407435768] HOST ALERT: nycnn-editxfer;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407435768] SERVICE ALERT: nycnn-ipvx13;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435778] SERVICE ALERT: nycnn-ipvx13;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435778] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407435788] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435788] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435808] HOST ALERT: nycnn-ipvx13;DOWN;SOFT;1;(Host Check Timed Out)
[1407435808] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435808] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407435808] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435808] SERVICE ALERT: nycnn-ipvx13;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435828] HOST ALERT: nycnn-ipvx13e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407435828] SERVICE ALERT: nycnn-ipvx13;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435838] HOST ALERT: nycnn-editxfer;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.120)
[1407435838] SERVICE ALERT: nycnn-ipvx13;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435848] SERVICE ALERT: nycnn-ipvx13;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407435848] SERVICE ALERT: nycnn-ipvx13e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407435868] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407435878] HOST ALERT: nycnn-editxfer3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1407435878] HOST FLAPPING ALERT: nycnn-editxfer3;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1407435888] HOST ALERT: nycnn-ipvx13e1;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.40 ms
[1407435888] HOST ALERT: nycnn-ipvx13;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.48 ms
[1407435888] SERVICE ALERT: nycnn-ipvx13e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407435888] SERVICE ALERT: nycnn-ipvx13;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407435898] SERVICE ALERT: nycnn-editxfer3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1407435898] SERVICE FLAPPING ALERT: nycnn-editxfer3;PING;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1407435898] SERVICE ALERT: nycnn-ipvx13;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407435898] SERVICE ALERT: nycnn-ipvx13e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407435908] HOST ALERT: nycnn-editxfer;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1407435908] HOST FLAPPING ALERT: nycnn-editxfer;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1407435908] SERVICE ALERT: nycnn-ipvx13e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435908] SERVICE ALERT: nycnn-ipvx13e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435918] SERVICE ALERT: nycnn-ipvx13;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407435918] SERVICE ALERT: nycnn-editxfer;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1407435918] SERVICE FLAPPING ALERT: nycnn-editxfer;PING;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1407435948] SERVICE ALERT: nycnn-ipvx13e1;Process: Siphon;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407435948] SERVICE ALERT: nycnn-ipvx14;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435948] SERVICE ALERT: nycnn-ipvx14e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435948] SERVICE ALERT: nycnn-ipvx13;Process: XCode;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407435958] SERVICE ALERT: nycnn-ipvx13e1;Process: XCode;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407435958] SERVICE ALERT: nycnn-ipvx13e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 360 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407435958] SERVICE ALERT: nycnn-ipvx14e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435968] SERVICE ALERT: nycnn-ipvx14e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435968] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435968] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407435978] SERVICE ALERT: nycnn-ipvx13;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 4.99%
[1407435978] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435978] SERVICE ALERT: nycnn-ipvx14;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407435978] HOST ALERT: nycnn-ipvx14;DOWN;SOFT;1;(Host Check Timed Out)
[1407435978] HOST ALERT: nycnn-ipvx14e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407435988] SERVICE ALERT: nycnn-ipvx14e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436008] SERVICE ALERT: nycnn-ipvx14;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436008] SERVICE ALERT: nycnn-ipvx14e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436018] SERVICE ALERT: nycnn-ipvx13e1;CPU Load;OK;SOFT;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.65%
[1407436018] SERVICE ALERT: nycnn-ipvx14e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436028] SERVICE ALERT: nycnn-ipvx14e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436028] SERVICE ALERT: nycnn-ipvx14e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436028] SERVICE ALERT: nycnn-ipvx14;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436038] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436038] SERVICE ALERT: nycnn-ipvx14;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436058] HOST ALERT: nycnn-ipvx14;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.47 ms
[1407436058] HOST ALERT: nycnn-ipvx14e1;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.41 ms
[1407436058] SERVICE ALERT: nycnn-ipvx14;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407436078] SERVICE ALERT: nycnn-ipvx14e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407436088] SERVICE ALERT: nycnn-ipvx14;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407436088] SERVICE ALERT: nycnn-ipvx13;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436088] SERVICE ALERT: nycnn-ipvx13;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407436108] SERVICE ALERT: nycnn-ipvx13;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.828GB (42%) - Free: 1.162GB (58%)
[1407436118] SERVICE ALERT: nycnn-ipvx14;Service: Apache;OK;SOFT;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407436118] SERVICE ALERT: nycnn-ipvx13;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436148] SERVICE ALERT: nycnn-ipvx13e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436188] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407436188] SERVICE ALERT: nycnn-ipvx15e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407436188] SERVICE ALERT: nycnn-ipvx15;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407436208] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436208] SERVICE ALERT: nycnn-ipvx6e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436208] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436248] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;1;(Host Check Timed Out)
[1407436248] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407436248] SERVICE ALERT: nycnn-ipvx6;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436248] SERVICE ALERT: nycnn-ipvx6e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436248] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436248] SERVICE ALERT: nycnn-ipvx15;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436258] SERVICE ALERT: nycnn-ipvx6;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436268] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436268] SERVICE ALERT: nycnn-ipvx15e1;Process: Siphon;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436268] SERVICE ALERT: nycnn-ipvx15;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436268] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436268] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436268] SERVICE ALERT: nycnn-ipvx6;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436268] SERVICE ALERT: nycnn-ipvx6e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436268] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436268] SERVICE ALERT: nycnn-ipvx6e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436268] SERVICE ALERT: nycnn-ipvx14e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436268] SERVICE ALERT: nycnn-ipvx15;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436278] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436278] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436278] SERVICE ALERT: nycnn-ipvx6;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436278] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436278] HOST ALERT: nycnn-ipvx15;DOWN;SOFT;1;(Host Check Timed Out)
[1407436298] HOST ALERT: nycnn-ipvx6;DOWN;SOFT;2;(Host Check Timed Out)
[1407436298] SERVICE ALERT: nycnn-ipvx14;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436298] SERVICE ALERT: nycnn-ipvx14e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.828GB (42%) - Free: 1.162GB (58%)
[1407436298] HOST ALERT: nycnn-ipvx15e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407436298] HOST ALERT: nycnn-ipvx6e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407436308] SERVICE ALERT: nycnn-ipvx14e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436308] SERVICE ALERT: nycnn-ipvx14e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436308] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436308] SERVICE ALERT: nycnn-ipvx15;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436318] SERVICE ALERT: nycnn-ipvx6;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407436328] SERVICE ALERT: nycnn-ipvx6e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407436328] SERVICE ALERT: nycnn-ipvx14;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436328] SERVICE ALERT: nycnn-ipvx14;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436328] SERVICE ALERT: nycnn-ipvx15e1;Process: Siphon;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407436328] HOST ALERT: nycnn-ipvx6;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407436328] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436328] SERVICE ALERT: nycnn-ipvx15e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436328] SERVICE ALERT: nycnn-ipvx15;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436328] SERVICE ALERT: nycnn-ipvx15;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436338] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407436338] HOST ALERT: nycnn-ipvx6e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407436348] HOST ALERT: nycnn-ipvx15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407436358] HOST ALERT: nycnn-ipvx15e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1407436368] SERVICE ALERT: nycnn-ipvx14e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.13%
[1407436378] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436378] SERVICE ALERT: nycnn-ipvx15e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1407436388] SERVICE ALERT: nycnn-ipvx15;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407436388] SERVICE ALERT: nycnn-ipvx14;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1407436428] SERVICE ALERT: nycnn-ipvx15e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 292 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436438] SERVICE ALERT: nycnn-ipvx6e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.49%
[1407436488] SERVICE ALERT: nycnn-ipvx15e1;CPU Load;OK;SOFT;2;OK (Sample Period 59 sec) - Average CPU Utilisation 0.31%
[1407436538] SERVICE ALERT: nycnn-ipvx6;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.817GB (41%) - Free: 1.173GB (59%)
[1407436538] SERVICE ALERT: nycnn-ipvx6e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407436548] SERVICE ALERT: nycnn-ipvx6;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436548] SERVICE ALERT: nycnn-ipvx15e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436548] SERVICE ALERT: nycnn-ipvx15;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407436558] SERVICE ALERT: nycnn-ipvx6e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436568] SERVICE ALERT: nycnn-ipvx6e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436568] SERVICE ALERT: nycnn-ipvx6;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436568] SERVICE ALERT: nycnn-ipvx6;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407436598] SERVICE ALERT: nycnn-ipvx15;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436598] SERVICE ALERT: nycnn-ipvx15;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407436608] SERVICE ALERT: nycnn-ipvx16;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407436608] SERVICE ALERT: nycnn-ipvx6;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.11%
[1407436608] SERVICE ALERT: nycnn-ipvx15e1;Process: Siphon;OK;HARD;3;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436608] SERVICE ALERT: nycnn-ipvx15e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407436618] SERVICE ALERT: nycnn-ipvx15;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436628] SERVICE ALERT: nycnn-ipvx15e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407436628] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436628] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436628] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436628] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436648] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;1;(Host Check Timed Out)
[1407436668] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407436668] HOST ALERT: nycnn-ipvx17;DOWN;SOFT;1;(Host Check Timed Out)
[1407436668] HOST ALERT: nycnn-ipvx17e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407436668] SERVICE ALERT: nycnn-ipvx16e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436668] SERVICE ALERT: nycnn-ipvx17;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436678] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 322 processes with STATE = RSZDT
[1407436678] SERVICE ALERT: nycnn-ipvx15;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1407436678] SERVICE ALERT: nycnn-ipvx17;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436678] SERVICE ALERT: nycnn-ipvx16e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436678] SERVICE ALERT: nycnn-ipvx16;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407436678] SERVICE ALERT: nycnn-ipvx16e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436688] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436688] SERVICE ALERT: nycnn-ipvx16;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436688] SERVICE ALERT: nycnn-ipvx16e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436688] SERVICE ALERT: nycnn-ipvx17;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436688] SERVICE ALERT: nycnn-ipvx17e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407436698] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.113)
[1407436698] SERVICE ALERT: nycnn-ipvx16e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436708] HOST ALERT: nycnn-ipvx16e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407436708] HOST ALERT: nycnn-ipvx16;DOWN;SOFT;2;(Host Check Timed Out)
[1407436708] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436708] SERVICE ALERT: nycnn-ipvx17;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436708] SERVICE ALERT: nycnn-ipvx16e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436718] HOST ALERT: nycnn-editxfer;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.120)
[1407436738] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 93 processes with STATE = RSZDT
[1407436738] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407436748] SERVICE ALERT: nycnn-ipvx17e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407436748] SERVICE ALERT: nycnn-ipvx17;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436758] HOST ALERT: nycnn-ipvx17;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.47 ms
[1407436758] HOST ALERT: nycnn-ipvx17e1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 3.71 ms
[1407436758] SERVICE ALERT: nycnn-ipvx17e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436758] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407436768] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407436778] HOST ALERT: nycnn-ipvx16;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407436778] HOST ALERT: nycnn-ipvx16e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407436788] HOST ALERT: nycnn-editxfer;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407436798] SERVICE ALERT: nycnn-ipvx17e1;Service: Apache;OK;SOFT;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407436808] SERVICE ALERT: nycnn-ipvx16;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407436808] SERVICE ALERT: nycnn-ipvx16e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407436808] SERVICE ALERT: nycnn-ipvx17;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407436808] SERVICE ALERT: nycnn-ipvx17e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407436818] SERVICE ALERT: nycnn-ipvx17e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436818] SERVICE ALERT: nycnn-ipvx16;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436838] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.113)
[1407436858] HOST ALERT: nycnn-editxfer;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.120)
[1407436878] SERVICE ALERT: nycnn-ipvx16;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.23%
[1407436878] SERVICE ALERT: nycnn-ipvx17e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.32%
[1407436878] SERVICE ALERT: nycnn-editxfer;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407436888] HOST ALERT: nycnn-editxfer;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1407436908] HOST ALERT: nycnn-editxfer3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1407436918] SERVICE ALERT: nycnn-editxfer3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1407436938] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407436948] SERVICE ALERT: nycnn-ipvx16e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407436958] SERVICE ALERT: nycnn-ipvx17;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436958] SERVICE ALERT: nycnn-ipvx17;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436968] SERVICE ALERT: nycnn-ipvx16e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407436968] SERVICE ALERT: nycnn-ipvx16;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407436968] SERVICE ALERT: nycnn-ipvx16e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.845GB (42%) - Free: 1.145GB (58%)
[1407436978] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436988] SERVICE ALERT: nycnn-ipvx5;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436988] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407436988] SERVICE ALERT: nycnn-ipvx16e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407436988] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407436988] SERVICE ALERT: nycnn-ipvx17;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.634GB (32%) - Free: 1.356GB (68%)
[1407436988] SERVICE ALERT: nycnn-ipvx16e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407436998] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437008] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437018] HOST ALERT: nycnn-ipvx5e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407437018] SERVICE ALERT: nycnn-ipvx17;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1407437028] HOST ALERT: nycnn-ipvx5;DOWN;SOFT;1;(Host Check Timed Out)
[1407437038] SERVICE ALERT: nycnn-ipvx17;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407437038] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437048] SERVICE ALERT: nycnn-ipvx5;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437048] SERVICE ALERT: nycnn-ipvx5e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437048] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437048] SERVICE ALERT: nycnn-ipvx17e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407437048] SERVICE ALERT: nycnn-ipvx16e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.12%
[1407437048] SERVICE ALERT: nycnn-ipvx5e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437058] SERVICE ALERT: nycnn-ipvx5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437068] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407437098] HOST ALERT: nycnn-ipvx5;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1407437098] HOST ALERT: nycnn-ipvx5e1;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.50 ms
[1407437098] SERVICE ALERT: nycnn-ipvx5e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437098] SERVICE ALERT: nycnn-ipvx5;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437108] SERVICE ALERT: nycnn-ipvx5e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407437108] SERVICE ALERT: nycnn-ipvx5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407437128] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437138] SERVICE ALERT: nycnn-ipvx5;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437138] HOST ALERT: nycnn-editxfer;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.120)
[1407437148] SERVICE ALERT: nycnn-ipvx5;Process: XCode;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407437148] SERVICE ALERT: nycnn-ipvx5e1;Process: Siphon;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437168] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437178] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.113)
[1407437178] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407437198] SERVICE ALERT: nycnn-ipvx5;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.68%
[1407437208] HOST ALERT: nycnn-editxfer;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407437218] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407437218] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407437238] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407437238] SERVICE ALERT: nycnn-ipvx8;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407437278] HOST ALERT: nycnn-editxfer;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.120)
[1407437288] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.113)
[1407437288] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437288] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437288] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437288] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437298] SERVICE ALERT: nycnn-ipvx9;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437308] SERVICE ALERT: nycnn-ipvx8;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437308] SERVICE ALERT: nycnn-ipvx9;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437308] SERVICE ALERT: nycnn-ipvx9e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437308] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437308] SERVICE ALERT: nycnn-ipvx8;CPU Load;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437328] HOST ALERT: nycnn-ipvx8;DOWN;SOFT;1;(Host Check Timed Out)
[1407437328] HOST ALERT: nycnn-ipvx8e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407437328] HOST ALERT: nycnn-ipvx9e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407437328] HOST ALERT: nycnn-ipvx9;DOWN;SOFT;1;(Host Check Timed Out)
[1407437328] SERVICE ALERT: nycnn-ipvx5e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437328] SERVICE ALERT: nycnn-ipvx5;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437328] SERVICE ALERT: nycnn-ipvx5e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.817GB (41%) - Free: 1.173GB (59%)
[1407437328] SERVICE ALERT: nycnn-ipvx5e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407437338] SERVICE ALERT: nycnn-editxfer3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407437348] SERVICE ALERT: nycnn-ipvx8e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437348] SERVICE ALERT: nycnn-ipvx9;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437348] SERVICE ALERT: nycnn-ipvx9e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437348] SERVICE ALERT: nycnn-ipvx9e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437348] HOST ALERT: nycnn-editxfer3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1407437348] HOST ALERT: nycnn-editxfer;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1407437348] SERVICE ALERT: nycnn-ipvx8;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437348] SERVICE ALERT: nycnn-ipvx8e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437348] SERVICE ALERT: nycnn-ipvx9e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437348] SERVICE ALERT: nycnn-ipvx9;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437358] SERVICE ALERT: nycnn-ipvx9;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437358] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;WARNING;HARD;3;WARNING (Sample Period 720 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437358] SERVICE ALERT: nycnn-editxfer;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1407437368] SERVICE ALERT: nycnn-ipvx8;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437368] SERVICE ALERT: nycnn-ipvx8;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437368] SERVICE ALERT: nycnn-ipvx9;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437368] SERVICE ALERT: nycnn-ipvx9e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437368] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407437368] SERVICE ALERT: nycnn-ipvx8;CPU Load;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407437388] SERVICE ALERT: nycnn-ipvx9e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437388] SERVICE ALERT: nycnn-ipvx8e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437388] SERVICE ALERT: nycnn-ipvx8e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437398] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437398] SERVICE ALERT: nycnn-ipvx9;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437398] HOST ALERT: nycnn-ipvx8;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407437398] HOST ALERT: nycnn-ipvx8e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1407437398] HOST ALERT: nycnn-ipvx9e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1407437398] HOST ALERT: nycnn-ipvx9;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1407437408] SERVICE ALERT: nycnn-ipvx8;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407437408] SERVICE ALERT: nycnn-ipvx8e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407437408] SERVICE ALERT: nycnn-ipvx9e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1407437408] SERVICE ALERT: nycnn-ipvx9;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407437468] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437478] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.113)
[1407437488] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437498] HOST ALERT: nycnn-editxfer;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.120)
[1407437518] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407437528] SERVICE ALERT: nycnn-editxfer4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437538] HOST ALERT: nycnn-editxfer4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.114)
[1407437538] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407437548] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407437568] HOST ALERT: nycnn-editxfer;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407437578] SERVICE ALERT: nycnn-editxfer4;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.114)
[1407437608] HOST ALERT: nycnn-editxfer4;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.114)
[1407437618] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.113)
[1407437628] SERVICE ALERT: nycnn-ipvx8e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.817GB (41%) - Free: 1.173GB (59%)
[1407437628] SERVICE ALERT: nycnn-ipvx9e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437628] SERVICE ALERT: nycnn-ipvx9e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407437628] SERVICE ALERT: nycnn-ipvx9;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407437638] HOST ALERT: nycnn-editxfer;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.120)
[1407437638] SERVICE ALERT: nycnn-ipvx9;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437648] SERVICE ALERT: nycnn-ipvx12e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437648] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437648] SERVICE ALERT: nycnn-ipvx8;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437648] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437648] SERVICE ALERT: nycnn-ipvx12e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437648] SERVICE ALERT: nycnn-ipvx8;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407437648] SERVICE ALERT: nycnn-ipvx9;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.822GB (41%) - Free: 1.168GB (59%)
[1407437648] SERVICE ALERT: nycnn-ipvx9e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.822GB (41%) - Free: 1.168GB (59%)
[1407437658] SERVICE ALERT: nycnn-ipvx5e1;CPU Load;OK;HARD;3;OK (Sample Period 300 sec) - Average CPU Utilisation 0.05%
[1407437658] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;WARNING;HARD;3;WARNING (Sample Period 720 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437658] SERVICE ALERT: nycnn-ipvx8;CPU Load;WARNING;HARD;3;WARNING (Sample Period 720 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437658] SERVICE ALERT: nycnn-editxfer;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1407437668] HOST ALERT: nycnn-editxfer;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407437678] SERVICE ALERT: nycnn-ipvx9e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407437678] HOST ALERT: nycnn-editxfer4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1407437678] HOST ALERT: nycnn-ipvx12e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407437678] HOST ALERT: nycnn-ipvx12;DOWN;SOFT;1;(Host Check Timed Out)
[1407437688] HOST ALERT: nycnn-editxfer3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1407437688] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437688] SERVICE ALERT: nycnn-ipvx8e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437688] SERVICE ALERT: nycnn-ipvx8e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407437688] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437698] SERVICE ALERT: nycnn-ipvx9;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437698] SERVICE ALERT: nycnn-ipvx9;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1407437698] SERVICE ALERT: nycnn-editxfer3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1407437708] SERVICE ALERT: nycnn-editxfer4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1407437708] SERVICE ALERT: nycnn-ipvx12e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437708] SERVICE ALERT: nycnn-ipvx12;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437708] SERVICE ALERT: nycnn-ipvx12e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437708] SERVICE ALERT: nycnn-ipvx12e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437708] SERVICE ALERT: nycnn-ipvx12e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437708] SERVICE ALERT: nycnn-ipvx12;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437718] SERVICE ALERT: nycnn-ipvx12;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437738] SERVICE ALERT: nycnn-ipvx11e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407437738] SERVICE ALERT: nycnn-ipvx11;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407437748] SERVICE ALERT: nycnn-ipvx8e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.01%
[1407437748] SERVICE ALERT: nycnn-ipvx12e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437748] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437748] SERVICE ALERT: nycnn-ipvx12;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437768] SERVICE ALERT: nycnn-ipvx12;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437768] SERVICE ALERT: nycnn-ipvx12;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437768] HOST ALERT: nycnn-ipvx12e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.43 ms
[1407437768] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437768] SERVICE ALERT: nycnn-ipvx12;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 10.49 ms
[1407437768] SERVICE ALERT: nycnn-ipvx12e1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 44%, RTA = 0.30 ms
[1407437768] HOST ALERT: nycnn-ipvx12;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.27 ms
[1407437768] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407437768] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407437768] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;1;(Host Check Timed Out)
[1407437808] SERVICE ALERT: nycnn-ipvx11e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437808] SERVICE ALERT: nycnn-ipvx11;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407437818] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437818] SERVICE ALERT: nycnn-ipvx12;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1407437828] SERVICE ALERT: nycnn-ipvx12e1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407437828] SERVICE ALERT: nycnn-ipvx11;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437828] SERVICE ALERT: nycnn-ipvx11e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407437828] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437858] HOST ALERT: nycnn-ipvx11e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407437868] HOST ALERT: nycnn-ipvx11;DOWN;SOFT;2;(Host Check Timed Out)
[1407437868] SERVICE ALERT: nycnn-ipvx11;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437878] SERVICE ALERT: nycnn-ipvx11e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407437888] SERVICE ALERT: nycnn-ipvx11e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 44%, RTA = 0.44 ms
[1407437888] SERVICE ALERT: nycnn-ipvx11;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437888] SERVICE ALERT: nycnn-ipvx11;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407437928] HOST ALERT: nycnn-ipvx11e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407437928] SERVICE ALERT: nycnn-ipvx11;CPU Load;WARNING;HARD;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437938] HOST ALERT: nycnn-ipvx11;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1407437938] SERVICE ALERT: nycnn-ipvx11;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407437948] SERVICE ALERT: nycnn-ipvx11e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407437958] SERVICE ALERT: nycnn-ipvx9e1;CPU Load;OK;HARD;3;OK (Sample Period 300 sec) - Average CPU Utilisation 0.06%
[1407437958] SERVICE ALERT: nycnn-ipvx8;CPU Load;OK;HARD;3;OK (Sample Period 300 sec) - Average CPU Utilisation 0.05%
[1407437978] SERVICE ALERT: nycnn-ipvx12e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407437988] SERVICE ALERT: nycnn-ipvx12e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437998] SERVICE ALERT: nycnn-ipvx12e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407437998] SERVICE ALERT: nycnn-ipvx12e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407437998] SERVICE ALERT: nycnn-ipvx12;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407437998] SERVICE ALERT: nycnn-ipvx12;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438018] SERVICE ALERT: nycnn-ipvx11e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438038] SERVICE ALERT: nycnn-ipvx12e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438038] SERVICE ALERT: nycnn-ipvx12;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438038] SERVICE ALERT: nycnn-ipvx12;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438048] SERVICE ALERT: nycnn-ipvx12;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.845GB (42%) - Free: 1.145GB (58%)
[1407438048] SERVICE ALERT: nycnn-ipvx12e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.16%
[1407438058] SERVICE ALERT: nycnn-ipvx12;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.12%
[1407438068] SERVICE ALERT: nycnn-ipvx12;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438078] SERVICE ALERT: nycnn-ipvx11e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.28%
[1407438098] SERVICE ALERT: nycnn-ipvx11e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.827GB (42%) - Free: 1.163GB (58%)
[1407438098] SERVICE ALERT: nycnn-ipvx11;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438108] SERVICE ALERT: nycnn-ipvx11;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438118] SERVICE ALERT: nycnn-ipvx11e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438148] SERVICE ALERT: nycnn-ipvx11;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438158] SERVICE ALERT: nycnn-ipvx11e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438188] SERVICE ALERT: nycnn-ipvx11;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.824GB (41%) - Free: 1.166GB (59%)
[1407438188] SERVICE ALERT: nycnn-ipvx11;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438198] SERVICE ALERT: nycnn-ipvx20;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438228] SERVICE ALERT: nycnn-ipvx11;CPU Load;OK;HARD;1;OK (Sample Period 300 sec) - Average CPU Utilisation 0.74%
[1407438228] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438228] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;1;(Host Check Timed Out)
[1407438238] SERVICE ALERT: nycnn-ipvx20;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438248] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438248] SERVICE ALERT: nycnn-ipvx21;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438248] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438248] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438248] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438248] SERVICE ALERT: nycnn-ipvx20;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1407438248] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438248] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438248] SERVICE ALERT: nycnn-ipvx21e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438248] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438258] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438258] HOST ALERT: nycnn-ipvx20e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438258] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.113)
[1407438258] SERVICE ALERT: nycnn-ipvx21;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438268] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438268] SERVICE ALERT: nycnn-ipvx20;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438278] HOST ALERT: nycnn-editxfer;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.120)
[1407438288] HOST ALERT: nycnn-ipvx21;DOWN;SOFT;1;(Host Check Timed Out)
[1407438288] HOST ALERT: nycnn-ipvx19;DOWN;SOFT;1;(Host Check Timed Out)
[1407438288] HOST ALERT: nycnn-ipvx19e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438288] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438298] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438298] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407438308] SERVICE ALERT: nycnn-ipvx20e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx21e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438308] SERVICE ALERT: nycnn-ipvx21;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx19e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx20e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.113)
[1407438308] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx21;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx19;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438308] SERVICE ALERT: nycnn-ipvx19e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438308] SERVICE ALERT: nycnn-ipvx20e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438308] SERVICE ALERT: nycnn-ipvx21;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438308] SERVICE ALERT: nycnn-ipvx21e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438308] SERVICE ALERT: nycnn-ipvx21e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438308] SERVICE ALERT: nycnn-ipvx21;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438318] SERVICE ALERT: nycnn-ipvx21e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438318] SERVICE ALERT: nycnn-ipvx20e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438318] SERVICE ALERT: nycnn-editxfer;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407438328] HOST ALERT: nycnn-ipvx20;DOWN;SOFT;2;(Host Check Timed Out)
[1407438328] SERVICE ALERT: nycnn-ipvx21;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438328] HOST ALERT: nycnn-editxfer;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.120)
[1407438338] SERVICE ALERT: nycnn-ipvx19;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407438338] SERVICE ALERT: nycnn-ipvx19e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407438338] HOST ALERT: nycnn-ipvx19;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1407438338] HOST ALERT: nycnn-ipvx19e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1407438338] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407438348] SERVICE ALERT: nycnn-ipvx20e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407438348] HOST ALERT: nycnn-ipvx20e1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 1.28 ms
[1407438348] SERVICE ALERT: nycnn-ipvx22;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438348] SERVICE ALERT: nycnn-ipvx20;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438348] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438348] HOST ALERT: nycnn-ipvx21e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407438358] SERVICE ALERT: nycnn-ipvx21;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407438358] SERVICE ALERT: nycnn-ipvx19;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407438368] SERVICE ALERT: nycnn-ipvx19e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1407438368] SERVICE ALERT: nycnn-ipvx20;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407438368] SERVICE ALERT: nycnn-ipvx20e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407438368] SERVICE ALERT: nycnn-ipvx21;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.52 ms
[1407438368] SERVICE ALERT: nycnn-ipvx21e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.88 ms
[1407438368] HOST ALERT: nycnn-ipvx21;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.31 ms
[1407438368] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438368] SERVICE ALERT: nycnn-ipvx22;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438368] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438378] SERVICE ALERT: nycnn-ipvx22e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438378] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438378] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.113)
[1407438378] SERVICE ALERT: nycnn-ipvx22;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438378] HOST ALERT: nycnn-ipvx20;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407438378] HOST ALERT: nycnn-ipvx21e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407438378] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438388] HOST ALERT: nycnn-ipvx22;DOWN;SOFT;1;(Host Check Timed Out)
[1407438388] HOST ALERT: nycnn-ipvx22e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438398] HOST ALERT: nycnn-editxfer;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.164.219.120)
[1407438398] SERVICE ALERT: nycnn-ipvx19e1;SafeNet Service;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438408] SERVICE ALERT: nycnn-ipvx22;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438408] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438428] SERVICE ALERT: nycnn-ipvx22;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438428] SERVICE ALERT: nycnn-ipvx22;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438428] SERVICE ALERT: nycnn-ipvx22e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438438] SERVICE ALERT: nycnn-ipvx22e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438438] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407438438] SERVICE ALERT: nycnn-ipvx22;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438448] HOST ALERT: nycnn-editxfer3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1407438458] SERVICE ALERT: nycnn-ipvx19;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438458] HOST ALERT: nycnn-ipvx22;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407438458] HOST ALERT: nycnn-ipvx22e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1407438468] HOST ALERT: nycnn-editxfer;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1407438478] SERVICE ALERT: nycnn-editxfer3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407438478] SERVICE ALERT: nycnn-ipvx22;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407438488] SERVICE ALERT: nycnn-ipvx22e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407438498] SERVICE ALERT: nycnn-editxfer;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1407438518] SERVICE ALERT: nycnn-ipvx19;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.35%
[1407438528] SERVICE ALERT: nycnn-ipvx20;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.835GB (42%) - Free: 1.155GB (58%)
[1407438538] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 287 processes with STATE = RSZDT
[1407438558] SERVICE ALERT: nycnn-ipvx20;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438578] SERVICE ALERT: nycnn-ipvx23e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438578] SERVICE ALERT: nycnn-ipvx20;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438578] SERVICE ALERT: nycnn-ipvx23;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438588] SERVICE ALERT: nycnn-ipvx20e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438588] SERVICE ALERT: nycnn-ipvx20e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.832GB (42%) - Free: 1.158GB (58%)
[1407438588] SERVICE ALERT: nycnn-ipvx19e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438588] SERVICE ALERT: nycnn-ipvx21e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438588] SERVICE ALERT: nycnn-ipvx21;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438588] SERVICE ALERT: nycnn-ipvx20e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438588] SERVICE ALERT: nycnn-ipvx20;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438588] SERVICE ALERT: nycnn-ipvx19e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438588] SERVICE ALERT: nycnn-ipvx21;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438588] SERVICE ALERT: nycnn-ipvx21e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438598] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 102 processes with STATE = RSZDT
[1407438598] SERVICE ALERT: nycnn-ipvx21e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438598] SERVICE ALERT: nycnn-ipvx21;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438598] SERVICE ALERT: nycnn-ipvx21e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438598] SERVICE ALERT: nycnn-ipvx20e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438608] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.19 ms
[1407438608] SERVICE ALERT: nycnn-ipvx24;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438608] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438608] SERVICE ALERT: nycnn-ipvx21;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.829GB (42%) - Free: 1.161GB (58%)
[1407438608] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438608] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1407438618] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438618] SERVICE ALERT: nycnn-ipvx24e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438638] SERVICE ALERT: nycnn-ipvx19e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438638] SERVICE ALERT: nycnn-ipvx19;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.833GB (42%) - Free: 1.157GB (58%)
[1407438638] SERVICE ALERT: nycnn-ipvx20;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438638] SERVICE ALERT: nycnn-ipvx20e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438648] HOST ALERT: nycnn-ipvx23e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438648] HOST ALERT: nycnn-ipvx24;DOWN;SOFT;1;(Host Check Timed Out)
[1407438648] HOST ALERT: nycnn-ipvx23;DOWN;SOFT;1;(Host Check Timed Out)
[1407438648] HOST ALERT: nycnn-ipvx24e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438648] SERVICE ALERT: nycnn-ipvx23e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx24;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx23;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx25e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx21;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.11%
[1407438648] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx25e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438648] SERVICE ALERT: nycnn-ipvx24e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438658] SERVICE ALERT: nycnn-ipvx21;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438658] SERVICE ALERT: nycnn-ipvx23;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438658] SERVICE ALERT: nycnn-ipvx21e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.12%
[1407438658] SERVICE ALERT: nycnn-ipvx20e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.11%
[1407438668] SERVICE ALERT: nycnn-ipvx23e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438668] SERVICE ALERT: nycnn-ipvx24;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438668] SERVICE ALERT: nycnn-ipvx25;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438668] SERVICE ALERT: nycnn-ipvx22e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438668] SERVICE ALERT: nycnn-ipvx24;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438668] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438668] SERVICE ALERT: nycnn-ipvx23;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438668] SERVICE ALERT: nycnn-ipvx24e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438668] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438678] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438678] SERVICE ALERT: nycnn-ipvx23;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438678] SERVICE ALERT: nycnn-ipvx24e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438688] HOST ALERT: nycnn-ipvx25e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438698] SERVICE ALERT: nycnn-ipvx22;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438698] SERVICE ALERT: nycnn-ipvx19e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.05%
[1407438698] SERVICE ALERT: nycnn-ipvx20;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.01%
[1407438698] SERVICE ALERT: nycnn-ipvx22e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438708] HOST ALERT: nycnn-ipvx25;DOWN;SOFT;1;(Host Check Timed Out)
[1407438708] SERVICE ALERT: nycnn-ipvx25e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438708] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438708] SERVICE ALERT: nycnn-ipvx25e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438708] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438718] SERVICE ALERT: nycnn-ipvx22;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438718] SERVICE ALERT: nycnn-ipvx22e1;Memory Usage;OK;HARD;3;OK - Physical Memory: Total: 1.990GB - Used: 0.828GB (42%) - Free: 1.162GB (58%)
[1407438718] SERVICE ALERT: nycnn-ipvx22e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438718] SERVICE ALERT: nycnn-ipvx23e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438718] SERVICE ALERT: nycnn-ipvx22;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438728] SERVICE ALERT: nycnn-ipvx25;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407438728] SERVICE ALERT: nycnn-ipvx24;CPU Load;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407438728] SERVICE ALERT: nycnn-ipvx25e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438728] SERVICE ALERT: nycnn-ipvx24e1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 54%, RTA = 0.49 ms
[1407438728] HOST ALERT: nycnn-ipvx24;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.29 ms
[1407438728] HOST ALERT: nycnn-ipvx24e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.31 ms
[1407438728] SERVICE ALERT: nycnn-ipvx25;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407438738] HOST ALERT: nycnn-ipvx23e1;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 1.29 ms
[1407438738] HOST ALERT: nycnn-ipvx23;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.63 ms
[1407438748] SERVICE ALERT: nycnn-ipvx25;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438758] SERVICE ALERT: nycnn-ipvx22;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.12%
[1407438758] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438768] SERVICE ALERT: nycnn-ipvx25;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407438768] HOST ALERT: nycnn-ipvx25e1;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.47 ms
[1407438778] HOST ALERT: nycnn-ipvx25;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407438778] SERVICE ALERT: nycnn-ipvx23e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407438778] SERVICE ALERT: nycnn-ipvx22e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.06%
[1407438778] SERVICE ALERT: nycnn-ipvx24;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1407438778] SERVICE ALERT: nycnn-ipvx25e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407438788] SERVICE ALERT: nycnn-ipvx23;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1407438788] SERVICE ALERT: nycnn-ipvx24e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1407438788] SERVICE ALERT: nycnn-ipvx25;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1407438828] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 299 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438828] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438828] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438888] SERVICE ALERT: nycnn-ipvx25e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.44%
[1407438888] SERVICE ALERT: nycnn-ipvx23e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.31%
[1407438888] SERVICE ALERT: nycnn-ipvx24e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.43%
[1407438938] SERVICE ALERT: nycnn-ipvx23e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438938] SERVICE ALERT: nycnn-ipvx29e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438938] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438938] SERVICE ALERT: nycnn-ipvx23;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407438938] SERVICE ALERT: nycnn-ipvx24;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.829GB (42%) - Free: 1.161GB (58%)
[1407438938] SERVICE ALERT: nycnn-ipvx28e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438938] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438938] SERVICE ALERT: nycnn-ipvx28e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438938] SERVICE ALERT: nycnn-ipvx23;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438938] SERVICE ALERT: nycnn-ipvx24e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438938] SERVICE ALERT: nycnn-ipvx24e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438938] SERVICE ALERT: nycnn-ipvx23;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.840GB (42%) - Free: 1.150GB (58%)
[1407438948] SERVICE ALERT: nycnn-ipvx28;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438948] SERVICE ALERT: nycnn-ipvx24;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438948] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438948] SERVICE ALERT: nycnn-ipvx27;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438958] SERVICE ALERT: nycnn-ipvx29;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438958] SERVICE ALERT: nycnn-ipvx24;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438968] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438968] SERVICE ALERT: nycnn-ipvx30e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438968] SERVICE ALERT: nycnn-ipvx30;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407438968] SERVICE ALERT: nycnn-ipvx23;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407438968] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.18 ms
[1407438968] SERVICE ALERT: nycnn-ipvx24e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407438968] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438968] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438968] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438968] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407438968] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438968] SERVICE ALERT: nycnn-ipvx28e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438978] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;1;(Host Check Timed Out)
[1407438978] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438978] SERVICE ALERT: nycnn-ipvx28;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1407438978] SERVICE ALERT: nycnn-ipvx27e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438978] SERVICE ALERT: nycnn-ipvx27;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438978] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438978] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407438988] HOST ALERT: nycnn-ipvx27;DOWN;SOFT;1;(Host Check Timed Out)
[1407438998] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;1;(Host Check Timed Out)
[1407438998] HOST ALERT: nycnn-ipvx27e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407438998] SERVICE ALERT: nycnn-ipvx25e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.825GB (41%) - Free: 1.165GB (59%)
[1407438998] SERVICE ALERT: nycnn-ipvx25e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438998] SERVICE ALERT: nycnn-ipvx23;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1407438998] SERVICE ALERT: nycnn-ipvx25e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407438998] SERVICE ALERT: nycnn-ipvx25e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439008] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407439008] SERVICE ALERT: nycnn-ipvx23e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439008] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439008] SERVICE ALERT: nycnn-ipvx25;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439008] SERVICE ALERT: nycnn-ipvx29e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439008] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439008] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439008] SERVICE ALERT: nycnn-ipvx28e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439008] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439008] SERVICE ALERT: nycnn-ipvx28e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439008] SERVICE ALERT: nycnn-ipvx27;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439018] SERVICE ALERT: nycnn-ipvx28;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439028] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439028] SERVICE ALERT: nycnn-ipvx29;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439028] SERVICE ALERT: nycnn-ipvx27;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439028] SERVICE ALERT: nycnn-ipvx27e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439028] SERVICE ALERT: nycnn-ipvx28e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439028] SERVICE ALERT: nycnn-ipvx29e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439028] SERVICE ALERT: nycnn-ipvx24;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439028] SERVICE ALERT: nycnn-ipvx25;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439028] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439028] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439028] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439028] SERVICE ALERT: nycnn-ipvx29;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439028] SERVICE ALERT: nycnn-ipvx28e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439038] SERVICE ALERT: nycnn-ipvx27e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439038] SERVICE ALERT: nycnn-ipvx30e1;Service: Apache;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439038] SERVICE ALERT: nycnn-ipvx30;Service: Apache;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439038] SERVICE ALERT: nycnn-ipvx27;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439038] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439038] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439038] HOST ALERT: nycnn-ipvx30e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407439048] SERVICE ALERT: nycnn-ipvx25e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439048] SERVICE ALERT: nycnn-ipvx30;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439058] HOST ALERT: nycnn-ipvx30;DOWN;SOFT;1;(Host Check Timed Out)
[1407439068] SERVICE ALERT: nycnn-ipvx25;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439068] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439068] SERVICE ALERT: nycnn-ipvx27;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439068] HOST ALERT: nycnn-ipvx27e1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.44 ms
[1407439068] HOST ALERT: nycnn-ipvx27;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 1.22 ms
[1407439068] HOST ALERT: nycnn-ipvx29e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407439068] HOST ALERT: nycnn-ipvx28;DOWN;SOFT;2;(Host Check Timed Out)
[1407439068] HOST ALERT: nycnn-ipvx28e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407439078] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439078] SERVICE ALERT: nycnn-ipvx28e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439078] SERVICE ALERT: nycnn-ipvx27;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407439078] SERVICE ALERT: nycnn-ipvx27e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407439088] SERVICE ALERT: nycnn-ipvx25;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.07%
[1407439088] SERVICE ALERT: nycnn-ipvx24;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.02%
[1407439088] SERVICE ALERT: nycnn-ipvx28;PING;WARNING;HARD;1;PING WARNING - Packet loss = 28%, RTA = 0.46 ms
[1407439088] SERVICE ALERT: nycnn-ipvx30e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439088] SERVICE ALERT: nycnn-ipvx30;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439088] SERVICE ALERT: nycnn-ipvx28e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439088] SERVICE ALERT: nycnn-ipvx27e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439098] HOST ALERT: nycnn-ipvx29;DOWN;SOFT;2;(Host Check Timed Out)
[1407439098] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439098] SERVICE ALERT: nycnn-ipvx29e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439098] SERVICE ALERT: nycnn-ipvx30e1;Service: Apache;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407439098] SERVICE ALERT: nycnn-ipvx30;Service: Apache;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407439108] SERVICE ALERT: nycnn-ipvx30;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439128] SERVICE ALERT: nycnn-ipvx30;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439128] HOST ALERT: nycnn-ipvx30e1;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.45 ms
[1407439128] HOST ALERT: nycnn-ipvx30;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1407439138] HOST ALERT: nycnn-ipvx28;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1407439138] HOST ALERT: nycnn-ipvx29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1407439138] HOST ALERT: nycnn-ipvx29e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407439138] HOST ALERT: nycnn-ipvx28e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407439138] SERVICE ALERT: nycnn-ipvx29;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407439138] SERVICE ALERT: nycnn-ipvx28e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407439138] SERVICE ALERT: nycnn-ipvx29e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407439148] SERVICE ALERT: nycnn-ipvx30e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407439148] SERVICE ALERT: nycnn-ipvx30;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407439148] SERVICE ALERT: nycnn-ipvx30e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439148] SERVICE ALERT: nycnn-ipvx28;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407439148] SERVICE ALERT: nycnn-ipvx27e1;CPU Load;OK;SOFT;2;OK (Sample Period 59 sec) - Average CPU Utilisation 6.81%
[1407439188] SERVICE ALERT: nycnn-ipvx29;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439188] SERVICE ALERT: nycnn-ipvx30e1;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 1.990GB - Used: 0.824GB (41%) - Free: 1.166GB (59%)
[1407439198] SERVICE ALERT: nycnn-ipvx27;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439248] SERVICE ALERT: nycnn-ipvx29;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.42%
[1407439258] SERVICE ALERT: nycnn-ipvx27;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.31%
[1407439298] SERVICE ALERT: nycnn-ipvx28;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439298] SERVICE ALERT: nycnn-ipvx29e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439298] SERVICE ALERT: nycnn-ipvx27;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439298] SERVICE ALERT: nycnn-ipvx27;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.823GB (41%) - Free: 1.167GB (59%)
[1407439298] SERVICE ALERT: nycnn-ipvx28e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439298] SERVICE ALERT: nycnn-ipvx28e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439298] SERVICE ALERT: nycnn-ipvx28e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439308] SERVICE ALERT: nycnn-ipvx28;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439308] SERVICE ALERT: nycnn-ipvx27e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439318] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.822GB (41%) - Free: 1.168GB (59%)
[1407439318] SERVICE ALERT: nycnn-ipvx28e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.836GB (42%) - Free: 1.154GB (58%)
[1407439318] SERVICE ALERT: nycnn-ipvx29;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439328] SERVICE ALERT: nycnn-ipvx27e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439328] SERVICE ALERT: nycnn-ipvx27;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439328] SERVICE ALERT: nycnn-ipvx27e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439328] SERVICE ALERT: nycnn-ipvx29e1;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439358] SERVICE ALERT: nycnn-ipvx30e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439358] SERVICE ALERT: nycnn-ipvx29e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1407439358] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439368] SERVICE ALERT: nycnn-ipvx27;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439368] SERVICE ALERT: nycnn-ipvx28;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.10%
[1407439368] SERVICE ALERT: nycnn-ipvx28e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439368] SERVICE ALERT: nycnn-ipvx33;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439388] SERVICE ALERT: nycnn-ipvx30;Service: Apache;OK;HARD;3;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439388] SERVICE ALERT: nycnn-ipvx30e1;Service: Apache;OK;HARD;3;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439388] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439388] SERVICE ALERT: nycnn-ipvx29e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439388] SERVICE ALERT: nycnn-ipvx29;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439388] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439388] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439388] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439388] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439388] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439388] SERVICE ALERT: nycnn-ipvx28e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439388] SERVICE ALERT: nycnn-ipvx30;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439398] HOST ALERT: nycnn-ipvx33;DOWN;SOFT;1;(Host Check Timed Out)
[1407439398] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407439408] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439408] SERVICE ALERT: nycnn-ipvx32;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439408] SERVICE ALERT: nycnn-ipvx32e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439408] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439418] SERVICE ALERT: nycnn-ipvx30;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439418] SERVICE ALERT: nycnn-ipvx30e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.16%
[1407439428] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;1;(Host Check Timed Out)
[1407439428] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407439428] HOST ALERT: nycnn-ipvx32;DOWN;SOFT;1;(Host Check Timed Out)
[1407439428] HOST ALERT: nycnn-ipvx33e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407439428] SERVICE ALERT: nycnn-ipvx33;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439438] SERVICE ALERT: nycnn-ipvx31e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx32;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx31;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx33e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx35;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx33;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx35;Service: Apache;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx32e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx31;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx31e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx32;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx33;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx33e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439448] SERVICE ALERT: nycnn-ipvx28e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.03%
[1407439448] SERVICE ALERT: nycnn-ipvx30;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.09%
[1407439448] SERVICE ALERT: nycnn-ipvx34;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439448] SERVICE ALERT: nycnn-ipvx31e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439458] SERVICE ALERT: nycnn-ipvx31;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439458] SERVICE ALERT: nycnn-ipvx32e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439458] SERVICE ALERT: nycnn-ipvx32e1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439458] SERVICE ALERT: nycnn-ipvx31e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439468] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439468] SERVICE ALERT: nycnn-ipvx32;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439478] HOST ALERT: nycnn-ipvx33e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.31 ms
[1407439478] HOST ALERT: nycnn-ipvx33;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.45 ms
[1407439478] HOST ALERT: nycnn-ipvx35;DOWN;SOFT;1;(Host Check Timed Out)
[1407439478] HOST ALERT: nycnn-ipvx35e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407439478] HOST ALERT: nycnn-ipvx34;DOWN;SOFT;1;(Host Check Timed Out)
[1407439478] HOST ALERT: nycnn-ipvx34e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407439498] HOST ALERT: nycnn-ipvx32e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407439498] SERVICE ALERT: nycnn-ipvx33;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439498] SERVICE FLAPPING ALERT: nycnn-editxfer3;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1407439498] HOST FLAPPING ALERT: nycnn-editxfer3;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1407439508] SERVICE ALERT: nycnn-ipvx33;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407439508] SERVICE ALERT: nycnn-ipvx33e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407439508] SERVICE ALERT: nycnn-ipvx34;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx31;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx32;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx31;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx32;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx34e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx35;Service: Apache;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx35;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx35e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439508] SERVICE ALERT: nycnn-ipvx34;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx34;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439508] SERVICE ALERT: nycnn-ipvx34e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439508] SERVICE ALERT: nycnn-ipvx35;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407439508] SERVICE ALERT: nycnn-ipvx32e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439508] SERVICE ALERT: nycnn-ipvx35;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439508] SERVICE ALERT: nycnn-ipvx34;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439518] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439518] HOST ALERT: nycnn-ipvx32;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.94 ms
[1407439518] SERVICE ALERT: nycnn-ipvx31e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439518] SERVICE FLAPPING ALERT: nycnn-editxfer;PING;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1407439518] HOST FLAPPING ALERT: nycnn-editxfer;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1407439518] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439528] HOST ALERT: nycnn-ipvx31;DOWN;SOFT;2;(Host Check Timed Out)
[1407439528] HOST ALERT: nycnn-ipvx31e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407439538] SERVICE ALERT: nycnn-ipvx33e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439538] SERVICE ALERT: nycnn-ipvx31;CPU Load;WARNING;HARD;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439538] SERVICE ALERT: nycnn-ipvx32;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439538] SERVICE ALERT: nycnn-ipvx33;Process: Siphon;OK;SOFT;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439548] SERVICE ALERT: nycnn-ipvx35e1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439548] SERVICE ALERT: nycnn-ipvx35e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407439548] HOST ALERT: nycnn-ipvx35;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407439548] HOST ALERT: nycnn-ipvx35e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407439548] SERVICE ALERT: nycnn-ipvx34e1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439548] SERVICE ALERT: nycnn-ipvx34e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439558] HOST ALERT: nycnn-ipvx34;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.48 ms
[1407439558] HOST ALERT: nycnn-ipvx34e1;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.44 ms
[1407439558] SERVICE ALERT: nycnn-ipvx35e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407439558] SERVICE ALERT: nycnn-ipvx32e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1407439558] SERVICE ALERT: nycnn-ipvx31;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407439558] SERVICE ALERT: nycnn-ipvx31e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407439558] SERVICE ALERT: nycnn-ipvx32;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407439568] SERVICE ALERT: nycnn-ipvx34;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407439568] SERVICE ALERT: nycnn-ipvx34e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1407439568] SERVICE ALERT: nycnn-ipvx35;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407439568] SERVICE ALERT: nycnn-ipvx35;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439568] HOST ALERT: nycnn-ipvx32e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1407439568] HOST ALERT: nycnn-ipvx31;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407439568] HOST ALERT: nycnn-ipvx31e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407439598] SERVICE ALERT: nycnn-ipvx33e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.97%
[1407439598] SERVICE ALERT: nycnn-ipvx32;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 6.21%
[1407439608] SERVICE ALERT: nycnn-ipvx35;Process: XCode;OK;SOFT;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439628] SERVICE ALERT: nycnn-ipvx31e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439628] SERVICE ALERT: nycnn-ipvx34;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439628] SERVICE ALERT: nycnn-ipvx32e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439638] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439648] SERVICE ALERT: nycnn-ipvx35;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439658] SERVICE ALERT: nycnn-ipvx34e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439688] SERVICE ALERT: nycnn-ipvx32e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439688] SERVICE ALERT: nycnn-ipvx33;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439688] SERVICE ALERT: nycnn-ipvx31e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.28%
[1407439688] SERVICE ALERT: nycnn-ipvx34;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.32%
[1407439688] SERVICE ALERT: nycnn-ipvx32e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.29%
[1407439698] SERVICE ALERT: nycnn-ipvx35e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.54%
[1407439708] SERVICE ALERT: nycnn-ipvx35;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.55%
[1407439718] SERVICE ALERT: nycnn-ipvx34e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.41%
[1407439718] SERVICE ALERT: nycnn-ipvx33;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407439728] SERVICE ALERT: nycnn-ipvx31e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439728] SERVICE ALERT: nycnn-ipvx32;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.834GB (42%) - Free: 1.156GB (58%)
[1407439728] SERVICE ALERT: nycnn-ipvx33e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.843GB (42%) - Free: 1.146GB (58%)
[1407439728] SERVICE ALERT: nycnn-ipvx33;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439728] SERVICE ALERT: nycnn-ipvx31;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439738] SERVICE ALERT: nycnn-ipvx31e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439748] SERVICE ALERT: nycnn-ipvx31;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439748] SERVICE ALERT: nycnn-ipvx32e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.832GB (42%) - Free: 1.158GB (58%)
[1407439748] SERVICE ALERT: nycnn-ipvx32e1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439748] SERVICE ALERT: nycnn-ipvx31e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439748] SERVICE ALERT: nycnn-ipvx31e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439748] SERVICE ALERT: nycnn-ipvx32;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439778] SERVICE ALERT: nycnn-ipvx33;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.15%
[1407439788] SERVICE ALERT: nycnn-ipvx31;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439788] SERVICE ALERT: nycnn-ipvx32;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439788] SERVICE ALERT: nycnn-ipvx31;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439788] SERVICE ALERT: nycnn-ipvx32;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439788] SERVICE ALERT: nycnn-ipvx34;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439788] SERVICE ALERT: nycnn-ipvx34e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439788] SERVICE ALERT: nycnn-ipvx35;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.842GB (42%) - Free: 1.148GB (58%)
[1407439788] SERVICE ALERT: nycnn-ipvx35;Service: Apache;OK;HARD;2;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439798] SERVICE ALERT: nycnn-ipvx34;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439798] SERVICE ALERT: nycnn-ipvx35;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439798] SERVICE ALERT: nycnn-ipvx34;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.831GB (42%) - Free: 1.159GB (58%)
[1407439798] SERVICE ALERT: nycnn-ipvx34;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439798] SERVICE ALERT: nycnn-ipvx35e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439798] SERVICE ALERT: nycnn-ipvx31e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.829GB (42%) - Free: 1.161GB (58%)
[1407439798] SERVICE ALERT: nycnn-ipvx32e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439808] SERVICE ALERT: nycnn-ipvx35;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439822] HOST DOWNTIME ALERT: nycnn-ipvsphn5;STARTED; Host has entered a period of scheduled downtime
[1407439838] SERVICE ALERT: nycnn-ipvx31;CPU Load;OK;HARD;1;OK (Sample Period 300 sec) - Average CPU Utilisation 0.95%
[1407439838] SERVICE ALERT: nycnn-ipvx35e1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.841GB (42%) - Free: 1.149GB (58%)
[1407439838] SERVICE ALERT: nycnn-ipvx34e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407439838] SERVICE ALERT: nycnn-ipvx35e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407439838] SERVICE ALERT: nycnn-ipvx34e1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407439868] SERVICE ALERT: nycnn-ipvsphn5;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407439928] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407439948] SERVICE ALERT: nycnn-ipvsphn5;Process: Siphon;UNKNOWN;SOFT;2;UNKNOWN - Plugin Timed out (15 sec)
[1407439968] HOST ALERT: nycnn-ipvsphn5;DOWN;SOFT;1;(Host Check Timed Out)
[1407439988] SERVICE ALERT: nycnn-ipvsphn5;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407439988] SERVICE ALERT: nycnn-ipvsphn5;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407440008] SERVICE ALERT: nycnn-ipvsphn5;Process: Siphon;UNKNOWN;HARD;3;UNKNOWN - Plugin Timed out (15 sec)
[1407440058] HOST ALERT: nycnn-ipvsphn5;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.38 ms
[1407440098] SERVICE ALERT: nycnn-ipvsphn5;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1407440268] SERVICE ALERT: nycnn-ipvsphn5;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1407440288] SERVICE ALERT: nycnn-ipvsphn5;Process: Siphon;OK;HARD;3;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407440398] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 264 processes with STATE = RSZDT
[1407440458] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 81 processes with STATE = RSZDT
[1407440528] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407440528] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407440558] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;1;(Host Check Timed Out)
[1407440558] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407440588] SERVICE ALERT: nycnn-ipvx10;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407440588] SERVICE ALERT: nycnn-ipvx10e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407440648] SERVICE ALERT: nycnn-ipvx10e1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407440648] SERVICE ALERT: nycnn-ipvx10;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407440658] HOST ALERT: nycnn-ipvx10;DOWN;SOFT;2;(Host Check Timed Out)
[1407440658] HOST ALERT: nycnn-ipvx10e1;DOWN;SOFT;2;(Host Check Timed Out)
[1407440658] SERVICE ALERT: nycnn-ipvx10e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407440708] SERVICE ALERT: nycnn-ipvx10;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1407440708] SERVICE ALERT: nycnn-ipvx10e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1407440708] HOST ALERT: nycnn-ipvx10e1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1407440708] HOST ALERT: nycnn-ipvx10;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407440718] SERVICE ALERT: nycnn-ipvx10;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407440778] SERVICE ALERT: nycnn-ipvx10;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 2.08%
[1407440928] SERVICE ALERT: nycnn-ipvx10e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407440938] SERVICE ALERT: nycnn-ipvx10;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407440948] SERVICE ALERT: nycnn-ipvx10e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407440988] SERVICE ALERT: nycnn-ipvx10e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 1.26%
[1407442048] SERVICE ALERT: nycnn-ipvx4;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442078] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442088] HOST ALERT: nycnn-ipvx4;DOWN;SOFT;1;(Host Check Timed Out)
[1407442088] SERVICE ALERT: nycnn-ipvx4;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1407442088] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407442088] SERVICE ALERT: nycnn-ipvx4e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442108] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442108] SERVICE ALERT: nycnn-ipvx4e1;Process: XCode;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442108] SERVICE ALERT: nycnn-ipvx4;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407442108] HOST ALERT: nycnn-ipvx4e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407442138] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407442148] SERVICE ALERT: nycnn-ipvx4;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442148] SERVICE ALERT: nycnn-ipvx4e1;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442148] SERVICE ALERT: nycnn-ipvx4e1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407442148] SERVICE ALERT: nycnn-ipvx4e1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407442158] SERVICE ALERT: nycnn-ipvx4;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407442168] SERVICE ALERT: nycnn-ipvx4e1;Process: XCode;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407442178] HOST ALERT: nycnn-ipvx4;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.21 ms
[1407442178] HOST ALERT: nycnn-ipvx4e1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1407442208] SERVICE ALERT: nycnn-ipvx4;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1407442208] SERVICE ALERT: nycnn-ipvx4e1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407442258] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1407442318] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 87 processes with STATE = RSZDT
[1407442388] SERVICE ALERT: nycnn-ipvx4;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407442398] SERVICE ALERT: nycnn-ipvx4;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407442428] SERVICE ALERT: nycnn-ipvx4e1;SafeNet Service;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407442428] SERVICE ALERT: nycnn-ipvx4;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.828GB (42%) - Free: 1.162GB (58%)
[1407442428] SERVICE ALERT: nycnn-ipvx4e1;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407442438] SERVICE ALERT: nycnn-ipvx4e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 660 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407442448] SERVICE ALERT: nycnn-ipvx4;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407442448] SERVICE ALERT: nycnn-ipvx4e1;Process: XCode;OK;HARD;2;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407442458] SERVICE ALERT: nycnn-ipvx4;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 1.80%
[1407442498] SERVICE ALERT: nycnn-ipvx4e1;CPU Load;OK;HARD;3;OK (Sample Period 60 sec) - Average CPU Utilisation 1.78%
[1407443568] SERVICE ALERT: nycnn-ipvx3e1;Process: Siphon;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. You might have your username/password wrong or the user's access level is too low. Wmic error text on the next line.
[1407443598] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 63%, RTA = 0.21 ms
[1407443608] SERVICE ALERT: nycnn-ipvx3e1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407443628] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;1;(Host Check Timed Out)
[1407443648] HOST ALERT: nycnn-ipvx3e1;DOWN;SOFT;1;(Host Check Timed Out)
[1407443648] SERVICE ALERT: nycnn-ipvx3e1;Process: Siphon;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407443648] SERVICE ALERT: nycnn-ipvx3;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407443648] SERVICE ALERT: nycnn-ipvx3e1;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1407443658] SERVICE ALERT: nycnn-ipvx3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407443668] SERVICE ALERT: nycnn-ipvx3e1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407443718] SERVICE ALERT: nycnn-ipvx3e1;Process: XCode;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407443718] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407443728] HOST ALERT: nycnn-ipvx3;DOWN;SOFT;2;(Host Check Timed Out)
[1407443728] HOST ALERT: nycnn-ipvx3e1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.45 ms
[1407443748] SERVICE ALERT: nycnn-ipvx3;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407443748] SERVICE ALERT: nycnn-ipvx3e1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407443748] SERVICE ALERT: nycnn-ipvx3;Service: Apache;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407443758] HOST ALERT: nycnn-ipvx3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1407443768] SERVICE ALERT: nycnn-ipvx3e1;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1407443768] SERVICE ALERT: nycnn-ipvx3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1407443798] SERVICE ALERT: nycnn-ipvx3e1;CPU Load;WARNING;SOFT;2;WARNING (Sample Period 360 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407443858] SERVICE ALERT: nycnn-ipvx3e1;CPU Load;OK;SOFT;3;OK (Sample Period 60 sec) - Average CPU Utilisation 0.71%
[1407443928] SERVICE ALERT: nycnn-ipvx3e1;Process: Siphon;OK;HARD;2;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407443928] SERVICE ALERT: nycnn-ipvx3;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 1.990GB - Used: 0.821GB (41%) - Free: 1.169GB (59%)
[1407443948] SERVICE ALERT: nycnn-ipvx3e1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 1.990GB - Used: 0.820GB (41%) - Free: 1.170GB (59%)
[1407444008] SERVICE ALERT: nycnn-ipvx3e1;Process: XCode;OK;HARD;1;OK - Found 1 Instance(s) of "XCodeService.exe" running (0 excluded).  (List is on next line)
[1407444008] SERVICE ALERT: nycnn-ipvx3;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407444038] SERVICE ALERT: nycnn-ipvx3;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 600 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1407444038] SERVICE ALERT: nycnn-ipvx3;Service: Apache;OK;HARD;1;OK - Found 1 Services(s), 1 OK and 0 with problems (0 excluded). 'Apache2.2' (Apache2.2) is Running.
[1407444098] SERVICE ALERT: nycnn-ipvx3;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.08%
[1407446198] HOST DOWNTIME ALERT: nycnn-ipvsphn1;STARTED; Host has entered a period of scheduled downtime
[1407447016] HOST DOWNTIME ALERT: nycnn-ipvsphn5;STOPPED; Host has exited from a period of scheduled downtime
[1407448028] HOST DOWNTIME ALERT: nycnn-ipvx1;STOPPED; Host has exited from a period of scheduled downtime
[1407448095] HOST DOWNTIME ALERT: nycnn-ipvx10;STOPPED; Host has exited from a period of scheduled downtime
[1407448117] HOST DOWNTIME ALERT: nycnn-ipvx11;STOPPED; Host has exited from a period of scheduled downtime
[1407448165] HOST DOWNTIME ALERT: nycnn-ipvx12;STOPPED; Host has exited from a period of scheduled downtime
[1407448195] HOST DOWNTIME ALERT: nycnn-ipvx13;STOPPED; Host has exited from a period of scheduled downtime
[1407448221] HOST DOWNTIME ALERT: nycnn-ipvx14;STOPPED; Host has exited from a period of scheduled downtime
[1407448243] HOST DOWNTIME ALERT: nycnn-ipvx15;STOPPED; Host has exited from a period of scheduled downtime
[1407448263] HOST DOWNTIME ALERT: nycnn-ipvx16;STOPPED; Host has exited from a period of scheduled downtime
[1407448298] HOST DOWNTIME ALERT: nycnn-ipvx17;STOPPED; Host has exited from a period of scheduled downtime
[1407448318] HOST DOWNTIME ALERT: nycnn-ipvx18;STOPPED; Host has exited from a period of scheduled downtime
[1407448335] HOST DOWNTIME ALERT: nycnn-ipvx19;STOPPED; Host has exited from a period of scheduled downtime
[1407448359] HOST DOWNTIME ALERT: nycnn-ipvx2;STOPPED; Host has exited from a period of scheduled downtime
[1407448380] HOST DOWNTIME ALERT: nycnn-ipvx20;STOPPED; Host has exited from a period of scheduled downtime
[1407448398] HOST DOWNTIME ALERT: nycnn-ipvx21;STOPPED; Host has exited from a period of scheduled downtime
[1407448416] HOST DOWNTIME ALERT: nycnn-ipvx22;STOPPED; Host has exited from a period of scheduled downtime
[1407448434] HOST DOWNTIME ALERT: nycnn-ipvx23;STOPPED; Host has exited from a period of scheduled downtime
[1407448448] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407448448] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448451] HOST DOWNTIME ALERT: nycnn-ipvx24;STOPPED; Host has exited from a period of scheduled downtime
[1407448470] HOST DOWNTIME ALERT: nycnn-ipvx25;STOPPED; Host has exited from a period of scheduled downtime
[1407448478] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;1;(Host Check Timed Out)
[1407448487] HOST DOWNTIME ALERT: nycnn-ipvx26;STOPPED; Host has exited from a period of scheduled downtime
[1407448504] HOST DOWNTIME ALERT: nycnn-ipvx27;STOPPED; Host has exited from a period of scheduled downtime
[1407448505] HOST DOWNTIME ALERT: nycnn-ipvsphn4;STARTED; Host has entered a period of scheduled downtime
[1407448508] SERVICE ALERT: nycnn-ipvsphn1;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448508] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448508] SERVICE ALERT: nycnn-ipvsphn1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407448508] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1407448548] SERVICE ALERT: nycnn-ipvsphn1;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448573] HOST DOWNTIME ALERT: nycnn-ipvx28;STOPPED; Host has exited from a period of scheduled downtime
[1407448578] SERVICE ALERT: nycnn-ipvsphn1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448578] HOST ALERT: nycnn-ipvsphn1;DOWN;SOFT;2;(Host Check Timed Out)
[1407448592] HOST DOWNTIME ALERT: nycnn-ipvx29;STOPPED; Host has exited from a period of scheduled downtime
[1407448611] HOST DOWNTIME ALERT: nycnn-ipvx3;STOPPED; Host has exited from a period of scheduled downtime
[1407448618] HOST ALERT: nycnn-ipvsphn1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1407448628] HOST DOWNTIME ALERT: nycnn-ipvx30;STOPPED; Host has exited from a period of scheduled downtime
[1407448628] SERVICE ALERT: nycnn-ipvsphn1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1407448645] HOST DOWNTIME ALERT: nycnn-ipvx31;STOPPED; Host has exited from a period of scheduled downtime
[1407448661] HOST DOWNTIME ALERT: nycnn-ipvx32;STOPPED; Host has exited from a period of scheduled downtime
[1407448676] HOST DOWNTIME ALERT: nycnn-ipvx33;STOPPED; Host has exited from a period of scheduled downtime
[1407448688] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1407448694] HOST DOWNTIME ALERT: nycnn-ipvx34;STOPPED; Host has exited from a period of scheduled downtime
[1407448698] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;1;(Host Check Timed Out)
[1407448708] SERVICE ALERT: nycnn-ipvsphn4;SafeNet Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448710] HOST DOWNTIME ALERT: nycnn-ipvx35;STOPPED; Host has exited from a period of scheduled downtime
[1407448732] HOST DOWNTIME ALERT: nycnn-ipvx4;STOPPED; Host has exited from a period of scheduled downtime
[1407448748] SERVICE ALERT: nycnn-ipvsphn4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1407448751] HOST DOWNTIME ALERT: nycnn-ipvx5;STOPPED; Host has exited from a period of scheduled downtime
[1407448768] HOST DOWNTIME ALERT: nycnn-ipvx6;STOPPED; Host has exited from a period of scheduled downtime
[1407448786] HOST DOWNTIME ALERT: nycnn-ipvx8;STOPPED; Host has exited from a period of scheduled downtime
[1407448788] SERVICE ALERT: nycnn-ipvsphn1;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407448788] SERVICE ALERT: nycnn-ipvsphn1;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407448798] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448798] HOST ALERT: nycnn-ipvsphn4;DOWN;SOFT;2;(Host Check Timed Out)
[1407448798] SERVICE ALERT: nycnn-ipvsphn1;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1407448802] HOST DOWNTIME ALERT: nycnn-ipvx9;STOPPED; Host has exited from a period of scheduled downtime
[1407448808] SERVICE ALERT: nycnn-ipvsphn4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448808] SERVICE ALERT: nycnn-ipvsphn4;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448818] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1407448838] SERVICE ALERT: nycnn-ipvsphn4;Process: Siphon;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1407448838] SERVICE ALERT: nycnn-ipvsphn1;D:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=1.62GB (1.2%), Free=135.07GB (98.8%)
[1407448858] SERVICE ALERT: nycnn-ipvsphn1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.603GB (15%) - Free: 3.397GB (85%)
[1407448868] SERVICE ALERT: nycnn-ipvsphn4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1407448868] HOST ALERT: nycnn-ipvsphn4;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1407448988] SERVICE ALERT: nycnn-ipvsphn4;SafeNet Service;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407449078] SERVICE ALERT: nycnn-ipvsphn4;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.630GB (16%) - Free: 3.369GB (84%)
[1407449088] SERVICE ALERT: nycnn-ipvsphn4;D:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.62GB (100.0%)
[1407449088] SERVICE ALERT: nycnn-ipvsphn4;C:\ Drive Space;OK;HARD;1;OK - D: Total=136.69GB, Used=0.07GB (0.0%), Free=136.62GB (100.0%)
[1407449098] SERVICE ALERT: nycnn-ipvsphn4;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.00%
[1407449138] SERVICE ALERT: nycnn-ipvsphn4;Process: Siphon;OK;HARD;1;OK - Found 1 Instance(s) of "spnsrvnt.exe" running (0 excluded).  (List is on next line)
[1407449764] HOST DOWNTIME ALERT: nycnn-ipvsphn1;STOPPED; Host has exited from a period of scheduled downtime
[1407450612] HOST DOWNTIME ALERT: nycnn-ipvsphn4;STOPPED; Host has exited from a period of scheduled downtime
[1407462718] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1407462778] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 69 processes with STATE = RSZDT