[1406268251] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1406268311] SERVICE ALERT: nycnn-ipvx26e1;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.03%
[1406268401] SERVICE ALERT: nycnn-ipvx26;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1406268461] SERVICE ALERT: nycnn-ipvx26;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.02%
[1406268721] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1406268761] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;1;(Host Check Timed Out)
[1406268781] SERVICE ALERT: nycnn-arvtrtape;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1406268851] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;2;(Host Check Timed Out)
[1406268901] SERVICE ALERT: nycnn-arvtrtape;SSH;CRITICAL;HARD;1;CRITICAL - Socket timeout after 10 seconds
[1406268941] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;3;(Host Check Timed Out)
[1406269041] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;4;(Host Check Timed Out)
[1406269141] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;5;(Host Check Timed Out)
[1406269241] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;6;(Host Check Timed Out)
[1406269341] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;7;(Host Check Timed Out)
[1406269431] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;8;(Host Check Timed Out)
[1406269531] HOST ALERT: nycnn-arvtrtape;DOWN;SOFT;9;(Host Check Timed Out)
[1406269621] HOST ALERT: nycnn-arvtrtape;DOWN;HARD;10;(Host Check Timed Out)
[1406276331] SERVICE ALERT: nycnn-adbeame2;AME Job status;CRITICAL;SOFT;1;Job 10319396 remains in a queued state
[1406276391] SERVICE ALERT: nycnn-adbeame2;AME Job status;OK;SOFT;2;AME is idle:
[1406293271] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1406293281] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.164.219.113)
[1406293321] SERVICE ALERT: nycnn-editxfer3;PING;CRITICAL;HARD;2;CRITICAL - Host Unreachable (10.164.219.113)
[1406293351] HOST ALERT: nycnn-editxfer3;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.164.219.113)
[1406293381] SERVICE ALERT: nycnn-editxfer3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1406293391] HOST ALERT: nycnn-editxfer3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.16 ms
[1406309521] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;WARNING;SOFT;1;WARNING: at least one queue with depth 174, all MS2_EVENT queues in RUNNING state
[1406309581] SERVICE ALERT: msny;IMQ Depth Check - Internal GF Queues;OK;SOFT;2;OK: all MS2_EVENT queues have depth 0 or less, all MS2_EVENT queues in RUNNING state