[1416114344] HOST ALERT: nycnn-adbeame4;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1416114354] SERVICE ALERT: nycnn-adbeame4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416114354] SERVICE ALERT: nycnn-adbeame4;AME Job status;OK;HARD;1;AME is idle:
[1416115174] SERVICE ALERT: omnmas2db1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 90%, RTA = 0.30 ms
[1416115174] SERVICE ALERT: msny;MS2 - Grid A Mount;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1416115184] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115184] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115184] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115184] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115184] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: nycnn-mamdb1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnmas2app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnmas1db2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnmas2db-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnclda02e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnclda02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: nycnn-ms2db3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnmas2db2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115194] SERVICE ALERT: spider-tmp;/nymoneygfx Disk Space;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1416115204] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115204] SERVICE ALERT: nycnn-ipvstrm3;Process: IPS_FTP;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115204] SERVICE ALERT: nycnn-ipvstrm4;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115204] SERVICE ALERT: nycnn-ipvstrm3;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115204] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115204] HOST ALERT: omnmas2db1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115214] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115214] SERVICE ALERT: nycnn-ipvstrm1;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115214] SERVICE ALERT: nycnn-ipvstrm3;C:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115214] SERVICE ALERT: nycnn-ms2db1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115214] SERVICE ALERT: nycnn-ipvstrm3;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115214] SERVICE ALERT: nycnn-splunk;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115214] SERVICE ALERT: nycnn-mammq2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115214] SERVICE ALERT: nycnn-ipvstrm3;Process: LicMgr;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115214] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115214] SERVICE ALERT: nycnn-mamfog;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115214] SERVICE ALERT: omnmas2app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115214] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416115214] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115224] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1416115224] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1416115224] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1416115224] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1416115224] SERVICE ALERT: msny;MS2 - Grid A Mount;OK;SOFT;2;OK - File /Volumes/GridA/Temp/.gridamounted seems fine.
[1416115224] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;SOFT;1;(Service Check Timed Out)
[1416115234] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnclda01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnhbcba07-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnmas2app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnmas2db-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnclda02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnmas2db2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: omnspsm-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;1;(Host Check Timed Out)
[1416115234] SERVICE ALERT: omnmas2db1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115244] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115244] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115244] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115244] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115244] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115244] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] HOST ALERT: nycnn-splunk;DOWN;SOFT;1;(Host Check Timed Out)
[1416115254] HOST ALERT: nycnn-mammq2;DOWN;SOFT;1;(Host Check Timed Out)
[1416115254] HOST ALERT: omnclda01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115254] HOST ALERT: nycnn-mamfog;DOWN;SOFT;1;(Host Check Timed Out)
[1416115254] HOST ALERT: omnmas2app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115254] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416115254] SERVICE ALERT: nycnn-mamdb1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: nycnn-ms2fsd2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnmas2app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnmas1db2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnmas2db-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnclda02e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnclda02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: nycnn-ms2db3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnmas2db2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115254] SERVICE ALERT: omnspsm-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115264] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115264] SERVICE ALERT: nycnn-ipvstrm3;Process: IPS_FTP;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115264] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115264] SERVICE ALERT: nycnn-ipvstrm4;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115264] SERVICE ALERT: nycnn-ipvstrm3;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115264] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115274] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;2;(Host Check Timed Out)
[1416115274] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115274] SERVICE ALERT: nycnn-ipvstrm1;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115274] SERVICE ALERT: nycnn-ipvstrm3;C:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115274] SERVICE ALERT: nycnn-ms2db1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115274] SERVICE ALERT: nycnn-ipvstrm3;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115274] SERVICE ALERT: nycnn-splunk;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115274] SERVICE ALERT: nycnn-mammq2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115274] SERVICE ALERT: nycnn-ipvstrm3;Process: LicMgr;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115274] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115274] SERVICE ALERT: nycnn-mamfog;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115274] SERVICE ALERT: omnmas2app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115274] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416115274] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416115274] HOST ALERT: omnclda01-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115274] HOST ALERT: omnclda02e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115274] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115284] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;CRITICAL;HARD;1;(Service Check Timed Out)
[1416115294] SERVICE ALERT: nycnn-ipvstrm2;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115304] HOST ALERT: omnhbcba07-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115304] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115304] HOST ALERT: omnmas2db1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115314] SERVICE ALERT: spider-tmp;/nymoneygfx Disk Space;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1416115314] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1416115314] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1416115324] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1416115324] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1416115324] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115334] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnmas2app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnmas2db-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnclda02-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnmas2db2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: omnspsm-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;2;(Host Check Timed Out)
[1416115334] SERVICE ALERT: nycnn-ipvstrm2;Process: SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115334] SERVICE ALERT: nycnn-ipvstrm2;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115334] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115334] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115334] SERVICE ALERT: nycnn-ipvstrm3;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115334] SERVICE ALERT: nycnn-ipvstrm2;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115344] SERVICE ALERT: nycnn-ipvstrm2;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115344] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;CRITICAL;HARD;2;(Service Check Timed Out)
[1416115354] HOST ALERT: omnmas2app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115354] HOST ALERT: nycnn-splunk;DOWN;SOFT;2;(Host Check Timed Out)
[1416115354] HOST ALERT: nycnn-mammq2;DOWN;SOFT;2;(Host Check Timed Out)
[1416115354] HOST ALERT: omnclda01e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115354] HOST ALERT: nycnn-mamfog;DOWN;SOFT;2;(Host Check Timed Out)
[1416115354] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416115354] HOST ALERT: omnclda01-ny;UP;SOFT;3;PING OK - Packet loss = 61%, RTA = 0.20 ms
[1416115354] HOST ALERT: omnclda02e1-ny;UP;SOFT;3;PING OK - Packet loss = 64%, RTA = 0.22 ms
[1416115364] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115364] SERVICE ALERT: omnhbcba03-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416115364] SERVICE ALERT: omnhbcba04-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1416115364] SERVICE ALERT: omnhbcba07-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1416115364] SERVICE ALERT: omnhbcba09-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416115364] SERVICE ALERT: omnhbcba06-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416115364] SERVICE ALERT: omnhbcba12-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416115364] SERVICE ALERT: omnhbcba10-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1416115364] SERVICE ALERT: omnclda02e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.46 ms
[1416115364] SERVICE ALERT: omnclda02-ny;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.22 ms
[1416115364] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115364] HOST ALERT: omnmas1app2-ny;UP;SOFT;3;PING WARNING - Packet loss = 83%, RTA = 0.49 ms
[1416115364] HOST ALERT: nycnn-ms2db3;DOWN;SOFT;3;(Host Check Timed Out)
[1416115374] SERVICE ALERT: omnclda01-ny;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 50%, RTA = 0.28 ms
[1416115374] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.23 ms
[1416115374] SERVICE ALERT: nycnn-ipvstrm2;Process: MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115374] SERVICE ALERT: nycnn-ipvstrm1;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115374] HOST ALERT: omnhbcba03-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1416115374] HOST ALERT: omnhbcba04-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1416115374] HOST ALERT: omnhbcba07-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1416115374] HOST ALERT: omnhbcba09-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1416115374] HOST ALERT: omnhbcba06-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1416115374] HOST ALERT: omnhbcba12-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1416115374] HOST ALERT: omnhbcba10-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.17 ms
[1416115374] HOST ALERT: omnclda02-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1416115374] HOST ALERT: omnmas2db1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1416115384] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115384] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115384] SERVICE ALERT: omnclda01e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1416115394] SERVICE ALERT: omnmas2app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416115394] SERVICE ALERT: omnhbcba01-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416115394] HOST ALERT: omnclda01e1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1416115404] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115404] HOST ALERT: omnmas2app1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1416115404] HOST ALERT: omnhbcba01-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1416115404] HOST ALERT: omnmas2app2-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416115404] HOST ALERT: omnmas2db2-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1416115404] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING WARNING - Packet loss = 81%, RTA = 0.41 ms
[1416115404] SERVICE ALERT: nycnn-ipvstrm1;Process: SVServer;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115404] SERVICE ALERT: nycnn-ipvstrm2;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115404] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115404] SERVICE ALERT: nycnn-ipvstrm1;Process: IPV_FPT;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115404] SERVICE ALERT: nycnn-ipvstrm1;Process: MMSService;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115404] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;3;(Host Check Timed Out)
[1416115414] SERVICE ALERT: nycnn-ipvstrm3;Process: SVServer;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115414] SERVICE ALERT: omnmas2db1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1416115414] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 77%, RTA = 0.51 ms
[1416115414] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115414] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING WARNING - Packet loss = 80%, RTA = 0.25 ms
[1416115424] SERVICE ALERT: nycnn-ipvstrm3;Process: MMSService;UNKNOWN;SOFT;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416115424] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1416115424] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416115424] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1416115424] SERVICE ALERT: omnclda01-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1416115424] SERVICE ALERT: omnmas2app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1416115424] SERVICE ALERT: spider-tmp;/nymoneygfx Disk Space;OK;SOFT;3;DISK OK - free space: /nymoneygfx 3272903 MB (82% inode=96%):
[1416115424] HOST ALERT: omnmas2db-ny;UP;SOFT;3;PING WARNING - Packet loss = 83%, RTA = 0.33 ms
[1416115424] SERVICE ALERT: omnmas2db-ny;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 44%, RTA = 0.31 ms
[1416115434] HOST ALERT: nycnn-mamdb1;DOWN;SOFT;3;(Host Check Timed Out)
[1416115434] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;3;(Host Check Timed Out)
[1416115434] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1416115434] HOST ALERT: omnspsm-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1416115434] HOST ALERT: nycnn-ms2db1;DOWN;SOFT;3;(Host Check Timed Out)
[1416115434] SERVICE ALERT: omnmas2db2-ny;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.30 ms
[1416115444] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;UNKNOWN;HARD;2;Remote command execution failed: ssh: connect to host 10.164.219.199 port 22: Connection refused
[1416115444] HOST ALERT: nycnn-ms2db3;UP;SOFT;4;PING OK - Packet loss = 68%, RTA = 0.49 ms
[1416115444] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115454] HOST ALERT: nycnn-splunk;DOWN;SOFT;3;(Host Check Timed Out)
[1416115454] HOST ALERT: nycnn-mammq2;DOWN;SOFT;3;(Host Check Timed Out)
[1416115454] HOST ALERT: nycnn-mamfog;DOWN;SOFT;3;(Host Check Timed Out)
[1416115454] SERVICE ALERT: nycnn-splunk;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 1.13 ms
[1416115454] SERVICE ALERT: nycnn-ipvstrm4;Process: MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416115464] SERVICE ALERT: nycnn-ipvstrm1;Process: SVServer;OK;SOFT;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1416115464] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;OK;SOFT;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115464] SERVICE ALERT: nycnn-ipvstrm1;Process: IPV_FPT;OK;SOFT;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115464] SERVICE ALERT: nycnn-ipvstrm1;Process: MMSService;OK;SOFT;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416115474] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;OK;SOFT;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416115474] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1416115484] SERVICE ALERT: nycnn-ms2db3;Disk Space on /oracle;OK;HARD;2;DISK OK - free space: /oracle 858641 MB (87% inode=99%):
[1416115484] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1416115484] SERVICE ALERT: nycnn-ipvstrm3;Process: MMSService;OK;SOFT;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416115484] SERVICE ALERT: omnmas2db-ny;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416115484] SERVICE ALERT: nycnn-ms2db3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416115494] SERVICE ALERT: nycnn-mamdb1;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 1.75 ms
[1416115504] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;UNKNOWN;HARD;1;Remote command execution failed: ssh: connect to host 10.164.219.197 port 22: Connection refused
[1416115504] HOST ALERT: nycnn-mamdb1;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416115504] HOST ALERT: nycnn-ms2db1;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1416115504] SERVICE ALERT: nycnn-ms2db1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1416115504] SERVICE ALERT: nycnn-splunk;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416115514] HOST ALERT: nycnn-splunk;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416115534] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;4;(Host Check Timed Out)
[1416115534] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1416115534] HOST ALERT: omnspsm-ny;DOWN;SOFT;4;(Host Check Timed Out)
[1416115534] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1416115534] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416115544] SERVICE ALERT: nycnn-mamdb1;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416115544] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;OK;HARD;2;OK - New Volume Total=1397.19GB, Used=1130.75GB (80.9%), Free=266.44GB (19.1%)
[1416115544] SERVICE ALERT: nycnn-ipvstrm3;Process: IPS_FTP;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115554] HOST ALERT: nycnn-mammq2;DOWN;SOFT;4;(Host Check Timed Out)
[1416115554] HOST ALERT: nycnn-mamfog;DOWN;SOFT;4;(Host Check Timed Out)
[1416115554] SERVICE ALERT: nycnn-ipvstrm3;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1416115554] SERVICE ALERT: nycnn-ipvstrm4;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1416115554] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416115554] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;OK;HARD;2;OK - New Volume Total=1397.19GB, Used=1128.66GB (80.8%), Free=268.53GB (19.2%)
[1416115554] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;OK;HARD;2;OK (Sample Period 660 sec) - Average CPU Utilisation 0.00%
[1416115554] SERVICE ALERT: nycnn-ipvstrm1;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 4.000GB - Used: 0.748GB (19%) - Free: 3.251GB (81%)
[1416115554] SERVICE ALERT: nycnn-ipvstrm3;C:\ Drive Space;OK;HARD;2;OK - New Volume Total=1397.19GB, Used=1130.75GB (80.9%), Free=266.44GB (19.1%)
[1416115554] SERVICE ALERT: nycnn-ipvstrm3;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 4.000GB - Used: 0.761GB (19%) - Free: 3.239GB (81%)
[1416115554] SERVICE ALERT: nycnn-ipvstrm3;Process: LicMgr;OK;HARD;2;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416115564] SERVICE ALERT: nycnn-ms2db1;Disk Space on /oracle;OK;HARD;1;DISK OK - free space: /oracle 895706 MB (91% inode=99%):
[1416115564] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;OK;HARD;2;OK - Proxy Total=1397.19GB, Used=1128.84GB (80.8%), Free=268.35GB (19.2%)
[1416115574] SERVICE ALERT: nycnn-ipvstrm2;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.750GB (19%) - Free: 3.250GB (81%)
[1416115604] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115614] SERVICE ALERT: nycnn-ipvstrm2;Process: SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1416115614] SERVICE ALERT: nycnn-ipvstrm2;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416115624] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416115624] SERVICE ALERT: nycnn-ipvstrm2;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115624] SERVICE ALERT: nycnn-ipvstrm3;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115624] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1129.82GB (80.9%), Free=267.37GB (19.1%)
[1416115624] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;5;(Host Check Timed Out)
[1416115624] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1416115624] HOST ALERT: omnspsm-ny;DOWN;SOFT;5;(Host Check Timed Out)
[1416115634] SERVICE ALERT: nycnn-ipvstrm2;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115644] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416115654] HOST ALERT: nycnn-mammq2;DOWN;SOFT;5;(Host Check Timed Out)
[1416115654] HOST ALERT: nycnn-mamfog;DOWN;SOFT;5;(Host Check Timed Out)
[1416115654] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.12%
[1416115654] SERVICE ALERT: nycnn-ipvstrm2;Process: MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416115654] SERVICE ALERT: nycnn-ipvstrm1;C:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1128.66GB (80.8%), Free=268.53GB (19.2%)
[1416115674] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1128.84GB (80.8%), Free=268.35GB (19.2%)
[1416115674] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416115684] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416115704] SERVICE ALERT: nycnn-ipvstrm2;C:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1129.82GB (80.9%), Free=267.37GB (19.1%)
[1416115704] SERVICE ALERT: nycnn-ipvstrm3;Process: SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1416115724] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;6;(Host Check Timed Out)
[1416115724] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1416115724] HOST ALERT: omnspsm-ny;DOWN;SOFT;6;(Host Check Timed Out)
[1416115734] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.834GB (21%) - Free: 3.166GB (79%)
[1416115734] SERVICE ALERT: nycnn-ipvstrm4;Process: MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416115754] HOST ALERT: nycnn-mammq2;DOWN;SOFT;6;(Host Check Timed Out)
[1416115754] HOST ALERT: nycnn-mamfog;DOWN;SOFT;6;(Host Check Timed Out)
[1416115824] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;7;(Host Check Timed Out)
[1416115824] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1416115824] HOST ALERT: omnspsm-ny;DOWN;SOFT;7;(Host Check Timed Out)
[1416115854] HOST ALERT: nycnn-mammq2;DOWN;SOFT;7;(Host Check Timed Out)
[1416115854] HOST ALERT: nycnn-mamfog;DOWN;SOFT;7;(Host Check Timed Out)
[1416115924] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;8;(Host Check Timed Out)
[1416115924] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1416115924] HOST ALERT: omnspsm-ny;DOWN;SOFT;8;(Host Check Timed Out)
[1416115954] HOST ALERT: nycnn-mammq2;DOWN;SOFT;8;(Host Check Timed Out)
[1416115954] HOST ALERT: nycnn-mamfog;DOWN;SOFT;8;(Host Check Timed Out)
[1416116014] HOST ALERT: nycnn-ms2fsd2;DOWN;SOFT;9;(Host Check Timed Out)
[1416116014] HOST ALERT: omnmas1db2-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1416116014] HOST ALERT: omnspsm-ny;DOWN;SOFT;9;(Host Check Timed Out)
[1416116054] HOST ALERT: nycnn-mammq2;DOWN;SOFT;9;(Host Check Timed Out)
[1416116054] HOST ALERT: nycnn-mamfog;DOWN;SOFT;9;(Host Check Timed Out)
[1416116114] HOST ALERT: nycnn-ms2fsd2;DOWN;HARD;10;(Host Check Timed Out)
[1416116114] HOST ALERT: omnmas1db2-ny;DOWN;HARD;10;(Host Check Timed Out)
[1416116114] HOST ALERT: omnspsm-ny;DOWN;HARD;10;(Host Check Timed Out)
[1416116154] HOST ALERT: nycnn-mammq2;DOWN;HARD;10;(Host Check Timed Out)
[1416116154] HOST ALERT: nycnn-mamfog;DOWN;HARD;10;(Host Check Timed Out)
[1416117244] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117254] SERVICE ALERT: nycnn-ipvstrm4;Process: MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117254] SERVICE ALERT: nycnn-splunk;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117254] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117254] SERVICE ALERT: omnmas2app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117254] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117264] HOST ALERT: omnmas2db1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117274] SERVICE ALERT: omnmas2db1-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117284] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;1;(Host Check Timed Out)
[1416117284] SERVICE ALERT: nycnn-ipvstrm1;Process: SVServer;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117284] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117284] SERVICE ALERT: nycnn-ipvstrm1;Process: IPV_FPT;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117284] SERVICE ALERT: nycnn-ipvstrm1;Process: MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117284] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117284] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117284] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117284] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117284] SERVICE ALERT: nycnn-ipvstrm4;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117284] HOST ALERT: nycnn-splunk;DOWN;SOFT;1;(Host Check Timed Out)
[1416117284] HOST ALERT: omnclda01e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117284] HOST ALERT: omnmas2app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117284] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117284] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117284] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117294] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117294] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117294] HOST ALERT: omnmas2app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117294] HOST ALERT: omnmas2db2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117294] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;1;(Host Check Timed Out)
[1416117294] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnmas2app2-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnmas2db-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: omnmas2db2-ny;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1416117294] SERVICE ALERT: nycnn-ipvstrm3;Process: MMSService;UNKNOWN;SOFT;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117294] SERVICE ALERT: spider-tmp;/nymoneygfx Disk Space;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1416117304] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;1;(Host Check Timed Out)
[1416117304] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;1;(Host Check Timed Out)
[1416117304] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117314] HOST ALERT: omnmas2db-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117314] SERVICE ALERT: nycnn-ipvstrm4;Process: MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117314] SERVICE ALERT: nycnn-splunk;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117314] SERVICE ALERT: omnclda01e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117314] SERVICE ALERT: omnmas2app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117314] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117324] HOST ALERT: omnhbcba07-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117324] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117324] HOST ALERT: omnclda01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117324] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416117344] SERVICE ALERT: nycnn-ipvstrm1;Process: SVServer;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117344] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117344] SERVICE ALERT: nycnn-ipvstrm1;Process: IPV_FPT;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117344] SERVICE ALERT: nycnn-ipvstrm1;Process: MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117344] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117344] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117344] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117344] SERVICE ALERT: nycnn-ipvstrm1;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117344] SERVICE ALERT: nycnn-ipvstrm2;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117344] SERVICE ALERT: nycnn-ipvstrm3;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117354] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117354] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117354] SERVICE ALERT: omnclda01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117354] SERVICE ALERT: omnmas2db-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117354] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416117354] SERVICE ALERT: nycnn-ipvstrm3;Process: MMSService;UNKNOWN;HARD;2;UNKNOWN - Plugin Timed out (15 sec)
[1416117354] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;2;(Host Check Timed Out)
[1416117354] HOST ALERT: omnmas2db1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117354] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117364] SERVICE ALERT: nycnn-ipvstrm3;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117364] SERVICE ALERT: nycnn-ipvstrm3;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117364] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117364] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117364] SERVICE ALERT: nycnn-ipvstrm4;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117374] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117374] SERVICE ALERT: nycnn-ipvstrm1;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117374] SERVICE ALERT: nycnn-ipvstrm3;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117374] SERVICE ALERT: nycnn-ipvstrm3;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117374] SERVICE ALERT: nycnn-ipvstrm3;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117374] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117374] HOST ALERT: nycnn-splunk;DOWN;SOFT;2;(Host Check Timed Out)
[1416117374] HOST ALERT: omnclda01e1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117374] HOST ALERT: omnmas2app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117374] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117384] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117384] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117394] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117394] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117394] HOST ALERT: omnmas2app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117394] HOST ALERT: omnmas2db2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117394] HOST ALERT: nycnn-ipvstrm1;DOWN;SOFT;2;(Host Check Timed Out)
[1416117394] SERVICE ALERT: nycnn-ipvstrm2;Memory Usage;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117394] HOST ALERT: omnhbcba07-ny;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.23 ms
[1416117404] HOST ALERT: nycnn-ipvstrm2;DOWN;SOFT;2;(Host Check Timed Out)
[1416117404] HOST ALERT: nycnn-ipvstrm3;DOWN;SOFT;2;(Host Check Timed Out)
[1416117404] SERVICE ALERT: omnhbcba03-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1416117404] SERVICE ALERT: omnhbcba04-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416117404] SERVICE ALERT: omnhbcba07-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1416117404] SERVICE ALERT: omnhbcba09-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117404] SERVICE ALERT: omnhbcba12-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1416117404] SERVICE ALERT: omnhbcba06-ny;PING;WARNING;HARD;2;PING WARNING - Packet loss = 37%, RTA = 0.19 ms
[1416117404] HOST ALERT: omnhbcba06-ny;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.21 ms
[1416117404] SERVICE ALERT: spider-tmp;/nymoneygfx Disk Space;OK;SOFT;2;DISK OK - free space: /nymoneygfx 3272903 MB (82% inode=96%):
[1416117414] HOST ALERT: omnmas2db-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117414] HOST ALERT: omnclda01-ny;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.26 ms
[1416117414] HOST ALERT: omnhbcba03-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416117414] HOST ALERT: omnhbcba04-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416117414] HOST ALERT: omnhbcba09-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1416117414] HOST ALERT: omnhbcba12-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416117414] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416117424] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117424] HOST ALERT: omnmas2db1-ny;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 0.53 ms
[1416117434] SERVICE ALERT: omnclda01e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 16%, RTA = 0.62 ms
[1416117434] SERVICE ALERT: omnmas2app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1416117434] SERVICE ALERT: nycnn-ipvstrm2;Process: SVServer;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117434] SERVICE ALERT: nycnn-ipvstrm2;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117434] SERVICE ALERT: omnhbcba01-ny;PING;WARNING;HARD;2;PING WARNING - Packet loss = 37%, RTA = 0.62 ms
[1416117434] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117434] SERVICE ALERT: nycnn-ipvstrm3;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117434] SERVICE ALERT: nycnn-ipvstrm2;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117434] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117444] HOST ALERT: omnclda01e1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1416117444] HOST ALERT: omnmas2app1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416117444] HOST ALERT: omnhbcba01-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416117444] SERVICE ALERT: nycnn-ipvstrm2;Process: IPS_FTP;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117454] HOST ALERT: nycnn-ipvstrm4;DOWN;SOFT;3;(Host Check Timed Out)
[1416117454] SERVICE ALERT: omnmas2db1-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416117464] HOST ALERT: omnmas2app2-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416117464] HOST ALERT: omnmas2db2-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117464] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117464] SERVICE ALERT: omnhbcba06-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1416117464] SERVICE ALERT: omnclda01-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117464] SERVICE ALERT: omnmas2app2-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1416117464] SERVICE ALERT: nycnn-ipvstrm1;C:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - The WMI query had problems. The error text from wmic is: [librpc/rpc/dcerpc_connect.c:329:dcerpc_pipe_connect_ncacn_ip_tcp_recv()] failed NT status (c0000236) in dcerpc_pipe_connect_ncacn_ip_tcp_recv
[1416117464] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117464] HOST ALERT: nycnn-ipvstrm1;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 0.37 ms
[1416117464] HOST FLAPPING ALERT: nycnn-ipvstrm1;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1416117464] SERVICE ALERT: nycnn-ipvstrm1;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 54%, RTA = 0.23 ms
[1416117464] SERVICE ALERT: omnmas2db2-ny;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1416117464] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117474] HOST ALERT: nycnn-splunk;DOWN;SOFT;3;(Host Check Timed Out)
[1416117474] SERVICE ALERT: omnmas2db-ny;PING;WARNING;HARD;2;PING WARNING - Packet loss = 50%, RTA = 0.31 ms
[1416117474] SERVICE ALERT: nycnn-ipvstrm2;Process: MMSService;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117474] HOST ALERT: omnmas1app2-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117474] HOST ALERT: nycnn-ipvstrm3;UP;SOFT;3;PING OK - Packet loss = 37%, RTA = 0.40 ms
[1416117474] HOST FLAPPING ALERT: nycnn-ipvstrm3;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1416117474] HOST ALERT: nycnn-ipvstrm2;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 0.33 ms
[1416117474] HOST FLAPPING ALERT: nycnn-ipvstrm2;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1416117484] HOST ALERT: omnmas2db-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1416117484] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117484] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117494] SERVICE ALERT: omnhbcba01-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117504] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;UNKNOWN;HARD;1;UNKNOWN - Plugin Timed out (15 sec)
[1416117524] SERVICE ALERT: nycnn-ipvstrm1;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416117524] SERVICE ALERT: nycnn-ipvstrm2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416117524] SERVICE ALERT: nycnn-ipvstrm3;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1416117524] SERVICE ALERT: nycnn-ipvstrm4;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416117524] HOST ALERT: nycnn-ipvstrm4;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117524] HOST FLAPPING ALERT: nycnn-ipvstrm4;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1416117524] SERVICE ALERT: omnmas2db-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1416117544] HOST ALERT: nycnn-splunk;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416117544] SERVICE ALERT: nycnn-splunk;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416117594] SERVICE ALERT: nycnn-ipvstrm4;Memory Usage;OK;HARD;2;OK - Physical Memory: Total: 4.000GB - Used: 0.785GB (20%) - Free: 3.214GB (80%)
[1416117594] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Memory Usage;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1416117594] SERVICE ALERT: nycnn-ipvstrm4;Process: MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416117594] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: MMSService;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1416117624] SERVICE ALERT: nycnn-ipvstrm1;Process: SVServer;OK;HARD;2;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1416117624] SERVICE ALERT: nycnn-ipvstrm1;Process: IPS_FTP;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117624] SERVICE ALERT: nycnn-ipvstrm1;Process: IPV_FPT;OK;HARD;2;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117624] SERVICE ALERT: nycnn-ipvstrm1;Process: MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416117634] SERVICE ALERT: nycnn-ipvstrm4;Process: SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1416117634] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: SVServer;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117634] SERVICE ALERT: nycnn-ipvstrm1;SSH Service;OK;HARD;2;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416117634] SERVICE ALERT: nycnn-ipvstrm1;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416117634] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;Process: LicMgr;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117644] SERVICE ALERT: nycnn-ipvstrm3;Process: MMSService;OK;HARD;2;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416117644] SERVICE ALERT: nycnn-ipvstrm3;D:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1130.75GB (80.9%), Free=266.44GB (19.1%)
[1416117644] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;D:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117644] SERVICE ALERT: nycnn-ipvstrm3;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117644] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Process: IPS_FTP;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm4;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.28%
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm3;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.49%
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm4;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;SSH Service;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm1;D:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1128.66GB (80.8%), Free=268.53GB (19.2%)
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;D:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm2;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.09%
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm3;C:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1130.75GB (80.9%), Free=266.44GB (19.1%)
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;C:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm1;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.743GB (19%) - Free: 3.256GB (81%)
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;Memory Usage;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm3;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.794GB (20%) - Free: 3.205GB (80%)
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Memory Usage;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117654] SERVICE ALERT: nycnn-ipvstrm3;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416117654] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Process: LicMgr;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117664] SERVICE ALERT: nycnn-ipvstrm4;C:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1128.84GB (80.8%), Free=268.35GB (19.2%)
[1416117664] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;C:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117674] SERVICE ALERT: nycnn-ipvstrm2;Memory Usage;OK;HARD;1;OK - Physical Memory: Total: 4.000GB - Used: 0.786GB (20%) - Free: 3.214GB (80%)
[1416117674] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Memory Usage;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117704] SERVICE ALERT: nycnn-ipvstrm4;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117704] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: IPS_FTP;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117714] SERVICE ALERT: nycnn-ipvstrm2;Process: SVServer;OK;HARD;1;OK - Found 1 Instance(s) of "SVServer.exe" running (0 excluded).  (List is on next line)
[1416117714] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: SVServer;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117714] SERVICE ALERT: nycnn-ipvstrm2;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416117714] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: LicMgr;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117724] SERVICE ALERT: nycnn-ipvstrm2;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1129.82GB (80.9%), Free=267.37GB (19.1%)
[1416117724] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;D:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117724] SERVICE ALERT: nycnn-ipvstrm2;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416117724] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;SSH Service;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117724] SERVICE ALERT: nycnn-ipvstrm2;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117724] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: IPV_FPT;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117724] SERVICE ALERT: nycnn-ipvstrm3;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117724] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Process: IPV_FPT;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117734] SERVICE ALERT: nycnn-ipvstrm2;Process: IPS_FTP;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117734] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: IPS_FTP;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117744] SERVICE ALERT: nycnn-ipvstrm3;SSH Service;OK;HARD;1;OK - Found 1 Instance(s) of "sshserver.exe" running (0 excluded).  (List is on next line)
[1416117744] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;SSH Service;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117754] SERVICE ALERT: nycnn-ipvstrm1;CPU Load;OK;HARD;1;OK (Sample Period 600 sec) - Average CPU Utilisation 0.08%
[1416117754] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117754] SERVICE ALERT: nycnn-ipvstrm1;C:\ Drive Space;OK;HARD;1;OK - New Volume Total=1397.19GB, Used=1128.66GB (80.8%), Free=268.53GB (19.2%)
[1416117754] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;C:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117754] SERVICE ALERT: nycnn-ipvstrm2;Process: MMSService;OK;HARD;1;OK - Found 1 Instance(s) of "MMSService.exe" running (0 excluded).  (List is on next line)
[1416117754] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: MMSService;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117774] SERVICE ALERT: nycnn-ipvstrm4;D:\ Drive Space;OK;HARD;1;OK - Proxy Total=1397.19GB, Used=1128.84GB (80.8%), Free=268.35GB (19.2%)
[1416117774] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;D:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117774] SERVICE ALERT: nycnn-ipvstrm4;Process: LicMgr;OK;HARD;1;OK - Found 1 Instance(s) of "LicenceManager.exe" running (0 excluded).  (List is on next line)
[1416117774] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: LicMgr;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416117784] SERVICE ALERT: nycnn-ipvstrm4;Process: IPV_FPT;OK;HARD;1;OK - Found 1 Instance(s) of "iftpsvc.exe" running (0 excluded).  (List is on next line)
[1416117784] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: IPV_FPT;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1416120594] HOST FLAPPING ALERT: nycnn-ipvstrm4;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1416120874] HOST FLAPPING ALERT: nycnn-ipvstrm1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1416120884] HOST FLAPPING ALERT: nycnn-ipvstrm2;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1416120884] HOST FLAPPING ALERT: nycnn-ipvstrm3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1416121194] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416121234] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416121254] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416121334] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416121424] HOST ALERT: omnmas1app2-ny;UP;SOFT;3;PING OK - Packet loss = 78%, RTA = 0.52 ms
[1416121424] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416121684] SERVICE ALERT: nycnn-mammq2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416121694] HOST ALERT: nycnn-mammq2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416121724] SERVICE ALERT: nycnn-ms2fsd2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1416121734] HOST ALERT: nycnn-ms2fsd2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1416121814] SERVICE ALERT: nycnn-mamfog;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1416121824] HOST ALERT: nycnn-mamfog;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1416122394] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416122424] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416122454] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416122474] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416122544] HOST ALERT: omnmas1app1-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1416122564] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1416123094] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1416123144] SERVICE ALERT: msny;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.002 second response time
[1416123214] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1416123264] SERVICE ALERT: msny;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.002 second response time
[1416123294] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123294] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: MMSService;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123334] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: SVServer;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123334] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;Process: LicMgr;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123344] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;D:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123344] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Process: IPS_FTP;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;SSH Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;D:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123354] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Process: LicMgr;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123364] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123374] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123384] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1416123404] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: IPS_FTP;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123414] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: SVServer;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123414] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: LicMgr;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123424] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;Process: IPV_FPT;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123424] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;SSH Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123424] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: IPV_FPT;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123424] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;D:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123434] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: IPS_FTP;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123444] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1416123444] SERVICE FLAPPING ALERT: nycnn-ipvstrm3;SSH Service;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123454] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123454] SERVICE FLAPPING ALERT: nycnn-ipvstrm2;Process: MMSService;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123454] SERVICE FLAPPING ALERT: nycnn-ipvstrm1;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123474] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;D:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123474] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: LicMgr;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123484] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;WARNING;SOFT;1;WARNING - [Triggered by _MemUsed%>90] - Physical Memory: Total: 1.990GB - Used: 1.809GB (91%) - Free: 0.181GB (9%)
[1416123484] SERVICE ALERT: nycnn-ipvx27;Memory Usage;WARNING;SOFT;1;WARNING - [Triggered by _MemUsed%>90] - Physical Memory: Total: 1.990GB - Used: 1.826GB (92%) - Free: 0.164GB (8%)
[1416123484] SERVICE FLAPPING ALERT: nycnn-ipvstrm4;Process: IPV_FPT;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1416123494] HOST ALERT: omnspsm-ny;UP;HARD;1;PING OK - Packet loss = 61%, RTA = 1.51 ms
[1416123504] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1416123534] SERVICE ALERT: omnspsm-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1416123544] SERVICE ALERT: nycnn-ipvx27e1;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 1.990GB - Used: 0.795GB (40%) - Free: 1.195GB (60%)
[1416123544] SERVICE ALERT: nycnn-ipvx27;Memory Usage;OK;SOFT;2;OK - Physical Memory: Total: 1.990GB - Used: 0.795GB (40%) - Free: 1.195GB (60%)
[1416123584] SERVICE ALERT: omnmas1db2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1416123594] HOST ALERT: omnmas1db2-ny;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416123624] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.123 second response time
[1416123894] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;WARNING;SOFT;1;WARNING (Sample Period 300 sec) - [Triggered by _AvgCPU>90] - Average CPU Utilisation 100.00%
[1416123954] SERVICE ALERT: nycnn-ipvsphn2;CPU Load;OK;SOFT;2;OK (Sample Period 60 sec) - Average CPU Utilisation 0.07%
[1416125724] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1416125784] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1416125844] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1416126494] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416126504] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;1;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1416126524] HOST ALERT: omnmdir04e1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416126534] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416126554] SERVICE ALERT: omnmdir04e1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416126564] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;SOFT;2;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1416126574] HOST ALERT: omnmdir04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416126594] SERVICE ALERT: omnmdir04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416126604] HOST ALERT: omnmdir04e1-ny;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.21 ms
[1416126614] SERVICE ALERT: omnmdir04e1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.15 ms
[1416126624] SERVICE ALERT: msny;MS2 - Grid B Mount;CRITICAL;HARD;3;CRITICAL - File /Volumes/GridB/Temp/.gridbmounted not found or not a regular file
[1416126644] HOST ALERT: omnmdir04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1416126644] SERVICE ALERT: omnmdir04-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1416127164] SERVICE ALERT: msny;MS2 - Grid B Mount;OK;HARD;3;OK - File /Volumes/GridB/Temp/.gridbmounted seems fine.
[1416127644] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.063 second response time
[1416132144] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1416132204] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1416132264] SERVICE ALERT: msny;HTTP;CRITICAL;HARD;3;Connection refused
[1416135624] SERVICE ALERT: msny;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.008 second response time
[1416136514] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416136554] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416136574] SERVICE ALERT: omnhbcba01-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416136614] SERVICE ALERT: omnhbcba05-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416136644] HOST ALERT: omnhbcba05-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416136654] HOST ALERT: omnhbcba01-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416136674] SERVICE ALERT: omnhbcba05-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416136694] HOST ALERT: omnhbcba05-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416136724] HOST ALERT: omnhbcba01-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416136734] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416136754] SERVICE ALERT: omnhbcba01-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416136774] HOST ALERT: omnhbcba12-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416136794] HOST ALERT: omnhbcba05-ny;DOWN;SOFT;3;(Host Check Timed Out)
[1416136794] SERVICE ALERT: omnhbcba12-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416136844] SERVICE ALERT: omnhbcba05-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416136854] HOST ALERT: omnhbcba12-ny;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.17 ms
[1416136854] HOST ALERT: omnhbcba05-ny;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416136904] SERVICE ALERT: omnhbcba12-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1416136934] SERVICE ALERT: nycnn-arvtrtape;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416136944] HOST ALERT: nycnn-arvtrtape;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1416137034] SERVICE ALERT: omnhbcba02-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137034] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137064] HOST ALERT: omnhbcba02-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137064] HOST ALERT: omnhbcba07-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137084] SERVICE ALERT: nycnn-arvtrmops2;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416137094] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137094] SERVICE ALERT: omnhbcba02-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137094] SERVICE ALERT: omnhbcba07-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137094] SERVICE ALERT: nycnn-arvtrtape;SSH;OK;HARD;2;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1416137134] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137144] HOST ALERT: omnhbcba07-ny;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.25 ms
[1416137144] SERVICE ALERT: omnhbcba07-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416137154] SERVICE ALERT: omnhbcba09-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137164] HOST ALERT: omnhbcba02-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416137174] SERVICE ALERT: nycnn-arvtrmops2;SSH;OK;HARD;1;SSH OK - OpenSSH_4.3 (protocol 2.0)
[1416137204] SERVICE ALERT: omnhbcba02-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1416137214] HOST ALERT: omnhbcba02-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416137224] HOST ALERT: omnhbcba09-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416137264] SERVICE ALERT: omnhbcba09-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1416137274] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137274] HOST ALERT: omnhbcba09-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416137304] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137334] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137334] SERVICE ALERT: omnhbcba03-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137354] SERVICE ALERT: omnhbcba08-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137374] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137384] HOST ALERT: omnhbcba08-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137394] SERVICE ALERT: omnhbcba10-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137404] HOST ALERT: omnhbcba03-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416137414] SERVICE ALERT: omnhbcba08-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137434] HOST ALERT: omnhbcba10-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416137444] SERVICE ALERT: omnhbcba03-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 1.54 ms
[1416137454] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137454] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.35 ms
[1416137454] HOST ALERT: omnhbcba03-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1416137484] HOST ALERT: omnhbcba08-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416137494] HOST ALERT: omnhbcba04-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137494] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137504] HOST ALERT: omnhbcba10-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1416137504] SERVICE ALERT: omnhbcba10-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416137514] SERVICE ALERT: omnhbcba04-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137514] SERVICE ALERT: omnhbcba06-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137524] SERVICE ALERT: omnhbcba08-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1416137534] HOST ALERT: omnhbcba08-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416137564] HOST ALERT: omnhbcba04-ny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1416137564] SERVICE ALERT: omnhbcba04-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1416137574] HOST ALERT: omnhbcba06-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416137624] SERVICE ALERT: omnhbcba06-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1416137624] SERVICE ALERT: omnmdk01h0-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 14%, RTA = 2211.33 ms
[1416137634] HOST ALERT: omnhbcba06-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1416137684] SERVICE ALERT: omnmdk01h0-ny;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 69.71 ms
[1416137794] SERVICE ALERT: omnhbcba11-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416137824] HOST ALERT: omnhbcba11-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416137854] SERVICE ALERT: omnhbcba11-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416137924] HOST ALERT: omnhbcba11-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416137974] SERVICE ALERT: omnhbcba11-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1416137984] HOST ALERT: omnhbcba11-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1416138264] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;1;Connection refused
[1416138324] SERVICE ALERT: msny;HTTP;CRITICAL;SOFT;2;Connection refused
[1416138384] SERVICE ALERT: msny;HTTP;OK;SOFT;3;HTTP OK: HTTP/1.1 200 OK - 5504 bytes in 0.006 second response time
[1416139294] HOST ALERT: nycnn-arvtrmops2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1416140034] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416140064] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416140094] SERVICE ALERT: omnmas1app1-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416140164] HOST ALERT: omnmas1app1-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416140254] HOST ALERT: omnmas1app1-ny;UP;SOFT;3;PING OK - Packet loss = 77%, RTA = 0.35 ms
[1416140264] SERVICE ALERT: omnmas1app1-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1416140934] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1416140974] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;1;(Host Check Timed Out)
[1416140994] SERVICE ALERT: omnmas1app2-ny;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1416141074] HOST ALERT: omnmas1app2-ny;DOWN;SOFT;2;(Host Check Timed Out)
[1416141144] HOST ALERT: omnmas1app2-ny;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1416141164] SERVICE ALERT: omnmas1app2-ny;PING;OK;HARD;2;PING OK - Packet loss = 0%, RTA = 0.27 ms